<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>System Results</title>
<meta name="ROBOTS" content="INDEX,NOFOLLOW" />
<script type="text/javascript">//<![CDATA[

  var tstp_dump;
  function openSoTSTP (dump) {   
    var tstp_url = 'http://www.tptp.org/cgi-bin/SystemOnTSTP';
    var tstp_browser = window.open(tstp_url, '_blank');
    tstp_dump = dump;
  }
  function getTSTPDump () {
    return tstp_dump;
  }

//]]></script>
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>

% SZS start RequiredInformation<BR>
% Congratulations - you have become a registered power user of SystemOnTPTP, at IP address 191.6.38.118.<BR>
% Please consider donating to the TPTP project - see www.tptp.org for details.<BR>
% When you donate this message will disappear.<BR>
% If you do not donate a random delay might be added to your processing time.<BR>
% SZS end RequiredInformation<BR>
<PRE>
Getting system recommendations for generic problem /tmp/Bo1eORq5Ot/SOT_4blGUI
This is a FOF_THM_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.95
Vampire---TF1-4.5    recommended  0.95
Vampire---4.5        recommended  0.95
Vampire---SAT-4.5    recommended  0.95
ET---2.0             recommended  0.80
E---2.5              recommended  0.78
E---FNT-2.5          recommended  0.78
CSE_E---1.0          recommended  0.76
E---FNT-2.5          recommended  0.74
Isabelle---2TH0      recommended  0.69
Isabelle---2TF0      recommended  0.69
Isabelle---2FOF      recommended  0.69
Isabelle---2016      recommended  0.69
Isabelle-HOT---2016  recommended  0.68
CVC4---1.7           recommended  0.68
CVC4---SAT-1.7       recommended  0.68
Leo-III---1.4        recommended  0.61
Vampire---THF-4.5    recommended  0.59
Vampire---TF1-4.5    recommended  0.59
Vampire---SAT-4.5    recommended  0.59
iProver---3.0        recommended  0.58
iProver---SAT-3.0    recommended  0.58
Fampire---1.3        recommended  0.57
SInE---0.4           recommended  0.53
iProver---SAT-3.0    recommended  0.52
ePrincess---1.0      recommended  0.50
Princess---170717    recommended  0.48
SPASS---3.9          recommended  0.48
Z3---4.4.1           recommended  0.46
Zipperpin---1.5      recommended  0.40
CVC4---SAT-1.7       recommended  0.40
Prover9---1109a      recommended  0.39
Equinox---6.0.1a     recommended  0.37
LEO-II---1.7.0       recommended  0.35
nanoCoP---1.1        recommended  0.34
Muscadet---4.5       recommended  0.25
Twee---2.2           recommended  0.23
Metis---2.4          recommended  0.23
Geo-III---2018C      recommended  0.18
iProverMo---2.5-0.1  recommended  0.14
Beagle---0.9.47      subsumed     0.41
leanCoP---2.2        subsumed     0.36
E-Darwin---1.5       subsumed     0.31
CSE---1.0            subsumed     0.28
Darwin---1.4.5       subsumed     0.28
Bliksem---1.12       subsumed     0.21
Zenon---0.7.1        subsumed     0.16
ZenonModulo---0.4.2  subsumed     0.15
SOS---2.0            subsumed     0.11
Otter---3.3          subsumed     0.04
This is a FOF_CSA_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.98
Vampire---TF1-4.5    recommended  0.98
Vampire---SAT-4.5    recommended  0.98
Paradox---4.0        recommended  0.60
Nitpick---2016       subsumed     0.55
DarwinFM---1.4.5     subsumed     0.43
CVC4---SAT-1.7       subsumed     0.40
Vampire---THF-4.5    subsumed     0.34
Vampire---TF1-4.5    subsumed     0.34
Vampire---4.5        subsumed     0.34
Vampire---SAT-4.5    subsumed     0.34
E---2.5              subsumed     0.34
E---FNT-2.5          subsumed     0.34
E---FNT-2.5          subsumed     0.34
iProver---SAT-3.0    subsumed     0.32
Metis---2.4          subsumed     0.32
SPASS---3.9          subsumed     0.28
Beagle---0.9.47      subsumed     0.26
Geo-III---2018C      subsumed     0.25
Z3---4.4.1           subsumed     0.25
Princess---170717    subsumed     0.25
Equinox---6.0.1a     subsumed     0.19
iProver---3.0        subsumed     0.11
iProver---SAT-3.0    subsumed     0.11
Darwin---1.4.5       subsumed     0.06
E-Darwin---1.5       subsumed     0.04
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use ET,2.0
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use CSE_E,1.0
CSE_E---1.0 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use Isabelle,2TH0
Isabelle---2016 can be used
Check if we can use Isabelle,2TF0
Isabelle---2016 can be used
Check if we can use Isabelle,2FOF
Isabelle---2016 can be used
Check if we can use Isabelle,2016
Isabelle---2016 can be used
Check if we can use Isabelle-HOT,2016
Isabelle-HOT---2016 can be used
Check if we can use CVC4,1.7
CVC4---SAT-1.7 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Leo-III,1.4
Leo-III---1.4 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Fampire,1.3
Check if we can use SInE,0.4
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use ePrincess,1.0
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Zipperpin,1.5
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Prover9,1109a
Prover9---1109a can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use LEO-II,1.7.0
LEO-II---1.7.0 can be used
Check if we can use nanoCoP,1.1
nanoCoP---1.1 can be used
Check if we can use Muscadet,4.5
Muscadet---4.5 can be used
Check if we can use Twee,2.2
Twee---2.2 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use iProverMo,2.5-0.1
iProverMo---2.5-0.1 can be used
Check if we can use Beagle,0.9.47
Check if we can use leanCoP,2.2
leanCoP---2.2 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Check if we can use CSE,1.0
CSE---1.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use Bliksem,1.12
Bliksem---1.12 can be used
Check if we can use Zenon,0.7.1
Zenon---0.7.1 can be used
Check if we can use ZenonModulo,0.4.2
Check if we can use SOS,2.0
SOS---2.0 can be used
Check if we can use Otter,3.3
Otter---3.3 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use Paradox,4.0
Paradox---4.0 can be used
Check if we can use Nitpick,2016
Nitpick---2016 can be used
Check if we can use DarwinFM,1.4.5
DarwinFM---1.4.5 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Beagle,0.9.47
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
E---FNT-2.5 will be used
E---FNT-2.5 will be used
Vampire---SAT-4.5 is given 499 seconds in -cE mode
Vampire---SAT-4.5 is given 250 seconds in -cE mode
Vampire---SAT-4.5 is given 125 seconds in -cE mode
Vampire---SAT-4.5 is given 62 seconds in -cE mode
E---FNT-2.5 is given 31 seconds in -cE mode
E---FNT-2.5 is given 32 seconds in -cE mode
Vampire---SAT-4.5 has started (PID 97283)
Vampire---SAT-4.5 has started (PID 97284)
Vampire---SAT-4.5 has started (PID 97285)
Vampire---SAT-4.5 has started (PID 97286)
E---FNT-2.5 has started (PID 97288)
E---FNT-2.5 has started (PID 97290)
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  499  /tmp/Bo1eORq5Ot/SOT_4blGUI
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  250  /tmp/Bo1eORq5Ot/SOT_4blGUI
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  125  /tmp/Bo1eORq5Ot/SOT_4blGUI
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  62  /tmp/Bo1eORq5Ot/SOT_4blGUI
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  31  /tmp/Bo1eORq5Ot/SOT_4blGUI
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  32  /tmp/Bo1eORq5Ot/SOT_4blGUI
% START OF SYSTEM OUTPUT
% START OF SYSTEM OUTPUT
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_4blG
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_4blG
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_4blG
% SZS output start Proof for SOT_4blG
fof(f161,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f96,f100,f112,f116,f132,f136,f141,f147,f155,f159,f160])).
fof(f160,plain,(
  spl11_14 | ~spl11_3),
  inference(avatar_split_clause,[],[f52,f98,f153])).
fof(f153,plain,(
  spl11_14 <=> individual(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_14])])).
fof(f98,plain,(
  spl11_3 <=> type_(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_3])])).
fof(f52,plain,(
  ~type_(sK0) | individual(sK1)),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ? [X0] : (individual(X0) & type_(X0)) | ? [X1] : (~individual(X1) & ~type_(X1))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~(~? [X0] : (individual(X0) & type_(X0)) & ! [X1] : (individual(X1) | type_(X1)))),
  inference(rectify,[],[f5])).
fof(f5,negated_conjecture,(
  ~(~? [X0] : (individual(X0) & type_(X0)) & ! [X0] : (individual(X0) | type_(X0)))),
  inference(negated_conjecture,[],[f4])).
fof(f4,conjecture,(
  ~? [X0] : (individual(X0) & type_(X0)) & ! [X0] : (individual(X0) | type_(X0))),
  file('/tmp/Bo1eORq5Ot/SOT_4blGUI',typesAndIndividualsPartitionEntity)).
fof(f159,plain,(
  ~spl11_1 | ~spl11_6 | ~spl11_12 | ~spl11_14),
  inference(avatar_contradiction_clause,[],[f158])).
fof(f158,plain,(
  $false | (~spl11_1 | ~spl11_6 | ~spl11_12 | ~spl11_14)),
  inference(subsumption_resolution,[],[f149,f157])).
fof(f157,plain,(
  ( ! [X0] : (~iof(X0,sK1)) ) | (~spl11_6 | ~spl11_14)),
  inference(resolution,[],[f154,f111])).
fof(f111,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) ) | ~spl11_6),
  inference(avatar_component_clause,[],[f110])).
fof(f110,plain,(
  spl11_6 <=> ! [X1,X0] : (~iof(X1,X0) | ~individual(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_6])])).
fof(f154,plain,(
  individual(sK1) | ~spl11_14),
  inference(avatar_component_clause,[],[f153])).
fof(f149,plain,(
  iof(sK7(sK1),sK1) | (~spl11_1 | ~spl11_12)),
  inference(resolution,[],[f92,f135])).
fof(f135,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK7(X0),X0)) ) | ~spl11_12),
  inference(avatar_component_clause,[],[f134])).
fof(f134,plain,(
  spl11_12 <=> ! [X0] : (iof(sK7(X0),X0) | ~type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_12])])).
fof(f92,plain,(
  type_(sK1) | ~spl11_1),
  inference(avatar_component_clause,[],[f91])).
fof(f91,plain,(
  spl11_1 <=> type_(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_1])])).
fof(f155,plain,(
  spl11_14 | spl11_3 | ~spl11_13),
  inference(avatar_split_clause,[],[f148,f139,f98,f153])).
fof(f139,plain,(
  spl11_13 <=> ! [X0] : (individual(X0) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_13])])).
fof(f148,plain,(
  individual(sK1) | (spl11_3 | ~spl11_13)),
  inference(subsumption_resolution,[],[f50,f142])).
fof(f142,plain,(
  individual(sK0) | (spl11_3 | ~spl11_13)),
  inference(resolution,[],[f140,f99])).
fof(f99,plain,(
  ~type_(sK0) | spl11_3),
  inference(avatar_component_clause,[],[f98])).
fof(f140,plain,(
  ( ! [X0] : (type_(X0) | individual(X0)) ) | ~spl11_13),
  inference(avatar_component_clause,[],[f139])).
fof(f50,plain,(
  ~individual(sK0) | individual(sK1)),
  inference(cnf_transformation,[],[f36])).
fof(f147,plain,(
  spl11_2 | spl11_3 | ~spl11_13),
% Refutation found. Thanks to Tanya!
  inference(avatar_contradiction_clause,[],[f146])).
fof(f146,plain,(
  $false | (spl11_2 | spl11_3 | ~spl11_13)),
  inference(subsumption_resolution,[],[f142,f95])).
% SZS status Theorem for SOT_4blG
fof(f95,plain,(
  ~individual(sK0) | spl11_2),
% SZS output start Proof for SOT_4blG
  inference(avatar_component_clause,[],[f94])).
fof(f161,plain,(
fof(f94,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f96,f100,f112,f116,f132,f136,f141,f147,f155,f159,f160])).
  spl11_2 <=> individual(sK0)),
fof(f160,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl11_2])])).
  spl11_14 | ~spl11_3),
  inference(avatar_split_clause,[],[f52,f98,f153])).
fof(f141,plain,(
fof(f153,plain,(
  spl11_14 <=> individual(sK1)),
  spl11_13 | ~spl11_7 | ~spl11_11),
  introduced(avatar_definition,[new_symbols(naming,[spl11_14])])).
  inference(avatar_split_clause,[],[f137,f130,f114,f139])).
fof(f98,plain,(
  spl11_3 <=> type_(sK0)),
fof(f114,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl11_3])])).
fof(f52,plain,(
  spl11_7 <=> ! [X1,X0] : (~iof(X1,X0) | type_(X0))),
  ~type_(sK0) | individual(sK1)),
  inference(cnf_transformation,[],[f36])).
  introduced(avatar_definition,[new_symbols(naming,[spl11_7])])).
fof(f36,plain,(
  ? [X0] : (individual(X0) & type_(X0)) | ? [X1] : (~individual(X1) & ~type_(X1))),
fof(f130,plain,(
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  spl11_11 <=> ! [X0] : (iof(sK6(X0),X0) | individual(X0))),
  ~(~? [X0] : (individual(X0) & type_(X0)) & ! [X1] : (individual(X1) | type_(X1)))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_11])])).
  inference(rectify,[],[f5])).
fof(f5,negated_conjecture,(
fof(f137,plain,(
  ~(~? [X0] : (individual(X0) & type_(X0)) & ! [X0] : (individual(X0) | type_(X0)))),
  ( ! [X0] : (individual(X0) | type_(X0)) ) | (~spl11_7 | ~spl11_11)),
  inference(negated_conjecture,[],[f4])).
fof(f4,conjecture,(
  inference(resolution,[],[f131,f115])).
  ~? [X0] : (individual(X0) & type_(X0)) & ! [X0] : (individual(X0) | type_(X0))),
  file('/tmp/Bo1eORq5Ot/SOT_4blGUI',typesAndIndividualsPartitionEntity)).
fof(f115,plain,(
fof(f159,plain,(
  ~spl11_1 | ~spl11_6 | ~spl11_12 | ~spl11_14),
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) ) | ~spl11_7),
  inference(avatar_contradiction_clause,[],[f158])).
  inference(avatar_component_clause,[],[f114])).
fof(f158,plain,(
  $false | (~spl11_1 | ~spl11_6 | ~spl11_12 | ~spl11_14)),
fof(f131,plain,(
  inference(subsumption_resolution,[],[f149,f157])).
fof(f157,plain,(
  ( ! [X0] : (iof(sK6(X0),X0) | individual(X0)) ) | ~spl11_11),
  ( ! [X0] : (~iof(X0,sK1)) ) | (~spl11_6 | ~spl11_14)),
  inference(resolution,[],[f154,f111])).
  inference(avatar_component_clause,[],[f130])).
fof(f111,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) ) | ~spl11_6),
fof(f136,plain,(
  inference(avatar_component_clause,[],[f110])).
  spl11_12),
fof(f110,plain,(
  spl11_6 <=> ! [X1,X0] : (~iof(X1,X0) | ~individual(X0))),
  inference(avatar_split_clause,[],[f66,f134])).
  introduced(avatar_definition,[new_symbols(naming,[spl11_6])])).
fof(f154,plain,(
fof(f66,plain,(
  individual(sK1) | ~spl11_14),
  inference(avatar_component_clause,[],[f153])).
  ( ! [X0] : (iof(sK7(X0),X0) | ~type_(X0)) )),
fof(f149,plain,(
  iof(sK7(sK1),sK1) | (~spl11_1 | ~spl11_12)),
  inference(cnf_transformation,[],[f3])).
  inference(resolution,[],[f92,f135])).
fof(f135,plain,(
fof(f3,axiom,(
  ( ! [X0] : (~type_(X0) | iof(sK7(X0),X0)) ) | ~spl11_12),
  inference(avatar_component_clause,[],[f134])).
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
fof(f134,plain,(
  spl11_12 <=> ! [X0] : (iof(sK7(X0),X0) | ~type_(X0))),
  file('/tmp/Bo1eORq5Ot/SOT_4blGUI',a3)).
  introduced(avatar_definition,[new_symbols(naming,[spl11_12])])).
fof(f132,plain,(
fof(f92,plain,(
  type_(sK1) | ~spl11_1),
  spl11_11),
  inference(avatar_component_clause,[],[f91])).
fof(f91,plain,(
  inference(avatar_split_clause,[],[f63,f130])).
  spl11_1 <=> type_(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_1])])).
fof(f63,plain,(
fof(f155,plain,(
  spl11_14 | spl11_3 | ~spl11_13),
  ( ! [X0] : (iof(sK6(X0),X0) | individual(X0)) )),
  inference(avatar_split_clause,[],[f148,f139,f98,f153])).
  inference(cnf_transformation,[],[f41])).
fof(f139,plain,(
  spl11_13 <=> ! [X0] : (individual(X0) | type_(X0))),
fof(f41,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl11_13])])).
fof(f148,plain,(
  ! [X0] : (individual(X0) <=> ! [X1] : ~iof(X1,X0))),
  individual(sK1) | (spl11_3 | ~spl11_13)),
  inference(subsumption_resolution,[],[f50,f142])).
  inference(ennf_transformation,[],[f2])).
fof(f142,plain,(
  individual(sK0) | (spl11_3 | ~spl11_13)),
fof(f2,axiom,(
  inference(resolution,[],[f140,f99])).
fof(f99,plain,(
  ! [X0] : (individual(X0) <=> ~? [X1] : iof(X1,X0))),
  ~type_(sK0) | spl11_3),
  inference(avatar_component_clause,[],[f98])).
  file('/tmp/Bo1eORq5Ot/SOT_4blGUI',a2)).
fof(f140,plain,(
  ( ! [X0] : (type_(X0) | individual(X0)) ) | ~spl11_13),
fof(f116,plain,(
  inference(avatar_component_clause,[],[f139])).
fof(f50,plain,(
  spl11_7),
  ~individual(sK0) | individual(sK1)),
  inference(avatar_split_clause,[],[f65,f114])).
  inference(cnf_transformation,[],[f36])).
fof(f147,plain,(
fof(f65,plain,(
  spl11_2 | spl11_3 | ~spl11_13),
  inference(avatar_contradiction_clause,[],[f146])).
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) )),
fof(f146,plain,(
  $false | (spl11_2 | spl11_3 | ~spl11_13)),
  inference(cnf_transformation,[],[f3])).
  inference(subsumption_resolution,[],[f142,f95])).
fof(f95,plain,(
fof(f112,plain,(
  ~individual(sK0) | spl11_2),
  spl11_6),
  inference(avatar_component_clause,[],[f94])).
fof(f94,plain,(
  inference(avatar_split_clause,[],[f64,f110])).
  spl11_2 <=> individual(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_2])])).
fof(f64,plain,(
fof(f141,plain,(
  spl11_13 | ~spl11_7 | ~spl11_11),
  ( ! [X0,X1] : (~iof(X1,X0) | ~individual(X0)) )),
  inference(avatar_split_clause,[],[f137,f130,f114,f139])).
fof(f114,plain,(
  inference(cnf_transformation,[],[f41])).
  spl11_7 <=> ! [X1,X0] : (~iof(X1,X0) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_7])])).
fof(f100,plain,(
fof(f130,plain,(
  spl11_11 <=> ! [X0] : (iof(sK6(X0),X0) | individual(X0))),
  spl11_1 | ~spl11_3),
  introduced(avatar_definition,[new_symbols(naming,[spl11_11])])).
fof(f137,plain,(
  inference(avatar_split_clause,[],[f51,f98,f91])).
  ( ! [X0] : (individual(X0) | type_(X0)) ) | (~spl11_7 | ~spl11_11)),
  inference(resolution,[],[f131,f115])).
fof(f51,plain,(
fof(f115,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) ) | ~spl11_7),
  ~type_(sK0) | type_(sK1)),
  inference(avatar_component_clause,[],[f114])).
  inference(cnf_transformation,[],[f36])).
fof(f131,plain,(
  ( ! [X0] : (iof(sK6(X0),X0) | individual(X0)) ) | ~spl11_11),
fof(f96,plain,(
  inference(avatar_component_clause,[],[f130])).
fof(f136,plain,(
  spl11_1 | ~spl11_2),
  spl11_12),
  inference(avatar_split_clause,[],[f66,f134])).
  inference(avatar_split_clause,[],[f49,f94,f91])).
fof(f66,plain,(
  ( ! [X0] : (iof(sK7(X0),X0) | ~type_(X0)) )),
fof(f49,plain,(
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ~individual(sK0) | type_(sK1)),
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/Bo1eORq5Ot/SOT_4blGUI',a3)).
  inference(cnf_transformation,[],[f36])).
fof(f132,plain,(
  spl11_11),
% SZS output end Proof for SOT_4blG
  inference(avatar_split_clause,[],[f63,f130])).
fof(f63,plain,(
% ------------------------------
  ( ! [X0] : (iof(sK6(X0),X0) | individual(X0)) )),
  inference(cnf_transformation,[],[f41])).
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
fof(f41,plain,(
  ! [X0] : (individual(X0) <=> ! [X1] : ~iof(X1,X0))),
% Termination reason: Refutation
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(

  ! [X0] : (individual(X0) <=> ~? [X1] : iof(X1,X0))),
  file('/tmp/Bo1eORq5Ot/SOT_4blGUI',a2)).
% Memory used [KB]: 5373
fof(f116,plain,(
  spl11_7),
% Time elapsed: 0.005 s
  inference(avatar_split_clause,[],[f65,f114])).
% ------------------------------
fof(f65,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) )),
% ------------------------------
  inference(cnf_transformation,[],[f3])).
fof(f112,plain,(
% Success in time 0.047 s
  spl11_6),
  inference(avatar_split_clause,[],[f64,f110])).
fof(f64,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | ~individual(X0)) )),
  inference(cnf_transformation,[],[f41])).
fof(f100,plain,(
  spl11_1 | ~spl11_3),
  inference(avatar_split_clause,[],[f51,f98,f91])).
fof(f51,plain,(
  ~type_(sK0) | type_(sK1)),
  inference(cnf_transformation,[],[f36])).
fof(f96,plain,(
  spl11_1 | ~spl11_2),
  inference(avatar_split_clause,[],[f49,f94,f91])).
fof(f49,plain,(
  ~individual(sK0) | type_(sK1)),
  inference(cnf_transformation,[],[f36])).
% SZS output end Proof for SOT_4blG
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
% Success in time 0.068 s
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 15 seconds
# AutoSched0-Mode selected heuristic G_E___208_C02CMA_F1_SE_CS_SP_PS_S5PRR_RG_S04AN
# and selection function SelectComplexExceptUniqMaxHorn.
#
# Preprocessing time       : 0.020 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(typesAndIndividualsPartitionEntity, conjecture, (![X1]:(type_(X1)|individual(X1))&~(?[X1]:(type_(X1)&individual(X1)))), file('/tmp/Bo1eORq5Ot/SOT_4blGUI', typesAndIndividualsPartitionEntity)).
fof(a3, axiom, ![X1]:(type_(X1)<=>?[X2]:iof(X2,X1)), file('/tmp/Bo1eORq5Ot/SOT_4blGUI', a3)).
fof(a2, axiom, ![X1]:(individual(X1)<=>~(?[X2]:iof(X2,X1))), file('/tmp/Bo1eORq5Ot/SOT_4blGUI', a2)).
fof(c_0_3, negated_conjecture, ~((![X1]:(type_(X1)|individual(X1))&~(?[X1]:(type_(X1)&individual(X1))))), inference(assume_negation,[status(cth)],[typesAndIndividualsPartitionEntity])).
fof(c_0_4, plain, ![X16, X18, X19]:((~type_(X16)|iof(esk2_1(X16),X16))&(~iof(X19,X18)|type_(X18))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a3])])])])])).
fof(c_0_5, plain, ![X12, X13, X14]:((~individual(X12)|~iof(X13,X12))&(iof(esk1_1(X14),X14)|individual(X14))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a2])])])])])).
fof(c_0_6, negated_conjecture, (((type_(esk4_0)|~type_(esk3_0))&(individual(esk4_0)|~type_(esk3_0)))&((type_(esk4_0)|~individual(esk3_0))&(individual(esk4_0)|~individual(esk3_0)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
cnf(c_0_7, plain, (type_(X2)|~iof(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, plain, (iof(esk1_1(X1),X1)|individual(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (~individual(X1)|~iof(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10, plain, (iof(esk2_1(X1),X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_11, negated_conjecture, (type_(esk4_0)|~type_(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12, plain, (type_(X1)|individual(X1)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_13, negated_conjecture, (type_(esk4_0)|~individual(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14, negated_conjecture, (individual(esk4_0)|~type_(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15, negated_conjecture, (individual(esk4_0)|~individual(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16, plain, (~type_(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_17, negated_conjecture, (type_(esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])).
cnf(c_0_18, negated_conjecture, (individual(esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_12]), c_0_15])).
cnf(c_0_19, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 20
# Proof object clause steps            : 13
# Proof object formula steps           : 7
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 78
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 73
# Processed clauses                    : 112
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 111
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 3
# Generated clauses                    : 21
# ...of the previous two non-trivial   : 14
# Contextual simplify-reflections      : 15
# Paramodulations                      : 20
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 28
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 25
# Current number of unprocessed clauses: 48
# ...number of literals in the above   : 159
# Current number of archived formulas  : 0
# Current number of archived clauses   : 83
# Clause-clause subsumption calls (NU) : 1550
# Rec. Clause-clause subsumption calls : 999
# Non-unit clause-clause subsumptions  : 16
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5366

# -------------------------------------------------
# User time                : 0.023 s
# System time              : 0.006 s
# Total time               : 0.029 s
# Maximum resident set size: 1600 pages
% START OF SYSTEM OUTPUT
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_4blG
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_4blG
% SZS output start Proof for SOT_4blG
fof(f161,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f96,f100,f112,f116,f132,f136,f141,f147,f155,f159,f160])).
fof(f160,plain,(
  spl11_14 | ~spl11_3),
  inference(avatar_split_clause,[],[f52,f98,f153])).
fof(f153,plain,(
  spl11_14 <=> individual(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_14])])).
fof(f98,plain,(
  spl11_3 <=> type_(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_3])])).
fof(f52,plain,(
  ~type_(sK0) | individual(sK1)),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ? [X0] : (individual(X0) & type_(X0)) | ? [X1] : (~individual(X1) & ~type_(X1))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~(~? [X0] : (individual(X0) & type_(X0)) & ! [X1] : (individual(X1) | type_(X1)))),
  inference(rectify,[],[f5])).
fof(f5,negated_conjecture,(
  ~(~? [X0] : (individual(X0) & type_(X0)) & ! [X0] : (individual(X0) | type_(X0)))),
  inference(negated_conjecture,[],[f4])).
fof(f4,conjecture,(
  ~? [X0] : (individual(X0) & type_(X0)) & ! [X0] : (individual(X0) | type_(X0))),
  file('/tmp/Bo1eORq5Ot/SOT_4blGUI',typesAndIndividualsPartitionEntity)).
fof(f159,plain,(
  ~spl11_1 | ~spl11_6 | ~spl11_12 | ~spl11_14),
  inference(avatar_contradiction_clause,[],[f158])).
fof(f158,plain,(
  $false | (~spl11_1 | ~spl11_6 | ~spl11_12 | ~spl11_14)),
  inference(subsumption_resolution,[],[f149,f157])).
fof(f157,plain,(
  ( ! [X0] : (~iof(X0,sK1)) ) | (~spl11_6 | ~spl11_14)),
  inference(resolution,[],[f154,f111])).
fof(f111,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) ) | ~spl11_6),
  inference(avatar_component_clause,[],[f110])).
fof(f110,plain,(
  spl11_6 <=> ! [X1,X0] : (~iof(X1,X0) | ~individual(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_6])])).
fof(f154,plain,(
  individual(sK1) | ~spl11_14),
  inference(avatar_component_clause,[],[f153])).
fof(f149,plain,(
  iof(sK7(sK1),sK1) | (~spl11_1 | ~spl11_12)),
  inference(resolution,[],[f92,f135])).
fof(f135,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK7(X0),X0)) ) | ~spl11_12),
  inference(avatar_component_clause,[],[f134])).
fof(f134,plain,(
  spl11_12 <=> ! [X0] : (iof(sK7(X0),X0) | ~type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_12])])).
fof(f92,plain,(
  type_(sK1) | ~spl11_1),
  inference(avatar_component_clause,[],[f91])).
fof(f91,plain,(
  spl11_1 <=> type_(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_1])])).
fof(f155,plain,(
  spl11_14 | spl11_3 | ~spl11_13),
  inference(avatar_split_clause,[],[f148,f139,f98,f153])).
fof(f139,plain,(
  spl11_13 <=> ! [X0] : (individual(X0) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_13])])).
fof(f148,plain,(
  individual(sK1) | (spl11_3 | ~spl11_13)),
  inference(subsumption_resolution,[],[f50,f142])).
fof(f142,plain,(
  individual(sK0) | (spl11_3 | ~spl11_13)),
  inference(resolution,[],[f140,f99])).
fof(f99,plain,(
  ~type_(sK0) | spl11_3),
  inference(avatar_component_clause,[],[f98])).
fof(f140,plain,(
  ( ! [X0] : (type_(X0) | individual(X0)) ) | ~spl11_13),
  inference(avatar_component_clause,[],[f139])).
fof(f50,plain,(
  ~individual(sK0) | individual(sK1)),
  inference(cnf_transformation,[],[f36])).
fof(f147,plain,(
  spl11_2 | spl11_3 | ~spl11_13),
  inference(avatar_contradiction_clause,[],[f146])).
fof(f146,plain,(
  $false | (spl11_2 | spl11_3 | ~spl11_13)),
  inference(subsumption_resolution,[],[f142,f95])).
fof(f95,plain,(
  ~individual(sK0) | spl11_2),
  inference(avatar_component_clause,[],[f94])).
fof(f94,plain,(
  spl11_2 <=> individual(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_2])])).
fof(f141,plain,(
  spl11_13 | ~spl11_7 | ~spl11_11),
  inference(avatar_split_clause,[],[f137,f130,f114,f139])).
fof(f114,plain,(
  spl11_7 <=> ! [X1,X0] : (~iof(X1,X0) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_7])])).
fof(f130,plain,(
  spl11_11 <=> ! [X0] : (iof(sK6(X0),X0) | individual(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_11])])).
fof(f137,plain,(
  ( ! [X0] : (individual(X0) | type_(X0)) ) | (~spl11_7 | ~spl11_11)),
  inference(resolution,[],[f131,f115])).
fof(f115,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) ) | ~spl11_7),
  inference(avatar_component_clause,[],[f114])).
fof(f131,plain,(
  ( ! [X0] : (iof(sK6(X0),X0) | individual(X0)) ) | ~spl11_11),
  inference(avatar_component_clause,[],[f130])).
fof(f136,plain,(
  spl11_12),
  inference(avatar_split_clause,[],[f66,f134])).
fof(f66,plain,(
  ( ! [X0] : (iof(sK7(X0),X0) | ~type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/Bo1eORq5Ot/SOT_4blGUI',a3)).
fof(f132,plain,(
  spl11_11),
  inference(avatar_split_clause,[],[f63,f130])).
fof(f63,plain,(
  ( ! [X0] : (iof(sK6(X0),X0) | individual(X0)) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : (individual(X0) <=> ! [X1] : ~iof(X1,X0))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0] : (individual(X0) <=> ~? [X1] : iof(X1,X0))),
  file('/tmp/Bo1eORq5Ot/SOT_4blGUI',a2)).
fof(f116,plain,(
  spl11_7),
  inference(avatar_split_clause,[],[f65,f114])).
fof(f65,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f112,plain,(
  spl11_6),
  inference(avatar_split_clause,[],[f64,f110])).
fof(f64,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | ~individual(X0)) )),
  inference(cnf_transformation,[],[f41])).
fof(f100,plain,(
  spl11_1 | ~spl11_3),
  inference(avatar_split_clause,[],[f51,f98,f91])).
fof(f51,plain,(
  ~type_(sK0) | type_(sK1)),
  inference(cnf_transformation,[],[f36])).
fof(f96,plain,(
  spl11_1 | ~spl11_2),
  inference(avatar_split_clause,[],[f49,f94,f91])).
fof(f49,plain,(
  ~individual(sK0) | type_(sK1)),
  inference(cnf_transformation,[],[f36])).
% SZS output end Proof for SOT_4blG
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
% Success in time 0.086 s
% START OF SYSTEM OUTPUT
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 15 seconds
# AutoSched0-Mode selected heuristic G_E___208_C02CMA_F1_SE_CS_SP_PS_S5PRR_RG_S04AN
# and selection function SelectComplexExceptUniqMaxHorn.
#
# Preprocessing time       : 0.016 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(typesAndIndividualsPartitionEntity, conjecture, (![X1]:(type_(X1)|individual(X1))&~(?[X1]:(type_(X1)&individual(X1)))), file('/tmp/Bo1eORq5Ot/SOT_4blGUI', typesAndIndividualsPartitionEntity)).
fof(a3, axiom, ![X1]:(type_(X1)<=>?[X2]:iof(X2,X1)), file('/tmp/Bo1eORq5Ot/SOT_4blGUI', a3)).
fof(a2, axiom, ![X1]:(individual(X1)<=>~(?[X2]:iof(X2,X1))), file('/tmp/Bo1eORq5Ot/SOT_4blGUI', a2)).
fof(c_0_3, negated_conjecture, ~((![X1]:(type_(X1)|individual(X1))&~(?[X1]:(type_(X1)&individual(X1))))), inference(assume_negation,[status(cth)],[typesAndIndividualsPartitionEntity])).
fof(c_0_4, plain, ![X16, X18, X19]:((~type_(X16)|iof(esk2_1(X16),X16))&(~iof(X19,X18)|type_(X18))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a3])])])])])).
fof(c_0_5, plain, ![X12, X13, X14]:((~individual(X12)|~iof(X13,X12))&(iof(esk1_1(X14),X14)|individual(X14))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a2])])])])])).
fof(c_0_6, negated_conjecture, (((type_(esk4_0)|~type_(esk3_0))&(individual(esk4_0)|~type_(esk3_0)))&((type_(esk4_0)|~individual(esk3_0))&(individual(esk4_0)|~individual(esk3_0)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
cnf(c_0_7, plain, (type_(X2)|~iof(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, plain, (iof(esk1_1(X1),X1)|individual(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (~individual(X1)|~iof(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10, plain, (iof(esk2_1(X1),X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_11, negated_conjecture, (type_(esk4_0)|~type_(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12, plain, (type_(X1)|individual(X1)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_13, negated_conjecture, (type_(esk4_0)|~individual(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14, negated_conjecture, (individual(esk4_0)|~type_(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15, negated_conjecture, (individual(esk4_0)|~individual(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16, plain, (~type_(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_17, negated_conjecture, (type_(esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])).
cnf(c_0_18, negated_conjecture, (individual(esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_12]), c_0_15])).
cnf(c_0_19, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 20
# Proof object clause steps            : 13
# Proof object formula steps           : 7
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 78
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 73
# Processed clauses                    : 112
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 111
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 3
# Generated clauses                    : 21
# ...of the previous two non-trivial   : 14
# Contextual simplify-reflections      : 15
# Paramodulations                      : 20
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 28
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 25
# Current number of unprocessed clauses: 48
# ...number of literals in the above   : 159
# Current number of archived formulas  : 0
# Current number of archived clauses   : 83
# Clause-clause subsumption calls (NU) : 1550
# Rec. Clause-clause subsumption calls : 999
# Non-unit clause-clause subsumptions  : 16
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5366

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.006 s
# Total time               : 0.026 s
# Maximum resident set size: 1600 pages
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_4blG
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_4blG
% SZS output start Proof for SOT_4blG
fof(f161,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f96,f100,f112,f116,f132,f136,f141,f147,f155,f159,f160])).
fof(f160,plain,(
  spl11_14 | ~spl11_3),
  inference(avatar_split_clause,[],[f52,f98,f153])).
fof(f153,plain,(
  spl11_14 <=> individual(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_14])])).
fof(f98,plain,(
  spl11_3 <=> type_(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_3])])).
fof(f52,plain,(
  ~type_(sK0) | individual(sK1)),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ? [X0] : (individual(X0) & type_(X0)) | ? [X1] : (~individual(X1) & ~type_(X1))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~(~? [X0] : (individual(X0) & type_(X0)) & ! [X1] : (individual(X1) | type_(X1)))),
  inference(rectify,[],[f5])).
fof(f5,negated_conjecture,(
  ~(~? [X0] : (individual(X0) & type_(X0)) & ! [X0] : (individual(X0) | type_(X0)))),
  inference(negated_conjecture,[],[f4])).
fof(f4,conjecture,(
  ~? [X0] : (individual(X0) & type_(X0)) & ! [X0] : (individual(X0) | type_(X0))),
  file('/tmp/Bo1eORq5Ot/SOT_4blGUI',typesAndIndividualsPartitionEntity)).
fof(f159,plain,(
  ~spl11_1 | ~spl11_6 | ~spl11_12 | ~spl11_14),
  inference(avatar_contradiction_clause,[],[f158])).
fof(f158,plain,(
  $false | (~spl11_1 | ~spl11_6 | ~spl11_12 | ~spl11_14)),
  inference(subsumption_resolution,[],[f149,f157])).
fof(f157,plain,(
  ( ! [X0] : (~iof(X0,sK1)) ) | (~spl11_6 | ~spl11_14)),
  inference(resolution,[],[f154,f111])).
fof(f111,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) ) | ~spl11_6),
  inference(avatar_component_clause,[],[f110])).
fof(f110,plain,(
  spl11_6 <=> ! [X1,X0] : (~iof(X1,X0) | ~individual(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_6])])).
fof(f154,plain,(
  individual(sK1) | ~spl11_14),
  inference(avatar_component_clause,[],[f153])).
fof(f149,plain,(
  iof(sK7(sK1),sK1) | (~spl11_1 | ~spl11_12)),
  inference(resolution,[],[f92,f135])).
fof(f135,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK7(X0),X0)) ) | ~spl11_12),
  inference(avatar_component_clause,[],[f134])).
fof(f134,plain,(
  spl11_12 <=> ! [X0] : (iof(sK7(X0),X0) | ~type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_12])])).
fof(f92,plain,(
  type_(sK1) | ~spl11_1),
  inference(avatar_component_clause,[],[f91])).
fof(f91,plain,(
  spl11_1 <=> type_(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_1])])).
fof(f155,plain,(
  spl11_14 | spl11_3 | ~spl11_13),
  inference(avatar_split_clause,[],[f148,f139,f98,f153])).
fof(f139,plain,(
  spl11_13 <=> ! [X0] : (individual(X0) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_13])])).
fof(f148,plain,(
  individual(sK1) | (spl11_3 | ~spl11_13)),
  inference(subsumption_resolution,[],[f50,f142])).
fof(f142,plain,(
  individual(sK0) | (spl11_3 | ~spl11_13)),
  inference(resolution,[],[f140,f99])).
fof(f99,plain,(
  ~type_(sK0) | spl11_3),
  inference(avatar_component_clause,[],[f98])).
fof(f140,plain,(
  ( ! [X0] : (type_(X0) | individual(X0)) ) | ~spl11_13),
  inference(avatar_component_clause,[],[f139])).
fof(f50,plain,(
  ~individual(sK0) | individual(sK1)),
  inference(cnf_transformation,[],[f36])).
fof(f147,plain,(
  spl11_2 | spl11_3 | ~spl11_13),
  inference(avatar_contradiction_clause,[],[f146])).
fof(f146,plain,(
  $false | (spl11_2 | spl11_3 | ~spl11_13)),
  inference(subsumption_resolution,[],[f142,f95])).
fof(f95,plain,(
  ~individual(sK0) | spl11_2),
  inference(avatar_component_clause,[],[f94])).
fof(f94,plain,(
  spl11_2 <=> individual(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_2])])).
fof(f141,plain,(
  spl11_13 | ~spl11_7 | ~spl11_11),
  inference(avatar_split_clause,[],[f137,f130,f114,f139])).
fof(f114,plain,(
  spl11_7 <=> ! [X1,X0] : (~iof(X1,X0) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_7])])).
fof(f130,plain,(
  spl11_11 <=> ! [X0] : (iof(sK6(X0),X0) | individual(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_11])])).
fof(f137,plain,(
  ( ! [X0] : (individual(X0) | type_(X0)) ) | (~spl11_7 | ~spl11_11)),
  inference(resolution,[],[f131,f115])).
fof(f115,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) ) | ~spl11_7),
  inference(avatar_component_clause,[],[f114])).
fof(f131,plain,(
  ( ! [X0] : (iof(sK6(X0),X0) | individual(X0)) ) | ~spl11_11),
  inference(avatar_component_clause,[],[f130])).
fof(f136,plain,(
  spl11_12),
  inference(avatar_split_clause,[],[f66,f134])).
fof(f66,plain,(
  ( ! [X0] : (iof(sK7(X0),X0) | ~type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/Bo1eORq5Ot/SOT_4blGUI',a3)).
fof(f132,plain,(
  spl11_11),
  inference(avatar_split_clause,[],[f63,f130])).
fof(f63,plain,(
  ( ! [X0] : (iof(sK6(X0),X0) | individual(X0)) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : (individual(X0) <=> ! [X1] : ~iof(X1,X0))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0] : (individual(X0) <=> ~? [X1] : iof(X1,X0))),
  file('/tmp/Bo1eORq5Ot/SOT_4blGUI',a2)).
fof(f116,plain,(
  spl11_7),
  inference(avatar_split_clause,[],[f65,f114])).
fof(f65,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f112,plain,(
  spl11_6),
  inference(avatar_split_clause,[],[f64,f110])).
fof(f64,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | ~individual(X0)) )),
  inference(cnf_transformation,[],[f41])).
fof(f100,plain,(
  spl11_1 | ~spl11_3),
  inference(avatar_split_clause,[],[f51,f98,f91])).
fof(f51,plain,(
  ~type_(sK0) | type_(sK1)),
  inference(cnf_transformation,[],[f36])).
fof(f96,plain,(
  spl11_1 | ~spl11_2),
  inference(avatar_split_clause,[],[f49,f94,f91])).
fof(f49,plain,(
  ~individual(sK0) | type_(sK1)),
  inference(cnf_transformation,[],[f36])).
% SZS output end Proof for SOT_4blG
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
% Success in time 0.053 s

% END OF SYSTEM OUTPUT
% RESULT: SOT_4blGUI - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.19 
% OUTPUT: SOT_4blGUI - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.19 

% END OF SYSTEM OUTPUT
% RESULT: SOT_4blGUI - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.07 
% OUTPUT: SOT_4blGUI - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.07 

% END OF SYSTEM OUTPUT
% RESULT: SOT_4blGUI - E---FNT-2.5 says Theorem - CPU = 0.00 WC = 0.03 
% OUTPUT: SOT_4blGUI - E---FNT-2.5 says CNFRefutation - CPU = 0.00 WC = 0.03 

% END OF SYSTEM OUTPUT
% RESULT: SOT_4blGUI - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.09 
% OUTPUT: SOT_4blGUI - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.09 

% END OF SYSTEM OUTPUT
% RESULT: SOT_4blGUI - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.05 
% OUTPUT: SOT_4blGUI - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.05 
% 10 wall clock seconds gone, waiting for system output
% 20 wall clock seconds gone, waiting for system output
% 30 wall clock seconds gone, waiting for system output
% 40 wall clock seconds gone, waiting for system output
% 50 wall clock seconds gone, waiting for system output
% 60 wall clock seconds gone, waiting for system output
% 70 wall clock seconds gone, waiting for system output
% 80 wall clock seconds gone, waiting for system output
% 90 wall clock seconds gone, waiting for system output
% 100 wall clock seconds gone, waiting for system output
% 110 wall clock seconds gone, waiting for system output
% 120 wall clock seconds gone, waiting for system output
% 130 wall clock seconds gone, waiting for system output
% 140 wall clock seconds gone, waiting for system output
% 150 wall clock seconds gone, waiting for system output
% 160 wall clock seconds gone, waiting for system output
% 170 wall clock seconds gone, waiting for system output
% 180 wall clock seconds gone, waiting for system output
% 190 wall clock seconds gone, waiting for system output
% 200 wall clock seconds gone, waiting for system output
% 210 wall clock seconds gone, waiting for system output
% 220 wall clock seconds gone, waiting for system output
% 230 wall clock seconds gone, waiting for system output
% 240 wall clock seconds gone, waiting for system output
% 250 wall clock seconds gone, waiting for system output
% 260 wall clock seconds gone, waiting for system output
% 270 wall clock seconds gone, waiting for system output
% 280 wall clock seconds gone, waiting for system output
% 290 wall clock seconds gone, waiting for system output
% 300 wall clock seconds gone, waiting for system output
% 310 wall clock seconds gone, waiting for system output
% 320 wall clock seconds gone, waiting for system output
% 330 wall clock seconds gone, waiting for system output
% 340 wall clock seconds gone, waiting for system output
% 350 wall clock seconds gone, waiting for system output
% 360 wall clock seconds gone, waiting for system output
% 370 wall clock seconds gone, waiting for system output
% 380 wall clock seconds gone, waiting for system output
% 390 wall clock seconds gone, waiting for system output
% 400 wall clock seconds gone, waiting for system output
% 410 wall clock seconds gone, waiting for system output
% 420 wall clock seconds gone, waiting for system output
% 430 wall clock seconds gone, waiting for system output
% 440 wall clock seconds gone, waiting for system output
% 450 wall clock seconds gone, waiting for system output
% 460 wall clock seconds gone, waiting for system output
% 470 wall clock seconds gone, waiting for system output
% 480 wall clock seconds gone, waiting for system output
% 490 wall clock seconds gone, waiting for system output
% 500 wall clock seconds gone, waiting for system output
% 510 wall clock seconds gone, waiting for system output
% 520 wall clock seconds gone, waiting for system output
% 530 wall clock seconds gone, waiting for system output
% 540 wall clock seconds gone, waiting for system output
% 550 wall clock seconds gone, waiting for system output
% 560 wall clock seconds gone, waiting for system output
% 570 wall clock seconds gone, waiting for system output
% 580 wall clock seconds gone, waiting for system output
% 590 wall clock seconds gone, waiting for system output
% 600 wall clock seconds gone, waiting for system output
% 610 wall clock seconds gone, waiting for system output
% 620 wall clock seconds gone, waiting for system output
% 630 wall clock seconds gone, waiting for system output
% 640 wall clock seconds gone, waiting for system output
% 650 wall clock seconds gone, waiting for system output
% 660 wall clock seconds gone, waiting for system output
% 670 wall clock seconds gone, waiting for system output
% 680 wall clock seconds gone, waiting for system output
% 690 wall clock seconds gone, waiting for system output
% 700 wall clock seconds gone, waiting for system output
% 710 wall clock seconds gone, waiting for system output
% 720 wall clock seconds gone, waiting for system output
% 730 wall clock seconds gone, waiting for system output
% 740 wall clock seconds gone, waiting for system output
% 750 wall clock seconds gone, waiting for system output
% 760 wall clock seconds gone, waiting for system output
% 770 wall clock seconds gone, waiting for system output
% 780 wall clock seconds gone, waiting for system output
% 790 wall clock seconds gone, waiting for system output
% 800 wall clock seconds gone, waiting for system output
% 810 wall clock seconds gone, waiting for system output
% 820 wall clock seconds gone, waiting for system output
% 830 wall clock seconds gone, waiting for system output
% 840 wall clock seconds gone, waiting for system output
% 850 wall clock seconds gone, waiting for system output
% 860 wall clock seconds gone, waiting for system output
% 870 wall clock seconds gone, waiting for system output
% 880 wall clock seconds gone, waiting for system output
% 890 wall clock seconds gone, waiting for system output
% 900 wall clock seconds gone, waiting for system output
% 910 wall clock seconds gone, waiting for system output
% 920 wall clock seconds gone, waiting for system output
% 930 wall clock seconds gone, waiting for system output
% 940 wall clock seconds gone, waiting for system output
% 950 wall clock seconds gone, waiting for system output
% 960 wall clock seconds gone, waiting for system output
% 970 wall clock seconds gone, waiting for system output
% 980 wall clock seconds gone, waiting for system output
% 990 wall clock seconds gone, waiting for system output
% 1000 wall clock seconds gone, waiting for system output
% 1010 wall clock seconds gone, waiting for system output
% 1020 wall clock seconds gone, waiting for system output
% 1030 wall clock seconds gone, waiting for system output
% 1040 wall clock seconds gone, waiting for system output
% 1050 wall clock seconds gone, waiting for system output
% 1060 wall clock seconds gone, waiting for system output
% 1070 wall clock seconds gone, waiting for system output
% 1080 wall clock seconds gone, waiting for system output
% 1090 wall clock seconds gone, waiting for system output
% 1100 wall clock seconds gone, waiting for system output
% 1110 wall clock seconds gone, waiting for system output
% 1120 wall clock seconds gone, waiting for system output
% 1130 wall clock seconds gone, waiting for system output
% 1140 wall clock seconds gone, waiting for system output
% 1150 wall clock seconds gone, waiting for system output
% 1160 wall clock seconds gone, waiting for system output
% 1170 wall clock seconds gone, waiting for system output
% 1180 wall clock seconds gone, waiting for system output
% 1190 wall clock seconds gone, waiting for system output
% 1200 wall clock seconds gone, waiting for system output
% 1210 wall clock seconds gone, waiting for system output
% 1220 wall clock seconds gone, waiting for system output
% 1230 wall clock seconds gone, waiting for system output
% 1240 wall clock seconds gone, waiting for system output
% 1250 wall clock seconds gone, waiting for system output
% 1260 wall clock seconds gone, waiting for system output
% 1270 wall clock seconds gone, waiting for system output
% 1280 wall clock seconds gone, waiting for system output
% 1290 wall clock seconds gone, waiting for system output
% 1300 wall clock seconds gone, waiting for system output
% 1310 wall clock seconds gone, waiting for system output
% 1320 wall clock seconds gone, waiting for system output
% 1330 wall clock seconds gone, waiting for system output
% 1340 wall clock seconds gone, waiting for system output
% 1350 wall clock seconds gone, waiting for system output
% 1360 wall clock seconds gone, waiting for system output
% 1370 wall clock seconds gone, waiting for system output
% 1380 wall clock seconds gone, waiting for system output
