<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>System Results</title>
<meta name="ROBOTS" content="INDEX,NOFOLLOW" />
<script type="text/javascript">//<![CDATA[

  var tstp_dump;
  function openSoTSTP (dump) {   
    var tstp_url = 'http://www.tptp.org/cgi-bin/SystemOnTSTP';
    var tstp_browser = window.open(tstp_url, '_blank');
    tstp_dump = dump;
  }
  function getTSTPDump () {
    return tstp_dump;
  }

//]]></script>
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>

% SZS start RequiredInformation<BR>
% Congratulations - you have become a registered power user of SystemOnTPTP, at IP address 191.6.38.118.<BR>
% Please consider donating to the TPTP project - see www.tptp.org for details.<BR>
% When you donate this message will disappear.<BR>
% If you do not donate a random delay might be added to your processing time.<BR>
% SZS end RequiredInformation<BR>
<PRE>
Getting system recommendations for generic problem /tmp/mI9PT9m0FF/SOT_U48AXO
This is a FOF_THM_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.95
Vampire---TF1-4.5    recommended  0.95
Vampire---4.5        recommended  0.95
Vampire---SAT-4.5    recommended  0.95
ET---2.0             recommended  0.80
E---2.5              recommended  0.78
E---FNT-2.5          recommended  0.78
CSE_E---1.0          recommended  0.76
E---FNT-2.5          recommended  0.74
Isabelle---2TH0      recommended  0.69
Isabelle---2TF0      recommended  0.69
Isabelle---2FOF      recommended  0.69
Isabelle---2016      recommended  0.69
Isabelle-HOT---2016  recommended  0.68
CVC4---1.7           recommended  0.68
CVC4---SAT-1.7       recommended  0.68
Leo-III---1.4        recommended  0.61
Vampire---THF-4.5    recommended  0.59
Vampire---TF1-4.5    recommended  0.59
Vampire---SAT-4.5    recommended  0.59
iProver---3.0        recommended  0.58
iProver---SAT-3.0    recommended  0.58
Fampire---1.3        recommended  0.57
SInE---0.4           recommended  0.53
iProver---SAT-3.0    recommended  0.52
ePrincess---1.0      recommended  0.50
Princess---170717    recommended  0.48
SPASS---3.9          recommended  0.48
Z3---4.4.1           recommended  0.46
Zipperpin---1.5      recommended  0.40
CVC4---SAT-1.7       recommended  0.40
Prover9---1109a      recommended  0.39
Equinox---6.0.1a     recommended  0.37
LEO-II---1.7.0       recommended  0.35
nanoCoP---1.1        recommended  0.34
Muscadet---4.5       recommended  0.25
Twee---2.2           recommended  0.23
Metis---2.4          recommended  0.23
Geo-III---2018C      recommended  0.18
iProverMo---2.5-0.1  recommended  0.14
Beagle---0.9.47      subsumed     0.41
leanCoP---2.2        subsumed     0.36
E-Darwin---1.5       subsumed     0.31
CSE---1.0            subsumed     0.28
Darwin---1.4.5       subsumed     0.28
Bliksem---1.12       subsumed     0.21
Zenon---0.7.1        subsumed     0.16
ZenonModulo---0.4.2  subsumed     0.15
SOS---2.0            subsumed     0.11
Otter---3.3          subsumed     0.04
This is a FOF_CSA_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.98
Vampire---TF1-4.5    recommended  0.98
Vampire---SAT-4.5    recommended  0.98
Paradox---4.0        recommended  0.60
Nitpick---2016       subsumed     0.55
DarwinFM---1.4.5     subsumed     0.43
CVC4---SAT-1.7       subsumed     0.40
Vampire---THF-4.5    subsumed     0.34
Vampire---TF1-4.5    subsumed     0.34
Vampire---4.5        subsumed     0.34
Vampire---SAT-4.5    subsumed     0.34
E---2.5              subsumed     0.34
E---FNT-2.5          subsumed     0.34
E---FNT-2.5          subsumed     0.34
iProver---SAT-3.0    subsumed     0.32
Metis---2.4          subsumed     0.32
SPASS---3.9          subsumed     0.28
Beagle---0.9.47      subsumed     0.26
Geo-III---2018C      subsumed     0.25
Z3---4.4.1           subsumed     0.25
Princess---170717    subsumed     0.25
Equinox---6.0.1a     subsumed     0.19
iProver---3.0        subsumed     0.11
iProver---SAT-3.0    subsumed     0.11
Darwin---1.4.5       subsumed     0.06
E-Darwin---1.5       subsumed     0.04
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use ET,2.0
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use CSE_E,1.0
CSE_E---1.0 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use Isabelle,2TH0
Isabelle---2016 can be used
Check if we can use Isabelle,2TF0
Isabelle---2016 can be used
Check if we can use Isabelle,2FOF
Isabelle---2016 can be used
Check if we can use Isabelle,2016
Isabelle---2016 can be used
Check if we can use Isabelle-HOT,2016
Isabelle-HOT---2016 can be used
Check if we can use CVC4,1.7
CVC4---SAT-1.7 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Leo-III,1.4
Leo-III---1.4 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Fampire,1.3
Check if we can use SInE,0.4
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use ePrincess,1.0
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Zipperpin,1.5
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Prover9,1109a
Prover9---1109a can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use LEO-II,1.7.0
LEO-II---1.7.0 can be used
Check if we can use nanoCoP,1.1
nanoCoP---1.1 can be used
Check if we can use Muscadet,4.5
Muscadet---4.5 can be used
Check if we can use Twee,2.2
Twee---2.2 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use iProverMo,2.5-0.1
iProverMo---2.5-0.1 can be used
Check if we can use Beagle,0.9.47
Check if we can use leanCoP,2.2
leanCoP---2.2 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Check if we can use CSE,1.0
CSE---1.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use Bliksem,1.12
Bliksem---1.12 can be used
Check if we can use Zenon,0.7.1
Zenon---0.7.1 can be used
Check if we can use ZenonModulo,0.4.2
Check if we can use SOS,2.0
SOS---2.0 can be used
Check if we can use Otter,3.3
Otter---3.3 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use Paradox,4.0
Paradox---4.0 can be used
Check if we can use Nitpick,2016
Nitpick---2016 can be used
Check if we can use DarwinFM,1.4.5
DarwinFM---1.4.5 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Beagle,0.9.47
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
E---FNT-2.5 will be used
E---FNT-2.5 will be used
Vampire---SAT-4.5 is given 499 seconds in -cE mode
Vampire---SAT-4.5 is given 250 seconds in -cE mode
Vampire---SAT-4.5 is given 125 seconds in -cE mode
Vampire---SAT-4.5 is given 62 seconds in -cE mode
E---FNT-2.5 is given 31 seconds in -cE mode
E---FNT-2.5 is given 32 seconds in -cE mode
Vampire---SAT-4.5 has started (PID 97211)
Vampire---SAT-4.5 has started (PID 97213)
Vampire---SAT-4.5 has started (PID 97215)
Vampire---SAT-4.5 has started (PID 97216)
E---FNT-2.5 has started (PID 97217)
E---FNT-2.5 has started (PID 97219)
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  499  /tmp/mI9PT9m0FF/SOT_U48AXO
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  250  /tmp/mI9PT9m0FF/SOT_U48AXO
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  125  /tmp/mI9PT9m0FF/SOT_U48AXO
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  62  /tmp/mI9PT9m0FF/SOT_U48AXO
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  31  /tmp/mI9PT9m0FF/SOT_U48AXO
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  32  /tmp/mI9PT9m0FF/SOT_U48AXO
% START OF SYSTEM OUTPUT
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_U48A
% START OF SYSTEM OUTPUT
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_U48A
% START OF SYSTEM OUTPUT
% START OF SYSTEM OUTPUT
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 15 seconds
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 15 seconds
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_U48A
# AutoSched0-Mode selected heuristic G_E___208_C02CMA_F1_SE_CS_SP_PS_S5PRR_RG_S04AN
# and selection function SelectComplexExceptUniqMaxHorn.
#
# Preprocessing time       : 0.032 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t2_powertypeUnique, conjecture, ![X11, X12]:(isPowertypeOf(X11,X12)=>~(?[X13]:(~(X13=X12)&isPowertypeOf(X11,X13)))), file('/tmp/mI9PT9m0FF/SOT_U48AXO', t2_powertypeUnique)).
fof(a10_isPowertypeOfDef, axiom, ![X7, X8]:(isPowertypeOf(X7,X8)<=>(type_(X7)&![X9]:(iof(X9,X7)<=>specializes(X9,X8)))), file('/tmp/mI9PT9m0FF/SOT_U48AXO', a10_isPowertypeOfDef)).
fof(a7_specDef, axiom, ![X7, X8]:(specializes(X7,X8)<=>(type_(X7)&![X4]:(iof(X4,X7)=>iof(X4,X8)))), file('/tmp/mI9PT9m0FF/SOT_U48AXO', a7_specDef)).
fof(a2, axiom, ![X1]:(individual(X1)<=>~(?[X2]:iof(X2,X1))), file('/tmp/mI9PT9m0FF/SOT_U48AXO', a2)).
fof(a3, axiom, ![X1]:(type_(X1)<=>?[X2]:iof(X2,X1)), file('/tmp/mI9PT9m0FF/SOT_U48AXO', a3)).
fof(a5_firstOrderTypeDef, axiom, ![X1]:(firstordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>individual(X2)))), file('/tmp/mI9PT9m0FF/SOT_U48AXO', a5_firstOrderTypeDef)).
fof(a8_properSpecDef, axiom, ![X7, X8]:(properSpecializes(X7,X8)<=>(specializes(X7,X8)&~(X7=X8))), file('/tmp/mI9PT9m0FF/SOT_U48AXO', a8_properSpecDef)).
fof(a9_equality, axiom, ![X7, X8]:((type_(X7)&type_(X8))=>(X7=X8<=>![X1]:(iof(X1,X7)<=>iof(X1,X8)))), file('/tmp/mI9PT9m0FF/SOT_U48AXO', a9_equality)).
fof(c_0_8, negated_conjecture, ~(![X11, X12]:(isPowertypeOf(X11,X12)=>~(?[X13]:(~(X13=X12)&isPowertypeOf(X11,X13))))), inference(assume_negation,[status(cth)],[t2_powertypeUnique])).
fof(c_0_9, plain, ![X65, X66, X67, X68, X69, X70]:(((type_(X65)|~isPowertypeOf(X65,X66))&((~iof(X67,X65)|specializes(X67,X66)|~isPowertypeOf(X65,X66))&(~specializes(X68,X66)|iof(X68,X65)|~isPowertypeOf(X65,X66))))&((~iof(esk13_2(X69,X70),X69)|~specializes(esk13_2(X69,X70),X70)|~type_(X69)|isPowertypeOf(X69,X70))&(iof(esk13_2(X69,X70),X69)|specializes(esk13_2(X69,X70),X70)|~type_(X69)|isPowertypeOf(X69,X70)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a10_isPowertypeOfDef])])])])])])).
fof(c_0_10, negated_conjecture, (isPowertypeOf(esk22_0,esk23_0)&(esk24_0!=esk23_0&isPowertypeOf(esk22_0,esk24_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])).
fof(c_0_11, plain, ![X52, X53, X54, X55, X56]:(((type_(X52)|~specializes(X52,X53))&(~iof(X54,X52)|iof(X54,X53)|~specializes(X52,X53)))&((iof(esk11_2(X55,X56),X55)|~type_(X55)|specializes(X55,X56))&(~iof(esk11_2(X55,X56),X56)|~type_(X55)|specializes(X55,X56)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a7_specDef])])])])])])).
cnf(c_0_12, plain, (specializes(X1,X3)|~iof(X1,X2)|~isPowertypeOf(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, negated_conjecture, (isPowertypeOf(esk22_0,esk24_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, plain, ![X15, X16, X17]:((~individual(X15)|~iof(X16,X15))&(iof(esk1_1(X17),X17)|individual(X17))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a2])])])])])).
fof(c_0_15, plain, ![X19, X21, X22]:((~type_(X19)|iof(esk2_1(X19),X19))&(~iof(X22,X21)|type_(X21))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a3])])])])])).
cnf(c_0_16, plain, (type_(X1)|~specializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, negated_conjecture, (specializes(X1,esk24_0)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_18, plain, (type_(X1)|~isPowertypeOf(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_19, plain, (~individual(X1)|~iof(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (iof(esk2_1(X1),X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, negated_conjecture, (type_(X1)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22, negated_conjecture, (type_(esk22_0)), inference(spm,[status(thm)],[c_0_18, c_0_13])).
fof(c_0_23, plain, ![X44, X45, X46]:(((type_(X44)|~firstordertype(X44))&(~iof(X45,X44)|individual(X45)|~firstordertype(X44)))&((iof(esk9_1(X46),X46)|~type_(X46)|firstordertype(X46))&(~individual(esk9_1(X46))|~type_(X46)|firstordertype(X46)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a5_firstOrderTypeDef])])])])])])).
cnf(c_0_24, plain, (~type_(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_25, negated_conjecture, (type_(esk2_1(esk22_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_20]), c_0_22])])).
cnf(c_0_26, plain, (individual(X1)|~iof(X1,X2)|~firstordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27, plain, (type_(X1)|~firstordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_28, plain, ![X58, X59]:(((specializes(X58,X59)|~properSpecializes(X58,X59))&(X58!=X59|~properSpecializes(X58,X59)))&(~specializes(X58,X59)|X58=X59|properSpecializes(X58,X59))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[a8_properSpecDef])])])])).
cnf(c_0_29, negated_conjecture, (isPowertypeOf(esk22_0,esk23_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_30, plain, (iof(esk9_1(X1),X1)|firstordertype(X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31, negated_conjecture, (~individual(esk2_1(esk22_0))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_32, plain, (individual(esk2_1(X1))|~firstordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_20]), c_0_27])).
cnf(c_0_33, plain, (X1=X2|properSpecializes(X1,X2)|~specializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, negated_conjecture, (specializes(X1,esk23_0)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_12, c_0_29])).
cnf(c_0_35, negated_conjecture, (firstordertype(esk22_0)|iof(esk9_1(esk22_0),esk22_0)), inference(spm,[status(thm)],[c_0_30, c_0_22])).
cnf(c_0_36, negated_conjecture, (~firstordertype(esk22_0)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_37, negated_conjecture, (X1=esk23_0|properSpecializes(X1,esk23_0)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_38, negated_conjecture, (iof(esk9_1(esk22_0),esk22_0)), inference(sr,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_39, plain, (specializes(X1,X2)|~properSpecializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_40, negated_conjecture, (esk9_1(esk22_0)=esk23_0|properSpecializes(esk9_1(esk22_0),esk23_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_41, negated_conjecture, (X1=esk24_0|properSpecializes(X1,esk24_0)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_33, c_0_17])).
cnf(c_0_42, plain, (iof(X1,X3)|~iof(X1,X2)|~specializes(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_43, negated_conjecture, (esk9_1(esk22_0)=esk23_0|specializes(esk9_1(esk22_0),esk23_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_44, negated_conjecture, (esk9_1(esk22_0)=esk24_0|properSpecializes(esk9_1(esk22_0),esk24_0)), inference(spm,[status(thm)],[c_0_41, c_0_38])).
cnf(c_0_45, negated_conjecture, (esk9_1(esk22_0)=esk23_0|iof(X1,esk23_0)|~iof(X1,esk9_1(esk22_0))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_46, negated_conjecture, (type_(esk9_1(esk22_0))), inference(spm,[status(thm)],[c_0_21, c_0_38])).
cnf(c_0_47, negated_conjecture, (esk9_1(esk22_0)=esk24_0|specializes(esk9_1(esk22_0),esk24_0)), inference(spm,[status(thm)],[c_0_39, c_0_44])).
cnf(c_0_48, plain, (type_(X2)|~iof(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_49, negated_conjecture, (esk9_1(esk22_0)=esk23_0|iof(esk2_1(esk9_1(esk22_0)),esk23_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_20]), c_0_46])])).
cnf(c_0_50, negated_conjecture, (esk9_1(esk22_0)=esk24_0|iof(X1,esk24_0)|~iof(X1,esk9_1(esk22_0))), inference(spm,[status(thm)],[c_0_42, c_0_47])).
cnf(c_0_51, negated_conjecture, (esk9_1(esk22_0)=esk23_0|type_(esk23_0)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_52, negated_conjecture, (esk9_1(esk22_0)=esk24_0|iof(esk2_1(esk9_1(esk22_0)),esk24_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_20]), c_0_46])])).
cnf(c_0_53, plain, (iof(esk11_2(X1,X2),X1)|specializes(X1,X2)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_54, negated_conjecture, (type_(esk23_0)), inference(spm,[status(thm)],[c_0_46, c_0_51])).
cnf(c_0_55, negated_conjecture, (esk9_1(esk22_0)=esk24_0|type_(esk24_0)), inference(spm,[status(thm)],[c_0_48, c_0_52])).
cnf(c_0_56, plain, (iof(X1,X3)|~specializes(X1,X2)|~isPowertypeOf(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_57, plain, (specializes(X1,X2)|~iof(esk11_2(X1,X2),X2)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_58, negated_conjecture, (specializes(esk23_0,X1)|iof(esk11_2(esk23_0,X1),esk23_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_59, negated_conjecture, (type_(esk24_0)), inference(spm,[status(thm)],[c_0_46, c_0_55])).
cnf(c_0_60, negated_conjecture, (iof(X1,esk22_0)|~specializes(X1,esk23_0)), inference(spm,[status(thm)],[c_0_56, c_0_29])).
cnf(c_0_61, negated_conjecture, (specializes(esk23_0,esk23_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_54])])).
cnf(c_0_62, negated_conjecture, (specializes(esk24_0,X1)|iof(esk11_2(esk24_0,X1),esk24_0)), inference(spm,[status(thm)],[c_0_53, c_0_59])).
cnf(c_0_63, negated_conjecture, (iof(esk23_0,esk22_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_64, negated_conjecture, (esk24_0!=esk23_0), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_65, negated_conjecture, (iof(X1,esk22_0)|~specializes(X1,esk24_0)), inference(spm,[status(thm)],[c_0_56, c_0_13])).
cnf(c_0_66, negated_conjecture, (specializes(esk24_0,esk24_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_62]), c_0_59])])).
fof(c_0_67, plain, ![X60, X61, X62, X63]:(((~iof(X62,X60)|iof(X62,X61)|X60!=X61|(~type_(X60)|~type_(X61)))&(~iof(X63,X61)|iof(X63,X60)|X60!=X61|(~type_(X60)|~type_(X61))))&((~iof(esk12_2(X60,X61),X60)|~iof(esk12_2(X60,X61),X61)|X60=X61|(~type_(X60)|~type_(X61)))&(iof(esk12_2(X60,X61),X60)|iof(esk12_2(X60,X61),X61)|X60=X61|(~type_(X60)|~type_(X61))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a9_equality])])])])])])).
cnf(c_0_68, negated_conjecture, (properSpecializes(esk23_0,esk24_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_63]), c_0_64])).
cnf(c_0_69, negated_conjecture, (iof(esk24_0,esk22_0)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_70, plain, (iof(esk12_2(X1,X2),X1)|iof(esk12_2(X1,X2),X2)|X1=X2|~type_(X1)|~type_(X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_71, negated_conjecture, (specializes(esk23_0,esk24_0)), inference(spm,[status(thm)],[c_0_39, c_0_68])).
cnf(c_0_72, negated_conjecture, (properSpecializes(esk24_0,esk23_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_69]), c_0_64])).
cnf(c_0_73, plain, (X1=X2|~iof(esk12_2(X1,X2),X1)|~iof(esk12_2(X1,X2),X2)|~type_(X1)|~type_(X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_74, negated_conjecture, (X1=esk24_0|iof(esk12_2(X1,esk24_0),esk24_0)|iof(esk12_2(X1,esk24_0),X1)|~type_(X1)), inference(spm,[status(thm)],[c_0_70, c_0_59])).
cnf(c_0_75, negated_conjecture, (iof(X1,esk24_0)|~iof(X1,esk23_0)), inference(spm,[status(thm)],[c_0_42, c_0_71])).
cnf(c_0_76, negated_conjecture, (specializes(esk24_0,esk23_0)), inference(spm,[status(thm)],[c_0_39, c_0_72])).
cnf(c_0_77, plain, (X1=X2|~iof(esk12_2(X1,X2),X2)|~iof(esk12_2(X1,X2),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_73, c_0_48]), c_0_48])).
cnf(c_0_78, negated_conjecture, (iof(esk12_2(esk23_0,esk24_0),esk24_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_54]), c_0_64]), c_0_75])).
cnf(c_0_79, negated_conjecture, (iof(X1,esk23_0)|~iof(X1,esk24_0)), inference(spm,[status(thm)],[c_0_42, c_0_76])).
cnf(c_0_80, negated_conjecture, (~iof(esk12_2(esk23_0,esk24_0),esk23_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_64])).
cnf(c_0_81, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_78]), c_0_80]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 82
# Proof object clause steps            : 65
# Proof object formula steps           : 17
# Proof object conjectures             : 48
# AutoSched0-Mode selected heuristic G_E___208_C02CMA_F1_SE_CS_SP_PS_S5PRR_RG_S04AN
# Proof object clause conjectures      : 45
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# and selection function SelectComplexExceptUniqMaxHorn.
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 43
# Proof object simplifying inferences  : 20
#
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Preprocessing time       : 0.031 s
# Initial clauses                      : 77
# Removed in clause preprocessing      : 5
# Presaturation interreduction done
# Initial clauses in saturation        : 72
# Processed clauses                    : 2014
# ...of these trivial                  : 3

# ...subsumed                          : 822
# ...remaining for further processing  : 1188
# Other redundant clauses eliminated   : 1
# Proof found!
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 35
# Backward-rewritten                   : 2
# SZS status Theorem
# Generated clauses                    : 4548
# SZS output start CNFRefutation
# ...of the previous two non-trivial   : 4237
fof(t2_powertypeUnique, conjecture, ![X11, X12]:(isPowertypeOf(X11,X12)=>~(?[X13]:(~(X13=X12)&isPowertypeOf(X11,X13)))), file('/tmp/mI9PT9m0FF/SOT_U48AXO', t2_powertypeUnique)).
# Contextual simplify-reflections      : 97
fof(a10_isPowertypeOfDef, axiom, ![X7, X8]:(isPowertypeOf(X7,X8)<=>(type_(X7)&![X9]:(iof(X9,X7)<=>specializes(X9,X8)))), file('/tmp/mI9PT9m0FF/SOT_U48AXO', a10_isPowertypeOfDef)).
# Paramodulations                      : 4546
fof(a7_specDef, axiom, ![X7, X8]:(specializes(X7,X8)<=>(type_(X7)&![X4]:(iof(X4,X7)=>iof(X4,X8)))), file('/tmp/mI9PT9m0FF/SOT_U48AXO', a7_specDef)).
# Factorizations                       : 0
fof(a2, axiom, ![X1]:(individual(X1)<=>~(?[X2]:iof(X2,X1))), file('/tmp/mI9PT9m0FF/SOT_U48AXO', a2)).
# Equation resolutions                 : 1
fof(a3, axiom, ![X1]:(type_(X1)<=>?[X2]:iof(X2,X1)), file('/tmp/mI9PT9m0FF/SOT_U48AXO', a3)).
fof(a5_firstOrderTypeDef, axiom, ![X1]:(firstordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>individual(X2)))), file('/tmp/mI9PT9m0FF/SOT_U48AXO', a5_firstOrderTypeDef)).
# Propositional unsat checks           : 0
fof(a8_properSpecDef, axiom, ![X7, X8]:(properSpecializes(X7,X8)<=>(specializes(X7,X8)&~(X7=X8))), file('/tmp/mI9PT9m0FF/SOT_U48AXO', a8_properSpecDef)).
#    Propositional check models        : 0
fof(a9_equality, axiom, ![X7, X8]:((type_(X7)&type_(X8))=>(X7=X8<=>![X1]:(iof(X1,X7)<=>iof(X1,X8)))), file('/tmp/mI9PT9m0FF/SOT_U48AXO', a9_equality)).
#    Propositional check unsatisfiable : 0
fof(c_0_8, negated_conjecture, ~(![X11, X12]:(isPowertypeOf(X11,X12)=>~(?[X13]:(~(X13=X12)&isPowertypeOf(X11,X13))))), inference(assume_negation,[status(cth)],[t2_powertypeUnique])).
#    Propositional clauses             : 0
fof(c_0_9, plain, ![X65, X66, X67, X68, X69, X70]:(((type_(X65)|~isPowertypeOf(X65,X66))&((~iof(X67,X65)|specializes(X67,X66)|~isPowertypeOf(X65,X66))&(~specializes(X68,X66)|iof(X68,X65)|~isPowertypeOf(X65,X66))))&((~iof(esk13_2(X69,X70),X69)|~specializes(esk13_2(X69,X70),X70)|~type_(X69)|isPowertypeOf(X69,X70))&(iof(esk13_2(X69,X70),X69)|specializes(esk13_2(X69,X70),X70)|~type_(X69)|isPowertypeOf(X69,X70)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a10_isPowertypeOfDef])])])])])])).
#    Propositional clauses after purity: 0
fof(c_0_10, negated_conjecture, (isPowertypeOf(esk22_0,esk23_0)&(esk24_0!=esk23_0&isPowertypeOf(esk22_0,esk24_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])).
#    Propositional unsat core size     : 0
fof(c_0_11, plain, ![X52, X53, X54, X55, X56]:(((type_(X52)|~specializes(X52,X53))&(~iof(X54,X52)|iof(X54,X53)|~specializes(X52,X53)))&((iof(esk11_2(X55,X56),X55)|~type_(X55)|specializes(X55,X56))&(~iof(esk11_2(X55,X56),X56)|~type_(X55)|specializes(X55,X56)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a7_specDef])])])])])])).
#    Propositional preprocessing time  : 0.000
cnf(c_0_12, plain, (specializes(X1,X3)|~iof(X1,X2)|~isPowertypeOf(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, negated_conjecture, (isPowertypeOf(esk22_0,esk24_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
#    Propositional encoding time       : 0.000
fof(c_0_14, plain, ![X15, X16, X17]:((~individual(X15)|~iof(X16,X15))&(iof(esk1_1(X17),X17)|individual(X17))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a2])])])])])).
#    Propositional solver time         : 0.000
fof(c_0_15, plain, ![X19, X21, X22]:((~type_(X19)|iof(esk2_1(X19),X19))&(~iof(X22,X21)|type_(X21))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a3])])])])])).
#    Success case prop preproc time    : 0.000
cnf(c_0_16, plain, (type_(X1)|~specializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
#    Success case prop encoding time   : 0.000
cnf(c_0_17, negated_conjecture, (specializes(X1,esk24_0)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
#    Success case prop solver time     : 0.000
cnf(c_0_18, plain, (type_(X1)|~isPowertypeOf(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_19, plain, (~individual(X1)|~iof(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
# Current number of processed clauses  : 1077
cnf(c_0_20, plain, (iof(esk2_1(X1),X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
#    Positive orientable unit clauses  : 39
cnf(c_0_21, negated_conjecture, (type_(X1)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
#    Positive unorientable unit clauses: 0
cnf(c_0_22, negated_conjecture, (type_(esk22_0)), inference(spm,[status(thm)],[c_0_18, c_0_13])).
#    Negative unit clauses             : 10
fof(c_0_23, plain, ![X44, X45, X46]:(((type_(X44)|~firstordertype(X44))&(~iof(X45,X44)|individual(X45)|~firstordertype(X44)))&((iof(esk9_1(X46),X46)|~type_(X46)|firstordertype(X46))&(~individual(esk9_1(X46))|~type_(X46)|firstordertype(X46)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a5_firstOrderTypeDef])])])])])])).
cnf(c_0_24, plain, (~type_(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
#    Non-unit-clauses                  : 1028
cnf(c_0_25, negated_conjecture, (type_(esk2_1(esk22_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_20]), c_0_22])])).
# Current number of unprocessed clauses: 2326
cnf(c_0_26, plain, (individual(X1)|~iof(X1,X2)|~firstordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
# ...number of literals in the above   : 8344
cnf(c_0_27, plain, (type_(X1)|~firstordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
# Current number of archived formulas  : 0
fof(c_0_28, plain, ![X58, X59]:(((specializes(X58,X59)|~properSpecializes(X58,X59))&(X58!=X59|~properSpecializes(X58,X59)))&(~specializes(X58,X59)|X58=X59|properSpecializes(X58,X59))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[a8_properSpecDef])])])])).
# Current number of archived clauses   : 111
cnf(c_0_29, negated_conjecture, (isPowertypeOf(esk22_0,esk23_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
# Clause-clause subsumption calls (NU) : 136510
cnf(c_0_30, plain, (iof(esk9_1(X1),X1)|firstordertype(X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
# Rec. Clause-clause subsumption calls : 112397
cnf(c_0_31, negated_conjecture, (~individual(esk2_1(esk22_0))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
# Non-unit clause-clause subsumptions  : 520
cnf(c_0_32, plain, (individual(esk2_1(X1))|~firstordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_20]), c_0_27])).
cnf(c_0_33, plain, (X1=X2|properSpecializes(X1,X2)|~specializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
# Unit Clause-clause subsumption calls : 1293
cnf(c_0_34, negated_conjecture, (specializes(X1,esk23_0)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_12, c_0_29])).
# Rewrite failures with RHS unbound    : 0
cnf(c_0_35, negated_conjecture, (firstordertype(esk22_0)|iof(esk9_1(esk22_0),esk22_0)), inference(spm,[status(thm)],[c_0_30, c_0_22])).
# BW rewrite match attempts            : 14
cnf(c_0_36, negated_conjecture, (~firstordertype(esk22_0)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
# BW rewrite match successes           : 14
cnf(c_0_37, negated_conjecture, (X1=esk23_0|properSpecializes(X1,esk23_0)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_38, negated_conjecture, (iof(esk9_1(esk22_0),esk22_0)), inference(sr,[status(thm)],[c_0_35, c_0_36])).
# Condensation attempts                : 0
cnf(c_0_39, plain, (specializes(X1,X2)|~properSpecializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
# Condensation successes               : 0
cnf(c_0_40, negated_conjecture, (esk9_1(esk22_0)=esk23_0|properSpecializes(esk9_1(esk22_0),esk23_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
# Termbank termtop insertions          : 72737

cnf(c_0_41, negated_conjecture, (X1=esk24_0|properSpecializes(X1,esk24_0)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_33, c_0_17])).
# -------------------------------------------------
cnf(c_0_42, plain, (iof(X1,X3)|~iof(X1,X2)|~specializes(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
# User time                : 0.192 s
cnf(c_0_43, negated_conjecture, (esk9_1(esk22_0)=esk23_0|specializes(esk9_1(esk22_0),esk23_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
# System time              : 0.014 s
cnf(c_0_44, negated_conjecture, (esk9_1(esk22_0)=esk24_0|properSpecializes(esk9_1(esk22_0),esk24_0)), inference(spm,[status(thm)],[c_0_41, c_0_38])).
# Total time               : 0.206 s
cnf(c_0_45, negated_conjecture, (esk9_1(esk22_0)=esk23_0|iof(X1,esk23_0)|~iof(X1,esk9_1(esk22_0))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
# Maximum resident set size: 1604 pages
cnf(c_0_46, negated_conjecture, (type_(esk9_1(esk22_0))), inference(spm,[status(thm)],[c_0_21, c_0_38])).
cnf(c_0_47, negated_conjecture, (esk9_1(esk22_0)=esk24_0|specializes(esk9_1(esk22_0),esk24_0)), inference(spm,[status(thm)],[c_0_39, c_0_44])).
cnf(c_0_48, plain, (type_(X2)|~iof(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_49, negated_conjecture, (esk9_1(esk22_0)=esk23_0|iof(esk2_1(esk9_1(esk22_0)),esk23_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_20]), c_0_46])])).
cnf(c_0_50, negated_conjecture, (esk9_1(esk22_0)=esk24_0|iof(X1,esk24_0)|~iof(X1,esk9_1(esk22_0))), inference(spm,[status(thm)],[c_0_42, c_0_47])).
cnf(c_0_51, negated_conjecture, (esk9_1(esk22_0)=esk23_0|type_(esk23_0)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_52, negated_conjecture, (esk9_1(esk22_0)=esk24_0|iof(esk2_1(esk9_1(esk22_0)),esk24_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_20]), c_0_46])])).
cnf(c_0_53, plain, (iof(esk11_2(X1,X2),X1)|specializes(X1,X2)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_54, negated_conjecture, (type_(esk23_0)), inference(spm,[status(thm)],[c_0_46, c_0_51])).
cnf(c_0_55, negated_conjecture, (esk9_1(esk22_0)=esk24_0|type_(esk24_0)), inference(spm,[status(thm)],[c_0_48, c_0_52])).
cnf(c_0_56, plain, (iof(X1,X3)|~specializes(X1,X2)|~isPowertypeOf(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_57, plain, (specializes(X1,X2)|~iof(esk11_2(X1,X2),X2)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_58, negated_conjecture, (specializes(esk23_0,X1)|iof(esk11_2(esk23_0,X1),esk23_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_59, negated_conjecture, (type_(esk24_0)), inference(spm,[status(thm)],[c_0_46, c_0_55])).
cnf(c_0_60, negated_conjecture, (iof(X1,esk22_0)|~specializes(X1,esk23_0)), inference(spm,[status(thm)],[c_0_56, c_0_29])).
cnf(c_0_61, negated_conjecture, (specializes(esk23_0,esk23_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_54])])).
cnf(c_0_62, negated_conjecture, (specializes(esk24_0,X1)|iof(esk11_2(esk24_0,X1),esk24_0)), inference(spm,[status(thm)],[c_0_53, c_0_59])).
cnf(c_0_63, negated_conjecture, (iof(esk23_0,esk22_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_64, negated_conjecture, (esk24_0!=esk23_0), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_65, negated_conjecture, (iof(X1,esk22_0)|~specializes(X1,esk24_0)), inference(spm,[status(thm)],[c_0_56, c_0_13])).
cnf(c_0_66, negated_conjecture, (specializes(esk24_0,esk24_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_62]), c_0_59])])).
fof(c_0_67, plain, ![X60, X61, X62, X63]:(((~iof(X62,X60)|iof(X62,X61)|X60!=X61|(~type_(X60)|~type_(X61)))&(~iof(X63,X61)|iof(X63,X60)|X60!=X61|(~type_(X60)|~type_(X61))))&((~iof(esk12_2(X60,X61),X60)|~iof(esk12_2(X60,X61),X61)|X60=X61|(~type_(X60)|~type_(X61)))&(iof(esk12_2(X60,X61),X60)|iof(esk12_2(X60,X61),X61)|X60=X61|(~type_(X60)|~type_(X61))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a9_equality])])])])])])).
cnf(c_0_68, negated_conjecture, (properSpecializes(esk23_0,esk24_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_63]), c_0_64])).
cnf(c_0_69, negated_conjecture, (iof(esk24_0,esk22_0)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_70, plain, (iof(esk12_2(X1,X2),X1)|iof(esk12_2(X1,X2),X2)|X1=X2|~type_(X1)|~type_(X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_71, negated_conjecture, (specializes(esk23_0,esk24_0)), inference(spm,[status(thm)],[c_0_39, c_0_68])).
cnf(c_0_72, negated_conjecture, (properSpecializes(esk24_0,esk23_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_69]), c_0_64])).
cnf(c_0_73, plain, (X1=X2|~iof(esk12_2(X1,X2),X1)|~iof(esk12_2(X1,X2),X2)|~type_(X1)|~type_(X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_74, negated_conjecture, (X1=esk24_0|iof(esk12_2(X1,esk24_0),esk24_0)|iof(esk12_2(X1,esk24_0),X1)|~type_(X1)), inference(spm,[status(thm)],[c_0_70, c_0_59])).
cnf(c_0_75, negated_conjecture, (iof(X1,esk24_0)|~iof(X1,esk23_0)), inference(spm,[status(thm)],[c_0_42, c_0_71])).
cnf(c_0_76, negated_conjecture, (specializes(esk24_0,esk23_0)), inference(spm,[status(thm)],[c_0_39, c_0_72])).
cnf(c_0_77, plain, (X1=X2|~iof(esk12_2(X1,X2),X2)|~iof(esk12_2(X1,X2),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_73, c_0_48]), c_0_48])).
cnf(c_0_78, negated_conjecture, (iof(esk12_2(esk23_0,esk24_0),esk24_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_54]), c_0_64]), c_0_75])).
cnf(c_0_79, negated_conjecture, (iof(X1,esk23_0)|~iof(X1,esk24_0)), inference(spm,[status(thm)],[c_0_42, c_0_76])).
cnf(c_0_80, negated_conjecture, (~iof(esk12_2(esk23_0,esk24_0),esk23_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_64])).
cnf(c_0_81, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_78]), c_0_80]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 82
# Proof object clause steps            : 65
# Proof object formula steps           : 17
# Proof object conjectures             : 48
# Proof object clause conjectures      : 45
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 43
# Proof object simplifying inferences  : 20
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 77
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 72
# Processed clauses                    : 2014
# ...of these trivial                  : 3
# ...subsumed                          : 822
# ...remaining for further processing  : 1188
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 35
# Backward-rewritten                   : 2
# Generated clauses                    : 4548
# ...of the previous two non-trivial   : 4237
# Contextual simplify-reflections      : 97
# Paramodulations                      : 4546
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1077
#    Positive orientable unit clauses  : 39
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 1028
# Current number of unprocessed clauses: 2326
# ...number of literals in the above   : 8344
# Current number of archived formulas  : 0
# Current number of archived clauses   : 111
# Clause-clause subsumption calls (NU) : 136510
# Rec. Clause-clause subsumption calls : 112397
# Non-unit clause-clause subsumptions  : 520
# Unit Clause-clause subsumption calls : 1293
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 14
# BW rewrite match successes           : 14
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 72737

# -------------------------------------------------
# User time                : 0.186 s
# System time              : 0.015 s
# Total time               : 0.201 s
# Maximum resident set size: 1608 pages
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_U48A

% END OF SYSTEM OUTPUT
% RESULT: SOT_U48AXO - E---FNT-2.5 says Theorem - CPU = 0.00 WC = 0.21 
% OUTPUT: SOT_U48AXO - E---FNT-2.5 says CNFRefutation - CPU = 0.00 WC = 0.21 

% END OF SYSTEM OUTPUT
% RESULT: SOT_U48AXO - Vampire---SAT-4.5 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_U48AXO - Vampire---SAT-4.5 says None - CPU = 0 WC = 0 

% END OF SYSTEM OUTPUT
% RESULT: SOT_U48AXO - E---FNT-2.5 says Theorem - CPU = 0.00 WC = 0.21 
% OUTPUT: SOT_U48AXO - E---FNT-2.5 says CNFRefutation - CPU = 0.00 WC = 0.21 

% END OF SYSTEM OUTPUT
% RESULT: SOT_U48AXO - Vampire---SAT-4.5 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_U48AXO - Vampire---SAT-4.5 says None - CPU = 0 WC = 0 
% 10 wall clock seconds gone, waiting for system output
% 20 wall clock seconds gone, waiting for system output
% 30 wall clock seconds gone, waiting for system output
% 40 wall clock seconds gone, waiting for system output
% 50 wall clock seconds gone, waiting for system output
% 60 wall clock seconds gone, waiting for system output

% END OF SYSTEM OUTPUT
% RESULT: SOT_U48AXO - Vampire---SAT-4.5 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_U48AXO - Vampire---SAT-4.5 says None - CPU = 0 WC = 0 
% 70 wall clock seconds gone, waiting for system output
% 80 wall clock seconds gone, waiting for system output
% 90 wall clock seconds gone, waiting for system output
% 100 wall clock seconds gone, waiting for system output
% 110 wall clock seconds gone, waiting for system output
% 120 wall clock seconds gone, waiting for system output
% 130 wall clock seconds gone, waiting for system output
% 140 wall clock seconds gone, waiting for system output
% 150 wall clock seconds gone, waiting for system output
% 160 wall clock seconds gone, waiting for system output
% 170 wall clock seconds gone, waiting for system output
% 180 wall clock seconds gone, waiting for system output
% 190 wall clock seconds gone, waiting for system output
% 200 wall clock seconds gone, waiting for system output
% 210 wall clock seconds gone, waiting for system output
% 220 wall clock seconds gone, waiting for system output
% 230 wall clock seconds gone, waiting for system output
% 240 wall clock seconds gone, waiting for system output
% 250 wall clock seconds gone, waiting for system output
% 260 wall clock seconds gone, waiting for system output
% 270 wall clock seconds gone, waiting for system output
% 280 wall clock seconds gone, waiting for system output
% 290 wall clock seconds gone, waiting for system output
% 300 wall clock seconds gone, waiting for system output
% 310 wall clock seconds gone, waiting for system output
% 320 wall clock seconds gone, waiting for system output
% 330 wall clock seconds gone, waiting for system output
% 340 wall clock seconds gone, waiting for system output
% 350 wall clock seconds gone, waiting for system output
% 360 wall clock seconds gone, waiting for system output
% 370 wall clock seconds gone, waiting for system output
% 380 wall clock seconds gone, waiting for system output
% 390 wall clock seconds gone, waiting for system output
% 400 wall clock seconds gone, waiting for system output
% 410 wall clock seconds gone, waiting for system output
% 420 wall clock seconds gone, waiting for system output
% 430 wall clock seconds gone, waiting for system output
% 440 wall clock seconds gone, waiting for system output
% 450 wall clock seconds gone, waiting for system output
% 460 wall clock seconds gone, waiting for system output
% 470 wall clock seconds gone, waiting for system output
% 480 wall clock seconds gone, waiting for system output
% 490 wall clock seconds gone, waiting for system output
% 500 wall clock seconds gone, waiting for system output
% 510 wall clock seconds gone, waiting for system output
% 520 wall clock seconds gone, waiting for system output
% 530 wall clock seconds gone, waiting for system output
% 540 wall clock seconds gone, waiting for system output
% 550 wall clock seconds gone, waiting for system output
% 560 wall clock seconds gone, waiting for system output
% 570 wall clock seconds gone, waiting for system output
% 580 wall clock seconds gone, waiting for system output
% 590 wall clock seconds gone, waiting for system output
% 600 wall clock seconds gone, waiting for system output
% 610 wall clock seconds gone, waiting for system output
% 620 wall clock seconds gone, waiting for system output
% 630 wall clock seconds gone, waiting for system output
% 640 wall clock seconds gone, waiting for system output
% 650 wall clock seconds gone, waiting for system output
% 660 wall clock seconds gone, waiting for system output
% 670 wall clock seconds gone, waiting for system output
% 680 wall clock seconds gone, waiting for system output
% 690 wall clock seconds gone, waiting for system output
% 700 wall clock seconds gone, waiting for system output
% 710 wall clock seconds gone, waiting for system output
% 720 wall clock seconds gone, waiting for system output
% 730 wall clock seconds gone, waiting for system output
% 740 wall clock seconds gone, waiting for system output
% 750 wall clock seconds gone, waiting for system output
% 760 wall clock seconds gone, waiting for system output
% 770 wall clock seconds gone, waiting for system output
% 780 wall clock seconds gone, waiting for system output
% 790 wall clock seconds gone, waiting for system output
% 800 wall clock seconds gone, waiting for system output
% 810 wall clock seconds gone, waiting for system output
% 820 wall clock seconds gone, waiting for system output
% 830 wall clock seconds gone, waiting for system output
% 840 wall clock seconds gone, waiting for system output
% 850 wall clock seconds gone, waiting for system output
% 860 wall clock seconds gone, waiting for system output
% 870 wall clock seconds gone, waiting for system output
% 880 wall clock seconds gone, waiting for system output
% 890 wall clock seconds gone, waiting for system output
% 900 wall clock seconds gone, waiting for system output
% 910 wall clock seconds gone, waiting for system output
% 920 wall clock seconds gone, waiting for system output
% 930 wall clock seconds gone, waiting for system output
% 940 wall clock seconds gone, waiting for system output
% 950 wall clock seconds gone, waiting for system output
% 960 wall clock seconds gone, waiting for system output
% 970 wall clock seconds gone, waiting for system output
% 980 wall clock seconds gone, waiting for system output
% 990 wall clock seconds gone, waiting for system output
% 1000 wall clock seconds gone, waiting for system output
% 1010 wall clock seconds gone, waiting for system output
% 1020 wall clock seconds gone, waiting for system output
% 1030 wall clock seconds gone, waiting for system output
% 1040 wall clock seconds gone, waiting for system output
% 1050 wall clock seconds gone, waiting for system output
% 1060 wall clock seconds gone, waiting for system output
% 1070 wall clock seconds gone, waiting for system output
% 1080 wall clock seconds gone, waiting for system output
% 1090 wall clock seconds gone, waiting for system output
% 1100 wall clock seconds gone, waiting for system output
% 1110 wall clock seconds gone, waiting for system output
% 1120 wall clock seconds gone, waiting for system output
% 1130 wall clock seconds gone, waiting for system output
% 1140 wall clock seconds gone, waiting for system output
% 1150 wall clock seconds gone, waiting for system output
% 1160 wall clock seconds gone, waiting for system output
% 1170 wall clock seconds gone, waiting for system output
% 1180 wall clock seconds gone, waiting for system output
% 1190 wall clock seconds gone, waiting for system output
% 1200 wall clock seconds gone, waiting for system output
% 1210 wall clock seconds gone, waiting for system output
% 1220 wall clock seconds gone, waiting for system output
% 1230 wall clock seconds gone, waiting for system output
% 1240 wall clock seconds gone, waiting for system output
% 1250 wall clock seconds gone, waiting for system output
% 1260 wall clock seconds gone, waiting for system output
% 1270 wall clock seconds gone, waiting for system output
% 1280 wall clock seconds gone, waiting for system output
% 1290 wall clock seconds gone, waiting for system output
% 1300 wall clock seconds gone, waiting for system output
% 1310 wall clock seconds gone, waiting for system output
% 1320 wall clock seconds gone, waiting for system output
% 1330 wall clock seconds gone, waiting for system output
% 1340 wall clock seconds gone, waiting for system output
% 1350 wall clock seconds gone, waiting for system output
% 1360 wall clock seconds gone, waiting for system output
% 1370 wall clock seconds gone, waiting for system output
% 1380 wall clock seconds gone, waiting for system output
