<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>System Results</title>
<meta name="ROBOTS" content="INDEX,NOFOLLOW" />
<script type="text/javascript">//<![CDATA[

  var tstp_dump;
  function openSoTSTP (dump) {   
    var tstp_url = 'http://www.tptp.org/cgi-bin/SystemOnTSTP';
    var tstp_browser = window.open(tstp_url, '_blank');
    tstp_dump = dump;
  }
  function getTSTPDump () {
    return tstp_dump;
  }

//]]></script>
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>

% SZS start RequiredInformation<BR>
% Congratulations - you have become a registered power user of SystemOnTPTP, at IP address 191.6.38.118.<BR>
% Please consider donating to the TPTP project - see www.tptp.org for details.<BR>
% When you donate this message will disappear.<BR>
% If you do not donate a random delay might be added to your processing time.<BR>
% SZS end RequiredInformation<BR>
<PRE>
Getting system recommendations for generic problem /tmp/WCAvVxDflM/SOT_fqme_H
This is a FOF_THM_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.95
Vampire---TF1-4.5    recommended  0.95
Vampire---4.5        recommended  0.95
Vampire---SAT-4.5    recommended  0.95
ET---2.0             recommended  0.80
E---2.5              recommended  0.78
E---FNT-2.5          recommended  0.78
CSE_E---1.0          recommended  0.76
E---FNT-2.5          recommended  0.74
Isabelle---2TH0      recommended  0.69
Isabelle---2TF0      recommended  0.69
Isabelle---2FOF      recommended  0.69
Isabelle---2016      recommended  0.69
Isabelle-HOT---2016  recommended  0.68
CVC4---1.7           recommended  0.68
CVC4---SAT-1.7       recommended  0.68
Leo-III---1.4        recommended  0.61
Vampire---THF-4.5    recommended  0.59
Vampire---TF1-4.5    recommended  0.59
Vampire---SAT-4.5    recommended  0.59
iProver---3.0        recommended  0.58
iProver---SAT-3.0    recommended  0.58
Fampire---1.3        recommended  0.57
SInE---0.4           recommended  0.53
iProver---SAT-3.0    recommended  0.52
ePrincess---1.0      recommended  0.50
Princess---170717    recommended  0.48
SPASS---3.9          recommended  0.48
Z3---4.4.1           recommended  0.46
Zipperpin---1.5      recommended  0.40
CVC4---SAT-1.7       recommended  0.40
Prover9---1109a      recommended  0.39
Equinox---6.0.1a     recommended  0.37
LEO-II---1.7.0       recommended  0.35
nanoCoP---1.1        recommended  0.34
Muscadet---4.5       recommended  0.25
Twee---2.2           recommended  0.23
Metis---2.4          recommended  0.23
Geo-III---2018C      recommended  0.18
iProverMo---2.5-0.1  recommended  0.14
Beagle---0.9.47      subsumed     0.41
leanCoP---2.2        subsumed     0.36
E-Darwin---1.5       subsumed     0.31
CSE---1.0            subsumed     0.28
Darwin---1.4.5       subsumed     0.28
Bliksem---1.12       subsumed     0.21
Zenon---0.7.1        subsumed     0.16
ZenonModulo---0.4.2  subsumed     0.15
SOS---2.0            subsumed     0.11
Otter---3.3          subsumed     0.04
This is a FOF_CSA_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.98
Vampire---TF1-4.5    recommended  0.98
Vampire---SAT-4.5    recommended  0.98
Paradox---4.0        recommended  0.60
Nitpick---2016       subsumed     0.55
DarwinFM---1.4.5     subsumed     0.43
CVC4---SAT-1.7       subsumed     0.40
Vampire---THF-4.5    subsumed     0.34
Vampire---TF1-4.5    subsumed     0.34
Vampire---4.5        subsumed     0.34
Vampire---SAT-4.5    subsumed     0.34
E---2.5              subsumed     0.34
E---FNT-2.5          subsumed     0.34
E---FNT-2.5          subsumed     0.34
iProver---SAT-3.0    subsumed     0.32
Metis---2.4          subsumed     0.32
SPASS---3.9          subsumed     0.28
Beagle---0.9.47      subsumed     0.26
Geo-III---2018C      subsumed     0.25
Z3---4.4.1           subsumed     0.25
Princess---170717    subsumed     0.25
Equinox---6.0.1a     subsumed     0.19
iProver---3.0        subsumed     0.11
iProver---SAT-3.0    subsumed     0.11
Darwin---1.4.5       subsumed     0.06
E-Darwin---1.5       subsumed     0.04
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use ET,2.0
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use CSE_E,1.0
CSE_E---1.0 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use Isabelle,2TH0
Isabelle---2016 can be used
Check if we can use Isabelle,2TF0
Isabelle---2016 can be used
Check if we can use Isabelle,2FOF
Isabelle---2016 can be used
Check if we can use Isabelle,2016
Isabelle---2016 can be used
Check if we can use Isabelle-HOT,2016
Isabelle-HOT---2016 can be used
Check if we can use CVC4,1.7
CVC4---SAT-1.7 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Leo-III,1.4
Leo-III---1.4 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Fampire,1.3
Check if we can use SInE,0.4
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use ePrincess,1.0
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Zipperpin,1.5
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Prover9,1109a
Prover9---1109a can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use LEO-II,1.7.0
LEO-II---1.7.0 can be used
Check if we can use nanoCoP,1.1
nanoCoP---1.1 can be used
Check if we can use Muscadet,4.5
Muscadet---4.5 can be used
Check if we can use Twee,2.2
Twee---2.2 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use iProverMo,2.5-0.1
iProverMo---2.5-0.1 can be used
Check if we can use Beagle,0.9.47
Check if we can use leanCoP,2.2
leanCoP---2.2 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Check if we can use CSE,1.0
CSE---1.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use Bliksem,1.12
Bliksem---1.12 can be used
Check if we can use Zenon,0.7.1
Zenon---0.7.1 can be used
Check if we can use ZenonModulo,0.4.2
Check if we can use SOS,2.0
SOS---2.0 can be used
Check if we can use Otter,3.3
Otter---3.3 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use Paradox,4.0
Paradox---4.0 can be used
Check if we can use Nitpick,2016
Nitpick---2016 can be used
Check if we can use DarwinFM,1.4.5
DarwinFM---1.4.5 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Beagle,0.9.47
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
E---FNT-2.5 will be used
E---FNT-2.5 will be used
Vampire---SAT-4.5 is given 499 seconds in -cE mode
Vampire---SAT-4.5 is given 250 seconds in -cE mode
Vampire---SAT-4.5 is given 125 seconds in -cE mode
Vampire---SAT-4.5 is given 62 seconds in -cE mode
E---FNT-2.5 is given 31 seconds in -cE mode
E---FNT-2.5 is given 32 seconds in -cE mode
Vampire---SAT-4.5 has started (PID 97302)
Vampire---SAT-4.5 has started (PID 97304)
Vampire---SAT-4.5 has started (PID 97305)
Vampire---SAT-4.5 has started (PID 97307)
E---FNT-2.5 has started (PID 97309)
E---FNT-2.5 has started (PID 97311)
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  125  /tmp/WCAvVxDflM/SOT_fqme_H
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  31  /tmp/WCAvVxDflM/SOT_fqme_H
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  499  /tmp/WCAvVxDflM/SOT_fqme_H
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  32  /tmp/WCAvVxDflM/SOT_fqme_H
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  250  /tmp/WCAvVxDflM/SOT_fqme_H
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  62  /tmp/WCAvVxDflM/SOT_fqme_H
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 15 seconds
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 15 seconds
# AutoSched0-Mode selected heuristic G_E___208_C02CMA_F1_SE_CS_SP_PS_S5PRR_RG_S04AN
# and selection function SelectComplexExceptUniqMaxHorn.
#
# Preprocessing time       : 0.015 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_basetypeUnique, conjecture, ![X11, X12]:(isPowertypeOf(X11,X12)=>~(?[X13]:(~(X13=X11)&isPowertypeOf(X13,X12)))), file('/tmp/WCAvVxDflM/SOT_fqme_H', t1_basetypeUnique)).
fof(a10_isPowertypeOfDef, axiom, ![X7, X8]:(isPowertypeOf(X7,X8)<=>(type_(X7)&![X9]:(iof(X9,X7)<=>specializes(X9,X8)))), file('/tmp/WCAvVxDflM/SOT_fqme_H', a10_isPowertypeOfDef)).
fof(a9_equality, axiom, ![X7, X8]:((type_(X7)&type_(X8))=>(X7=X8<=>![X1]:(iof(X1,X7)<=>iof(X1,X8)))), file('/tmp/WCAvVxDflM/SOT_fqme_H', a9_equality)).
fof(a3, axiom, ![X1]:(type_(X1)<=>?[X2]:iof(X2,X1)), file('/tmp/WCAvVxDflM/SOT_fqme_H', a3)).
fof(c_0_4, negated_conjecture, ~(![X11, X12]:(isPowertypeOf(X11,X12)=>~(?[X13]:(~(X13=X11)&isPowertypeOf(X13,X12))))), inference(assume_negation,[status(cth)],[t1_basetypeUnique])).
fof(c_0_5, plain, ![X65, X66, X67, X68, X69, X70]:(((type_(X65)|~isPowertypeOf(X65,X66))&((~iof(X67,X65)|specializes(X67,X66)|~isPowertypeOf(X65,X66))&(~specializes(X68,X66)|iof(X68,X65)|~isPowertypeOf(X65,X66))))&((~iof(esk13_2(X69,X70),X69)|~specializes(esk13_2(X69,X70),X70)|~type_(X69)|isPowertypeOf(X69,X70))&(iof(esk13_2(X69,X70),X69)|specializes(esk13_2(X69,X70),X70)|~type_(X69)|isPowertypeOf(X69,X70)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a10_isPowertypeOfDef])])])])])])).
fof(c_0_6, negated_conjecture, (isPowertypeOf(esk22_0,esk23_0)&(esk24_0!=esk22_0&isPowertypeOf(esk24_0,esk23_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_4])])])])).
fof(c_0_7, plain, ![X60, X61, X62, X63]:(((~iof(X62,X60)|iof(X62,X61)|X60!=X61|(~type_(X60)|~type_(X61)))&(~iof(X63,X61)|iof(X63,X60)|X60!=X61|(~type_(X60)|~type_(X61))))&((~iof(esk12_2(X60,X61),X60)|~iof(esk12_2(X60,X61),X61)|X60=X61|(~type_(X60)|~type_(X61)))&(iof(esk12_2(X60,X61),X60)|iof(esk12_2(X60,X61),X61)|X60=X61|(~type_(X60)|~type_(X61))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a9_equality])])])])])])).
cnf(c_0_8, plain, (type_(X1)|~isPowertypeOf(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, negated_conjecture, (isPowertypeOf(esk24_0,esk23_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (iof(X1,X3)|~specializes(X1,X2)|~isPowertypeOf(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11, plain, (specializes(X1,X3)|~iof(X1,X2)|~isPowertypeOf(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_12, negated_conjecture, (isPowertypeOf(esk22_0,esk23_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_13, plain, ![X19, X21, X22]:((~type_(X19)|iof(esk2_1(X19),X19))&(~iof(X22,X21)|type_(X21))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a3])])])])])).
cnf(c_0_14, plain, (iof(esk12_2(X1,X2),X1)|iof(esk12_2(X1,X2),X2)|X1=X2|~type_(X1)|~type_(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15, negated_conjecture, (type_(esk24_0)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_16, negated_conjecture, (iof(X1,esk24_0)|~specializes(X1,esk23_0)), inference(spm,[status(thm)],[c_0_10, c_0_9])).
cnf(c_0_17, negated_conjecture, (specializes(X1,esk23_0)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_18, plain, (X1=X2|~iof(esk12_2(X1,X2),X1)|~iof(esk12_2(X1,X2),X2)|~type_(X1)|~type_(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_19, plain, (type_(X2)|~iof(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, negated_conjecture, (X1=esk24_0|iof(esk12_2(X1,esk24_0),esk24_0)|iof(esk12_2(X1,esk24_0),X1)|~type_(X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_21, negated_conjecture, (type_(esk22_0)), inference(spm,[status(thm)],[c_0_8, c_0_12])).
cnf(c_0_22, negated_conjecture, (esk24_0!=esk22_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_23, negated_conjecture, (iof(X1,esk24_0)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_24, negated_conjecture, (iof(X1,esk22_0)|~specializes(X1,esk23_0)), inference(spm,[status(thm)],[c_0_10, c_0_12])).
cnf(c_0_25, negated_conjecture, (specializes(X1,esk23_0)|~iof(X1,esk24_0)), inference(spm,[status(thm)],[c_0_11, c_0_9])).
cnf(c_0_26, plain, (X1=X2|~iof(esk12_2(X1,X2),X2)|~iof(esk12_2(X1,X2),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_18, c_0_19]), c_0_19])).
cnf(c_0_27, negated_conjecture, (iof(esk12_2(esk22_0,esk24_0),esk24_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23])).
cnf(c_0_28, negated_conjecture, (iof(X1,esk22_0)|~iof(X1,esk24_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29, negated_conjecture, (~iof(esk12_2(esk22_0,esk24_0),esk22_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_22])).
cnf(c_0_30, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_27]), c_0_29]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 31
# Proof object clause steps            : 22
# Proof object formula steps           : 9
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 6
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 77
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 72
# Processed clauses                    : 2201
# ...of these trivial                  : 2
# ...subsumed                          : 860
# ...remaining for further processing  : 1338
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 47
# Backward-rewritten                   : 1
# Generated clauses                    : 6594
# ...of the previous two non-trivial   : 6168
# Contextual simplify-reflections      : 120
# Paramodulations                      : 6591
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1215
#    Positive orientable unit clauses  : 58
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 15
#    Non-unit-clauses                  : 1142
# Current number of unprocessed clauses: 4081
# ...number of literals in the above   : 14164
# Current number of archived formulas  : 0
# Current number of archived clauses   : 123
# Clause-clause subsumption calls (NU) : 135115
# Rec. Clause-clause subsumption calls : 107702
# Non-unit clause-clause subsumptions  : 559
# Unit Clause-clause subsumption calls : 2432
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 34
# BW rewrite match successes           : 24
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 113013

# -------------------------------------------------
# User time                : 0.148 s
# System time              : 0.008 s
# Total time               : 0.156 s
# Maximum resident set size: 1608 pages
% START OF SYSTEM OUTPUT
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_fqme
% START OF SYSTEM OUTPUT
# AutoSched0-Mode selected heuristic G_E___208_C02CMA_F1_SE_CS_SP_PS_S5PRR_RG_S04AN
# and selection function SelectComplexExceptUniqMaxHorn.
#
# Preprocessing time       : 0.015 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_basetypeUnique, conjecture, ![X11, X12]:(isPowertypeOf(X11,X12)=>~(?[X13]:(~(X13=X11)&isPowertypeOf(X13,X12)))), file('/tmp/WCAvVxDflM/SOT_fqme_H', t1_basetypeUnique)).
fof(a10_isPowertypeOfDef, axiom, ![X7, X8]:(isPowertypeOf(X7,X8)<=>(type_(X7)&![X9]:(iof(X9,X7)<=>specializes(X9,X8)))), file('/tmp/WCAvVxDflM/SOT_fqme_H', a10_isPowertypeOfDef)).
fof(a9_equality, axiom, ![X7, X8]:((type_(X7)&type_(X8))=>(X7=X8<=>![X1]:(iof(X1,X7)<=>iof(X1,X8)))), file('/tmp/WCAvVxDflM/SOT_fqme_H', a9_equality)).
fof(a3, axiom, ![X1]:(type_(X1)<=>?[X2]:iof(X2,X1)), file('/tmp/WCAvVxDflM/SOT_fqme_H', a3)).
fof(c_0_4, negated_conjecture, ~(![X11, X12]:(isPowertypeOf(X11,X12)=>~(?[X13]:(~(X13=X11)&isPowertypeOf(X13,X12))))), inference(assume_negation,[status(cth)],[t1_basetypeUnique])).
fof(c_0_5, plain, ![X65, X66, X67, X68, X69, X70]:(((type_(X65)|~isPowertypeOf(X65,X66))&((~iof(X67,X65)|specializes(X67,X66)|~isPowertypeOf(X65,X66))&(~specializes(X68,X66)|iof(X68,X65)|~isPowertypeOf(X65,X66))))&((~iof(esk13_2(X69,X70),X69)|~specializes(esk13_2(X69,X70),X70)|~type_(X69)|isPowertypeOf(X69,X70))&(iof(esk13_2(X69,X70),X69)|specializes(esk13_2(X69,X70),X70)|~type_(X69)|isPowertypeOf(X69,X70)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a10_isPowertypeOfDef])])])])])])).
fof(c_0_6, negated_conjecture, (isPowertypeOf(esk22_0,esk23_0)&(esk24_0!=esk22_0&isPowertypeOf(esk24_0,esk23_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_4])])])])).
fof(c_0_7, plain, ![X60, X61, X62, X63]:(((~iof(X62,X60)|iof(X62,X61)|X60!=X61|(~type_(X60)|~type_(X61)))&(~iof(X63,X61)|iof(X63,X60)|X60!=X61|(~type_(X60)|~type_(X61))))&((~iof(esk12_2(X60,X61),X60)|~iof(esk12_2(X60,X61),X61)|X60=X61|(~type_(X60)|~type_(X61)))&(iof(esk12_2(X60,X61),X60)|iof(esk12_2(X60,X61),X61)|X60=X61|(~type_(X60)|~type_(X61))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a9_equality])])])])])])).
cnf(c_0_8, plain, (type_(X1)|~isPowertypeOf(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, negated_conjecture, (isPowertypeOf(esk24_0,esk23_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (iof(X1,X3)|~specializes(X1,X2)|~isPowertypeOf(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11, plain, (specializes(X1,X3)|~iof(X1,X2)|~isPowertypeOf(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_12, negated_conjecture, (isPowertypeOf(esk22_0,esk23_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_13, plain, ![X19, X21, X22]:((~type_(X19)|iof(esk2_1(X19),X19))&(~iof(X22,X21)|type_(X21))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a3])])])])])).
cnf(c_0_14, plain, (iof(esk12_2(X1,X2),X1)|iof(esk12_2(X1,X2),X2)|X1=X2|~type_(X1)|~type_(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15, negated_conjecture, (type_(esk24_0)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_16, negated_conjecture, (iof(X1,esk24_0)|~specializes(X1,esk23_0)), inference(spm,[status(thm)],[c_0_10, c_0_9])).
cnf(c_0_17, negated_conjecture, (specializes(X1,esk23_0)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_18, plain, (X1=X2|~iof(esk12_2(X1,X2),X1)|~iof(esk12_2(X1,X2),X2)|~type_(X1)|~type_(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_19, plain, (type_(X2)|~iof(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, negated_conjecture, (X1=esk24_0|iof(esk12_2(X1,esk24_0),esk24_0)|iof(esk12_2(X1,esk24_0),X1)|~type_(X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_21, negated_conjecture, (type_(esk22_0)), inference(spm,[status(thm)],[c_0_8, c_0_12])).
cnf(c_0_22, negated_conjecture, (esk24_0!=esk22_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_23, negated_conjecture, (iof(X1,esk24_0)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_24, negated_conjecture, (iof(X1,esk22_0)|~specializes(X1,esk23_0)), inference(spm,[status(thm)],[c_0_10, c_0_12])).
cnf(c_0_25, negated_conjecture, (specializes(X1,esk23_0)|~iof(X1,esk24_0)), inference(spm,[status(thm)],[c_0_11, c_0_9])).
cnf(c_0_26, plain, (X1=X2|~iof(esk12_2(X1,X2),X2)|~iof(esk12_2(X1,X2),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_18, c_0_19]), c_0_19])).
cnf(c_0_27, negated_conjecture, (iof(esk12_2(esk22_0,esk24_0),esk24_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23])).
cnf(c_0_28, negated_conjecture, (iof(X1,esk22_0)|~iof(X1,esk24_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29, negated_conjecture, (~iof(esk12_2(esk22_0,esk24_0),esk22_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_22])).
cnf(c_0_30, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_27]), c_0_29]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 31
# Proof object clause steps            : 22
# Proof object formula steps           : 9
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 6
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 77
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 72
# Processed clauses                    : 2201
# ...of these trivial                  : 2
# ...subsumed                          : 860
# ...remaining for further processing  : 1338
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 47
# Backward-rewritten                   : 1
# Generated clauses                    : 6594
# ...of the previous two non-trivial   : 6168
# Contextual simplify-reflections      : 120
# Paramodulations                      : 6591
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1215
#    Positive orientable unit clauses  : 58
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 15
#    Non-unit-clauses                  : 1142
# Current number of unprocessed clauses: 4081
# ...number of literals in the above   : 14164
# Current number of archived formulas  : 0
# Current number of archived clauses   : 123
# Clause-clause subsumption calls (NU) : 135115
# Rec. Clause-clause subsumption calls : 107702
# Non-unit clause-clause subsumptions  : 559
# Unit Clause-clause subsumption calls : 2432
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 34
# BW rewrite match successes           : 24
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 113013

# -------------------------------------------------
# User time                : 0.153 s
# System time              : 0.011 s
# Total time               : 0.164 s
# Maximum resident set size: 1608 pages
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_fqme
% START OF SYSTEM OUTPUT
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_fqme
% SZS output start Proof for SOT_fqme
fof(f2433,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f102,f106,f110,f142,f156,f162,f213,f226,f426,f628,f995,f1673,f1679,f1692,f1713,f1721,f1725,f2422,f2429,f2432])).
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_fqme
fof(f2432,plain,(
  spl12_244 | ~spl12_3 | ~spl12_299),
  inference(avatar_split_clause,[],[f2424,f2420,f108,f1671])).
fof(f1671,plain,(
  spl12_244 <=> iof(sK8(sK2,sK0),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_244])])).
fof(f108,plain,(
  spl12_3 <=> isPowertypeOf(sK0,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_3])])).
fof(f2420,plain,(
  spl12_299 <=> ! [X0] : (~isPowertypeOf(X0,sK1) | iof(sK8(sK2,sK0),X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_299])])).
fof(f2424,plain,(
  iof(sK8(sK2,sK0),sK0) | (~spl12_3 | ~spl12_299)),
  inference(resolution,[],[f2421,f109])).
fof(f109,plain,(
  isPowertypeOf(sK0,sK1) | ~spl12_3),
  inference(avatar_component_clause,[],[f108])).
fof(f2421,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK1) | iof(sK8(sK2,sK0),X0)) ) | ~spl12_299),
  inference(avatar_component_clause,[],[f2420])).
fof(f2429,plain,(
  spl12_243 | ~spl12_2 | ~spl12_299),
  inference(avatar_split_clause,[],[f2423,f2420,f104,f1668])).
fof(f1668,plain,(
  spl12_243 <=> iof(sK8(sK2,sK0),sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_243])])).
fof(f104,plain,(
  spl12_2 <=> isPowertypeOf(sK2,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_2])])).
fof(f2423,plain,(
  iof(sK8(sK2,sK0),sK2) | (~spl12_2 | ~spl12_299)),
  inference(resolution,[],[f2421,f105])).
fof(f105,plain,(
  isPowertypeOf(sK2,sK1) | ~spl12_2),
  inference(avatar_component_clause,[],[f104])).
fof(f2422,plain,(
  spl12_299 | ~spl12_25 | ~spl12_251),
  inference(avatar_split_clause,[],[f1736,f1719,f211,f2420])).
fof(f211,plain,(
  spl12_25 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | ~specializes(X2,X1) | iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_25])])).
fof(f1719,plain,(
  spl12_251 <=> specializes(sK8(sK2,sK0),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_251])])).
fof(f1736,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK1) | iof(sK8(sK2,sK0),X0)) ) | (~spl12_25 | ~spl12_251)),
  inference(resolution,[],[f1720,f212])).
fof(f212,plain,(
  ( ! [X2,X0,X1] : (~specializes(X2,X1) | ~isPowertypeOf(X0,X1) | iof(X2,X0)) ) | ~spl12_25),
  inference(avatar_component_clause,[],[f211])).
fof(f1720,plain,(
  specializes(sK8(sK2,sK0),sK1) | ~spl12_251),
  inference(avatar_component_clause,[],[f1719])).
fof(f1725,plain,(
  spl12_251 | ~spl12_2 | ~spl12_250),
  inference(avatar_split_clause,[],[f1724,f1711,f104,f1719])).
fof(f1711,plain,(
  spl12_250 <=> ! [X0] : (specializes(sK8(sK2,sK0),X0) | ~isPowertypeOf(sK2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_250])])).
fof(f1724,plain,(
  specializes(sK8(sK2,sK0),sK1) | (~spl12_2 | ~spl12_250)),
  inference(resolution,[],[f1712,f105])).
fof(f1712,plain,(
  ( ! [X0] : (~isPowertypeOf(sK2,X0) | specializes(sK8(sK2,sK0),X0)) ) | ~spl12_250),
  inference(avatar_component_clause,[],[f1711])).
fof(f1721,plain,(
  spl12_251 | ~spl12_3 | ~spl12_247),
  inference(avatar_split_clause,[],[f1717,f1690,f108,f1719])).
fof(f1690,plain,(
  spl12_247 <=> ! [X0] : (specializes(sK8(sK2,sK0),X0) | ~isPowertypeOf(sK0,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_247])])).
fof(f1717,plain,(
  specializes(sK8(sK2,sK0),sK1) | (~spl12_3 | ~spl12_247)),
  inference(resolution,[],[f1691,f109])).
fof(f1691,plain,(
  ( ! [X0] : (~isPowertypeOf(sK0,X0) | specializes(sK8(sK2,sK0),X0)) ) | ~spl12_247),
  inference(avatar_component_clause,[],[f1690])).
fof(f1713,plain,(
  spl12_250 | ~spl12_28 | ~spl12_243),
  inference(avatar_split_clause,[],[f1693,f1668,f224,f1711])).
fof(f224,plain,(
  spl12_28 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_28])])).
fof(f1693,plain,(
  ( ! [X0] : (specializes(sK8(sK2,sK0),X0) | ~isPowertypeOf(sK2,X0)) ) | (~spl12_28 | ~spl12_243)),
  inference(resolution,[],[f1669,f225])).
fof(f225,plain,(
  ( ! [X2,X0,X1] : (~iof(X2,X0) | specializes(X2,X1) | ~isPowertypeOf(X0,X1)) ) | ~spl12_28),
  inference(avatar_component_clause,[],[f224])).
fof(f1669,plain,(
  iof(sK8(sK2,sK0),sK2) | ~spl12_243),
  inference(avatar_component_clause,[],[f1668])).
fof(f1692,plain,(
  spl12_247 | ~spl12_28 | ~spl12_244),
  inference(avatar_split_clause,[],[f1675,f1671,f224,f1690])).
fof(f1675,plain,(
  ( ! [X0] : (specializes(sK8(sK2,sK0),X0) | ~isPowertypeOf(sK0,X0)) ) | (~spl12_28 | ~spl12_244)),
  inference(resolution,[],[f1672,f225])).
fof(f1672,plain,(
  iof(sK8(sK2,sK0),sK0) | ~spl12_244),
  inference(avatar_component_clause,[],[f1671])).
fof(f1679,plain,(
  ~spl12_243 | spl12_1 | ~spl12_65 | ~spl12_244),
  inference(avatar_split_clause,[],[f1677,f1671,f424,f100,f1668])).
fof(f100,plain,(
  spl12_1 <=> sK0 = sK2),
  introduced(avatar_definition,[new_symbols(naming,[spl12_1])])).
fof(f424,plain,(
  spl12_65 <=> ! [X1,X0] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_65])])).
fof(f1677,plain,(
  ~iof(sK8(sK2,sK0),sK2) | (spl12_1 | ~spl12_65 | ~spl12_244)),
  inference(subsumption_resolution,[],[f1674,f101])).
fof(f101,plain,(
  sK0 != sK2 | spl12_1),
  inference(avatar_component_clause,[],[f100])).
fof(f1674,plain,(
  ~iof(sK8(sK2,sK0),sK2) | sK0 = sK2 | (~spl12_65 | ~spl12_244)),
  inference(resolution,[],[f1672,f425])).
fof(f425,plain,(
  ( ! [X0,X1] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) ) | ~spl12_65),
  inference(avatar_component_clause,[],[f424])).
fof(f1673,plain,(
  spl12_243 | spl12_244 | spl12_1 | ~spl12_14 | ~spl12_151),
  inference(avatar_split_clause,[],[f1014,f993,f154,f100,f1671,f1668])).
fof(f154,plain,(
  spl12_14 <=> type_(sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_14])])).
fof(f993,plain,(
  spl12_151 <=> ! [X0] : (~type_(X0) | iof(sK8(X0,sK0),sK0) | iof(sK8(X0,sK0),X0) | sK0 = X0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_151])])).
fof(f1014,plain,(
  iof(sK8(sK2,sK0),sK0) | iof(sK8(sK2,sK0),sK2) | (spl12_1 | ~spl12_14 | ~spl12_151)),
  inference(subsumption_resolution,[],[f1010,f101])).
fof(f1010,plain,(
  iof(sK8(sK2,sK0),sK0) | iof(sK8(sK2,sK0),sK2) | sK0 = sK2 | (~spl12_14 | ~spl12_151)),
  inference(resolution,[],[f994,f155])).
fof(f155,plain,(
  type_(sK2) | ~spl12_14),
  inference(avatar_component_clause,[],[f154])).
fof(f994,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK8(X0,sK0),sK0) | iof(sK8(X0,sK0),X0) | sK0 = X0) ) | ~spl12_151),
  inference(avatar_component_clause,[],[f993])).
fof(f995,plain,(
  spl12_151 | ~spl12_15 | ~spl12_97),
  inference(avatar_split_clause,[],[f652,f626,f160,f993])).
fof(f160,plain,(
  spl12_15 <=> type_(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_15])])).
fof(f626,plain,(
  spl12_97 <=> ! [X1,X0] : (~type_(X0) | ~type_(X1) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_97])])).
fof(f652,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK8(X0,sK0),sK0) | iof(sK8(X0,sK0),X0) | sK0 = X0) ) | (~spl12_15 | ~spl12_97)),
  inference(resolution,[],[f627,f161])).
fof(f161,plain,(
  type_(sK0) | ~spl12_15),
  inference(avatar_component_clause,[],[f160])).
fof(f627,plain,(
  ( ! [X0,X1] : (~type_(X1) | ~type_(X0) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1) ) | ~spl12_97),
  inference(avatar_component_clause,[],[f626])).
fof(f628,plain,(
  spl12_97),
  inference(avatar_split_clause,[],[f73,f626])).
fof(f73,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~type_(X1) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1] : ((X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))) | ~type_(X1) | ~type_(X0))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : ((X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))) | (~type_(X1) | ~type_(X0)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : ((type_(X1) & type_(X0)) => (X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7] : ((type_(X7) & type_(X6)) => (X6 = X7 <=> ! [X0] : (iof(X0,X6) <=> iof(X0,X7))))),
  file('/tmp/WCAvVxDflM/SOT_fqme_H',a9_equality)).
fof(f426,plain,(
  spl12_65),
  inference(avatar_split_clause,[],[f95,f424])).
fof(f95,plain,(
  ( ! [X0,X1] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(subsumption_resolution,[],[f94,f67])).
fof(f67,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/WCAvVxDflM/SOT_fqme_H',a3)).
fof(f94,plain,(
  ( ! [X0,X1] : (~type_(X1) | ~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(subsumption_resolution,[],[f74,f67])).
fof(f74,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~type_(X1) | ~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f46])).
fof(f226,plain,(
  spl12_28),
  inference(avatar_split_clause,[],[f88,f224])).
fof(f88,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : ((! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)) | ~isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) => (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) <=> (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : (isPowertypeOf(X6,X7) <=> (! [X8] : (iof(X8,X6) <=> specializes(X8,X7)) & type_(X6)))),
  file('/tmp/WCAvVxDflM/SOT_fqme_H',a10_isPowertypeOfDef)).
fof(f213,plain,(
  spl12_25),
  inference(avatar_split_clause,[],[f87,f211])).
fof(f87,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | ~specializes(X2,X1) | iof(X2,X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f162,plain,(
  spl12_15 | ~spl12_3 | ~spl12_11),
  inference(avatar_split_clause,[],[f152,f140,f108,f160])).
fof(f140,plain,(
  spl12_11 <=> ! [X1,X0] : (~isPowertypeOf(X0,X1) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_11])])).
fof(f152,plain,(
  type_(sK0) | (~spl12_3 | ~spl12_11)),
  inference(resolution,[],[f141,f109])).
fof(f141,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X0,X1) | type_(X0)) ) | ~spl12_11),
  inference(avatar_component_clause,[],[f140])).
fof(f156,plain,(
  spl12_14 | ~spl12_2 | ~spl12_11),
  inference(avatar_split_clause,[],[f151,f140,f104,f154])).
fof(f151,plain,(
  type_(sK2) | (~spl12_2 | ~spl12_11)),
  inference(resolution,[],[f141,f105])).
fof(f142,plain,(
  spl12_11),
  inference(avatar_split_clause,[],[f89,f140])).
fof(f89,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X0,X1) | type_(X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f110,plain,(
  spl12_3),
  inference(avatar_split_clause,[],[f55,f108])).
fof(f55,plain,(
  isPowertypeOf(sK0,sK1)),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ? [X0,X1] : (? [X2] : (isPowertypeOf(X2,X1) & X0 != X2) & isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1] : (isPowertypeOf(X0,X1) => ~? [X2] : (isPowertypeOf(X2,X1) & X0 != X2))),
  inference(rectify,[],[f22])).
fof(f22,negated_conjecture,(
  ~! [X10,X11] : (isPowertypeOf(X10,X11) => ~? [X12] : (isPowertypeOf(X12,X11) & X10 != X12))),
  inference(negated_conjecture,[],[f21])).
fof(f21,conjecture,(
  ! [X10,X11] : (isPowertypeOf(X10,X11) => ~? [X12] : (isPowertypeOf(X12,X11) & X10 != X12))),
  file('/tmp/WCAvVxDflM/SOT_fqme_H',t1_basetypeUnique)).
fof(f106,plain,(
  spl12_2),
  inference(avatar_split_clause,[],[f54,f104])).
fof(f54,plain,(
  isPowertypeOf(sK2,sK1)),
  inference(cnf_transformation,[],[f38])).
fof(f102,plain,(
  ~spl12_1),
  inference(avatar_split_clause,[],[f53,f100])).
fof(f53,plain,(
  sK0 != sK2),
  inference(cnf_transformation,[],[f38])).
% SZS output end Proof for SOT_fqme
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 7675
% Time elapsed: 0.089 s
% ------------------------------
% ------------------------------
% Success in time 0.135 s
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_fqme
% SZS output start Proof for SOT_fqme
fof(f2433,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f102,f106,f110,f142,f156,f162,f213,f226,f426,f628,f995,f1673,f1679,f1692,f1713,f1721,f1725,f2422,f2429,f2432])).
fof(f2432,plain,(
  spl12_244 | ~spl12_3 | ~spl12_299),
  inference(avatar_split_clause,[],[f2424,f2420,f108,f1671])).
fof(f1671,plain,(
  spl12_244 <=> iof(sK8(sK2,sK0),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_244])])).
fof(f108,plain,(
  spl12_3 <=> isPowertypeOf(sK0,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_3])])).
fof(f2420,plain,(
  spl12_299 <=> ! [X0] : (~isPowertypeOf(X0,sK1) | iof(sK8(sK2,sK0),X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_299])])).
fof(f2424,plain,(
  iof(sK8(sK2,sK0),sK0) | (~spl12_3 | ~spl12_299)),
  inference(resolution,[],[f2421,f109])).
fof(f109,plain,(
  isPowertypeOf(sK0,sK1) | ~spl12_3),
  inference(avatar_component_clause,[],[f108])).
fof(f2421,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK1) | iof(sK8(sK2,sK0),X0)) ) | ~spl12_299),
  inference(avatar_component_clause,[],[f2420])).
fof(f2429,plain,(
  spl12_243 | ~spl12_2 | ~spl12_299),
  inference(avatar_split_clause,[],[f2423,f2420,f104,f1668])).
fof(f1668,plain,(
  spl12_243 <=> iof(sK8(sK2,sK0),sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_243])])).
fof(f104,plain,(
  spl12_2 <=> isPowertypeOf(sK2,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_2])])).
fof(f2423,plain,(
  iof(sK8(sK2,sK0),sK2) | (~spl12_2 | ~spl12_299)),
  inference(resolution,[],[f2421,f105])).
fof(f105,plain,(
  isPowertypeOf(sK2,sK1) | ~spl12_2),
  inference(avatar_component_clause,[],[f104])).
fof(f2422,plain,(
  spl12_299 | ~spl12_25 | ~spl12_251),
  inference(avatar_split_clause,[],[f1736,f1719,f211,f2420])).
fof(f211,plain,(
  spl12_25 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | ~specializes(X2,X1) | iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_25])])).
fof(f1719,plain,(
  spl12_251 <=> specializes(sK8(sK2,sK0),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_251])])).
fof(f1736,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK1) | iof(sK8(sK2,sK0),X0)) ) | (~spl12_25 | ~spl12_251)),
  inference(resolution,[],[f1720,f212])).
fof(f212,plain,(
  ( ! [X2,X0,X1] : (~specializes(X2,X1) | ~isPowertypeOf(X0,X1) | iof(X2,X0)) ) | ~spl12_25),
  inference(avatar_component_clause,[],[f211])).
fof(f1720,plain,(
  specializes(sK8(sK2,sK0),sK1) | ~spl12_251),
  inference(avatar_component_clause,[],[f1719])).
fof(f1725,plain,(
  spl12_251 | ~spl12_2 | ~spl12_250),
  inference(avatar_split_clause,[],[f1724,f1711,f104,f1719])).
fof(f1711,plain,(
  spl12_250 <=> ! [X0] : (specializes(sK8(sK2,sK0),X0) | ~isPowertypeOf(sK2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_250])])).
fof(f1724,plain,(
  specializes(sK8(sK2,sK0),sK1) | (~spl12_2 | ~spl12_250)),
  inference(resolution,[],[f1712,f105])).
fof(f1712,plain,(
  ( ! [X0] : (~isPowertypeOf(sK2,X0) | specializes(sK8(sK2,sK0),X0)) ) | ~spl12_250),
  inference(avatar_component_clause,[],[f1711])).
fof(f1721,plain,(
  spl12_251 | ~spl12_3 | ~spl12_247),
  inference(avatar_split_clause,[],[f1717,f1690,f108,f1719])).
fof(f1690,plain,(
  spl12_247 <=> ! [X0] : (specializes(sK8(sK2,sK0),X0) | ~isPowertypeOf(sK0,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_247])])).
fof(f1717,plain,(
  specializes(sK8(sK2,sK0),sK1) | (~spl12_3 | ~spl12_247)),
  inference(resolution,[],[f1691,f109])).
fof(f1691,plain,(
  ( ! [X0] : (~isPowertypeOf(sK0,X0) | specializes(sK8(sK2,sK0),X0)) ) | ~spl12_247),
  inference(avatar_component_clause,[],[f1690])).
fof(f1713,plain,(
  spl12_250 | ~spl12_28 | ~spl12_243),
  inference(avatar_split_clause,[],[f1693,f1668,f224,f1711])).
fof(f224,plain,(
  spl12_28 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_28])])).
fof(f1693,plain,(
  ( ! [X0] : (specializes(sK8(sK2,sK0),X0) | ~isPowertypeOf(sK2,X0)) ) | (~spl12_28 | ~spl12_243)),
  inference(resolution,[],[f1669,f225])).
fof(f225,plain,(
  ( ! [X2,X0,X1] : (~iof(X2,X0) | specializes(X2,X1) | ~isPowertypeOf(X0,X1)) ) | ~spl12_28),
  inference(avatar_component_clause,[],[f224])).
fof(f1669,plain,(
  iof(sK8(sK2,sK0),sK2) | ~spl12_243),
  inference(avatar_component_clause,[],[f1668])).
fof(f1692,plain,(
  spl12_247 | ~spl12_28 | ~spl12_244),
  inference(avatar_split_clause,[],[f1675,f1671,f224,f1690])).
fof(f1675,plain,(
  ( ! [X0] : (specializes(sK8(sK2,sK0),X0) | ~isPowertypeOf(sK0,X0)) ) | (~spl12_28 | ~spl12_244)),
  inference(resolution,[],[f1672,f225])).
fof(f1672,plain,(
  iof(sK8(sK2,sK0),sK0) | ~spl12_244),
  inference(avatar_component_clause,[],[f1671])).
fof(f1679,plain,(
  ~spl12_243 | spl12_1 | ~spl12_65 | ~spl12_244),
  inference(avatar_split_clause,[],[f1677,f1671,f424,f100,f1668])).
fof(f100,plain,(
  spl12_1 <=> sK0 = sK2),
  introduced(avatar_definition,[new_symbols(naming,[spl12_1])])).
fof(f424,plain,(
  spl12_65 <=> ! [X1,X0] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_65])])).
fof(f1677,plain,(
  ~iof(sK8(sK2,sK0),sK2) | (spl12_1 | ~spl12_65 | ~spl12_244)),
  inference(subsumption_resolution,[],[f1674,f101])).
fof(f101,plain,(
  sK0 != sK2 | spl12_1),
  inference(avatar_component_clause,[],[f100])).
fof(f1674,plain,(
  ~iof(sK8(sK2,sK0),sK2) | sK0 = sK2 | (~spl12_65 | ~spl12_244)),
  inference(resolution,[],[f1672,f425])).
fof(f425,plain,(
  ( ! [X0,X1] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) ) | ~spl12_65),
  inference(avatar_component_clause,[],[f424])).
fof(f1673,plain,(
  spl12_243 | spl12_244 | spl12_1 | ~spl12_14 | ~spl12_151),
  inference(avatar_split_clause,[],[f1014,f993,f154,f100,f1671,f1668])).
fof(f154,plain,(
  spl12_14 <=> type_(sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_14])])).
fof(f993,plain,(
  spl12_151 <=> ! [X0] : (~type_(X0) | iof(sK8(X0,sK0),sK0) | iof(sK8(X0,sK0),X0) | sK0 = X0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_151])])).
fof(f1014,plain,(
  iof(sK8(sK2,sK0),sK0) | iof(sK8(sK2,sK0),sK2) | (spl12_1 | ~spl12_14 | ~spl12_151)),
  inference(subsumption_resolution,[],[f1010,f101])).
fof(f1010,plain,(
  iof(sK8(sK2,sK0),sK0) | iof(sK8(sK2,sK0),sK2) | sK0 = sK2 | (~spl12_14 | ~spl12_151)),
  inference(resolution,[],[f994,f155])).
fof(f155,plain,(
  type_(sK2) | ~spl12_14),
  inference(avatar_component_clause,[],[f154])).
fof(f994,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK8(X0,sK0),sK0) | iof(sK8(X0,sK0),X0) | sK0 = X0) ) | ~spl12_151),
  inference(avatar_component_clause,[],[f993])).
fof(f995,plain,(
  spl12_151 | ~spl12_15 | ~spl12_97),
  inference(avatar_split_clause,[],[f652,f626,f160,f993])).
fof(f160,plain,(
  spl12_15 <=> type_(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_15])])).
fof(f626,plain,(
  spl12_97 <=> ! [X1,X0] : (~type_(X0) | ~type_(X1) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_97])])).
fof(f652,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK8(X0,sK0),sK0) | iof(sK8(X0,sK0),X0) | sK0 = X0) ) | (~spl12_15 | ~spl12_97)),
  inference(resolution,[],[f627,f161])).
fof(f161,plain,(
  type_(sK0) | ~spl12_15),
  inference(avatar_component_clause,[],[f160])).
fof(f627,plain,(
  ( ! [X0,X1] : (~type_(X1) | ~type_(X0) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1) ) | ~spl12_97),
  inference(avatar_component_clause,[],[f626])).
fof(f628,plain,(
  spl12_97),
  inference(avatar_split_clause,[],[f73,f626])).
fof(f73,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~type_(X1) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1] : ((X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))) | ~type_(X1) | ~type_(X0))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : ((X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))) | (~type_(X1) | ~type_(X0)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : ((type_(X1) & type_(X0)) => (X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7] : ((type_(X7) & type_(X6)) => (X6 = X7 <=> ! [X0] : (iof(X0,X6) <=> iof(X0,X7))))),
  file('/tmp/WCAvVxDflM/SOT_fqme_H',a9_equality)).
fof(f426,plain,(
  spl12_65),
  inference(avatar_split_clause,[],[f95,f424])).
fof(f95,plain,(
  ( ! [X0,X1] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(subsumption_resolution,[],[f94,f67])).
fof(f67,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/WCAvVxDflM/SOT_fqme_H',a3)).
fof(f94,plain,(
  ( ! [X0,X1] : (~type_(X1) | ~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(subsumption_resolution,[],[f74,f67])).
fof(f74,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~type_(X1) | ~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f46])).
fof(f226,plain,(
  spl12_28),
  inference(avatar_split_clause,[],[f88,f224])).
fof(f88,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : ((! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)) | ~isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) => (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) <=> (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : (isPowertypeOf(X6,X7) <=> (! [X8] : (iof(X8,X6) <=> specializes(X8,X7)) & type_(X6)))),
  file('/tmp/WCAvVxDflM/SOT_fqme_H',a10_isPowertypeOfDef)).
fof(f213,plain,(
  spl12_25),
  inference(avatar_split_clause,[],[f87,f211])).
fof(f87,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | ~specializes(X2,X1) | iof(X2,X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f162,plain,(
  spl12_15 | ~spl12_3 | ~spl12_11),
  inference(avatar_split_clause,[],[f152,f140,f108,f160])).
fof(f140,plain,(
  spl12_11 <=> ! [X1,X0] : (~isPowertypeOf(X0,X1) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_11])])).
fof(f152,plain,(
  type_(sK0) | (~spl12_3 | ~spl12_11)),
  inference(resolution,[],[f141,f109])).
fof(f141,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X0,X1) | type_(X0)) ) | ~spl12_11),
  inference(avatar_component_clause,[],[f140])).
fof(f156,plain,(
  spl12_14 | ~spl12_2 | ~spl12_11),
  inference(avatar_split_clause,[],[f151,f140,f104,f154])).
fof(f151,plain,(
  type_(sK2) | (~spl12_2 | ~spl12_11)),
  inference(resolution,[],[f141,f105])).
fof(f142,plain,(
  spl12_11),
  inference(avatar_split_clause,[],[f89,f140])).
fof(f89,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X0,X1) | type_(X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f110,plain,(
  spl12_3),
  inference(avatar_split_clause,[],[f55,f108])).
fof(f55,plain,(
  isPowertypeOf(sK0,sK1)),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ? [X0,X1] : (? [X2] : (isPowertypeOf(X2,X1) & X0 != X2) & isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1] : (isPowertypeOf(X0,X1) => ~? [X2] : (isPowertypeOf(X2,X1) & X0 != X2))),
  inference(rectify,[],[f22])).
fof(f22,negated_conjecture,(
  ~! [X10,X11] : (isPowertypeOf(X10,X11) => ~? [X12] : (isPowertypeOf(X12,X11) & X10 != X12))),
  inference(negated_conjecture,[],[f21])).
fof(f21,conjecture,(
  ! [X10,X11] : (isPowertypeOf(X10,X11) => ~? [X12] : (isPowertypeOf(X12,X11) & X10 != X12))),
  file('/tmp/WCAvVxDflM/SOT_fqme_H',t1_basetypeUnique)).
fof(f106,plain,(
  spl12_2),
  inference(avatar_split_clause,[],[f54,f104])).
fof(f54,plain,(
  isPowertypeOf(sK2,sK1)),
  inference(cnf_transformation,[],[f38])).
fof(f102,plain,(
  ~spl12_1),
  inference(avatar_split_clause,[],[f53,f100])).
fof(f53,plain,(
  sK0 != sK2),
  inference(cnf_transformation,[],[f38])).
% SZS output end Proof for SOT_fqme
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 7675
% Time elapsed: 0.084 s
% ------------------------------
% ------------------------------
% Success in time 0.136 s
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_fqme
% SZS output start Proof for SOT_fqme
fof(f2433,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f102,f106,f110,f142,f156,f162,f213,f226,f426,f628,f995,f1673,f1679,f1692,f1713,f1721,f1725,f2422,f2429,f2432])).
fof(f2432,plain,(
  spl12_244 | ~spl12_3 | ~spl12_299),
  inference(avatar_split_clause,[],[f2424,f2420,f108,f1671])).
fof(f1671,plain,(
  spl12_244 <=> iof(sK8(sK2,sK0),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_244])])).
fof(f108,plain,(
  spl12_3 <=> isPowertypeOf(sK0,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_3])])).
fof(f2420,plain,(
  spl12_299 <=> ! [X0] : (~isPowertypeOf(X0,sK1) | iof(sK8(sK2,sK0),X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_299])])).
fof(f2424,plain,(
  iof(sK8(sK2,sK0),sK0) | (~spl12_3 | ~spl12_299)),
  inference(resolution,[],[f2421,f109])).
fof(f109,plain,(
  isPowertypeOf(sK0,sK1) | ~spl12_3),
  inference(avatar_component_clause,[],[f108])).
fof(f2421,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK1) | iof(sK8(sK2,sK0),X0)) ) | ~spl12_299),
  inference(avatar_component_clause,[],[f2420])).
fof(f2429,plain,(
  spl12_243 | ~spl12_2 | ~spl12_299),
  inference(avatar_split_clause,[],[f2423,f2420,f104,f1668])).
fof(f1668,plain,(
  spl12_243 <=> iof(sK8(sK2,sK0),sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_243])])).
fof(f104,plain,(
  spl12_2 <=> isPowertypeOf(sK2,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_2])])).
fof(f2423,plain,(
  iof(sK8(sK2,sK0),sK2) | (~spl12_2 | ~spl12_299)),
  inference(resolution,[],[f2421,f105])).
fof(f105,plain,(
  isPowertypeOf(sK2,sK1) | ~spl12_2),
  inference(avatar_component_clause,[],[f104])).
fof(f2422,plain,(
  spl12_299 | ~spl12_25 | ~spl12_251),
  inference(avatar_split_clause,[],[f1736,f1719,f211,f2420])).
fof(f211,plain,(
  spl12_25 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | ~specializes(X2,X1) | iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_25])])).
fof(f1719,plain,(
  spl12_251 <=> specializes(sK8(sK2,sK0),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_251])])).
fof(f1736,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK1) | iof(sK8(sK2,sK0),X0)) ) | (~spl12_25 | ~spl12_251)),
  inference(resolution,[],[f1720,f212])).
fof(f212,plain,(
  ( ! [X2,X0,X1] : (~specializes(X2,X1) | ~isPowertypeOf(X0,X1) | iof(X2,X0)) ) | ~spl12_25),
  inference(avatar_component_clause,[],[f211])).
fof(f1720,plain,(
  specializes(sK8(sK2,sK0),sK1) | ~spl12_251),
  inference(avatar_component_clause,[],[f1719])).
fof(f1725,plain,(
  spl12_251 | ~spl12_2 | ~spl12_250),
  inference(avatar_split_clause,[],[f1724,f1711,f104,f1719])).
fof(f1711,plain,(
  spl12_250 <=> ! [X0] : (specializes(sK8(sK2,sK0),X0) | ~isPowertypeOf(sK2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_250])])).
fof(f1724,plain,(
  specializes(sK8(sK2,sK0),sK1) | (~spl12_2 | ~spl12_250)),
  inference(resolution,[],[f1712,f105])).
fof(f1712,plain,(
  ( ! [X0] : (~isPowertypeOf(sK2,X0) | specializes(sK8(sK2,sK0),X0)) ) | ~spl12_250),
  inference(avatar_component_clause,[],[f1711])).
fof(f1721,plain,(
  spl12_251 | ~spl12_3 | ~spl12_247),
  inference(avatar_split_clause,[],[f1717,f1690,f108,f1719])).
fof(f1690,plain,(
  spl12_247 <=> ! [X0] : (specializes(sK8(sK2,sK0),X0) | ~isPowertypeOf(sK0,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_247])])).
fof(f1717,plain,(
  specializes(sK8(sK2,sK0),sK1) | (~spl12_3 | ~spl12_247)),
  inference(resolution,[],[f1691,f109])).
fof(f1691,plain,(
  ( ! [X0] : (~isPowertypeOf(sK0,X0) | specializes(sK8(sK2,sK0),X0)) ) | ~spl12_247),
  inference(avatar_component_clause,[],[f1690])).
fof(f1713,plain,(
  spl12_250 | ~spl12_28 | ~spl12_243),
  inference(avatar_split_clause,[],[f1693,f1668,f224,f1711])).
fof(f224,plain,(
  spl12_28 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_28])])).
fof(f1693,plain,(
  ( ! [X0] : (specializes(sK8(sK2,sK0),X0) | ~isPowertypeOf(sK2,X0)) ) | (~spl12_28 | ~spl12_243)),
  inference(resolution,[],[f1669,f225])).
fof(f225,plain,(
  ( ! [X2,X0,X1] : (~iof(X2,X0) | specializes(X2,X1) | ~isPowertypeOf(X0,X1)) ) | ~spl12_28),
  inference(avatar_component_clause,[],[f224])).
fof(f1669,plain,(
  iof(sK8(sK2,sK0),sK2) | ~spl12_243),
  inference(avatar_component_clause,[],[f1668])).
fof(f1692,plain,(
  spl12_247 | ~spl12_28 | ~spl12_244),
  inference(avatar_split_clause,[],[f1675,f1671,f224,f1690])).
fof(f1675,plain,(
  ( ! [X0] : (specializes(sK8(sK2,sK0),X0) | ~isPowertypeOf(sK0,X0)) ) | (~spl12_28 | ~spl12_244)),
  inference(resolution,[],[f1672,f225])).
fof(f1672,plain,(
  iof(sK8(sK2,sK0),sK0) | ~spl12_244),
  inference(avatar_component_clause,[],[f1671])).
fof(f1679,plain,(
  ~spl12_243 | spl12_1 | ~spl12_65 | ~spl12_244),
  inference(avatar_split_clause,[],[f1677,f1671,f424,f100,f1668])).
fof(f100,plain,(
  spl12_1 <=> sK0 = sK2),
  introduced(avatar_definition,[new_symbols(naming,[spl12_1])])).
fof(f424,plain,(
  spl12_65 <=> ! [X1,X0] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_65])])).
fof(f1677,plain,(
  ~iof(sK8(sK2,sK0),sK2) | (spl12_1 | ~spl12_65 | ~spl12_244)),
  inference(subsumption_resolution,[],[f1674,f101])).
fof(f101,plain,(
  sK0 != sK2 | spl12_1),
  inference(avatar_component_clause,[],[f100])).
fof(f1674,plain,(
  ~iof(sK8(sK2,sK0),sK2) | sK0 = sK2 | (~spl12_65 | ~spl12_244)),
  inference(resolution,[],[f1672,f425])).
fof(f425,plain,(
  ( ! [X0,X1] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) ) | ~spl12_65),
  inference(avatar_component_clause,[],[f424])).
fof(f1673,plain,(
  spl12_243 | spl12_244 | spl12_1 | ~spl12_14 | ~spl12_151),
  inference(avatar_split_clause,[],[f1014,f993,f154,f100,f1671,f1668])).
fof(f154,plain,(
  spl12_14 <=> type_(sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_14])])).
fof(f993,plain,(
  spl12_151 <=> ! [X0] : (~type_(X0) | iof(sK8(X0,sK0),sK0) | iof(sK8(X0,sK0),X0) | sK0 = X0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_151])])).
fof(f1014,plain,(
  iof(sK8(sK2,sK0),sK0) | iof(sK8(sK2,sK0),sK2) | (spl12_1 | ~spl12_14 | ~spl12_151)),
  inference(subsumption_resolution,[],[f1010,f101])).
fof(f1010,plain,(
  iof(sK8(sK2,sK0),sK0) | iof(sK8(sK2,sK0),sK2) | sK0 = sK2 | (~spl12_14 | ~spl12_151)),
  inference(resolution,[],[f994,f155])).
fof(f155,plain,(
  type_(sK2) | ~spl12_14),
  inference(avatar_component_clause,[],[f154])).
fof(f994,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK8(X0,sK0),sK0) | iof(sK8(X0,sK0),X0) | sK0 = X0) ) | ~spl12_151),
  inference(avatar_component_clause,[],[f993])).
fof(f995,plain,(
  spl12_151 | ~spl12_15 | ~spl12_97),
  inference(avatar_split_clause,[],[f652,f626,f160,f993])).
fof(f160,plain,(
  spl12_15 <=> type_(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_15])])).
fof(f626,plain,(
  spl12_97 <=> ! [X1,X0] : (~type_(X0) | ~type_(X1) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_97])])).
fof(f652,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK8(X0,sK0),sK0) | iof(sK8(X0,sK0),X0) | sK0 = X0) ) | (~spl12_15 | ~spl12_97)),
  inference(resolution,[],[f627,f161])).
fof(f161,plain,(
  type_(sK0) | ~spl12_15),
  inference(avatar_component_clause,[],[f160])).
fof(f627,plain,(
  ( ! [X0,X1] : (~type_(X1) | ~type_(X0) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1) ) | ~spl12_97),
  inference(avatar_component_clause,[],[f626])).
fof(f628,plain,(
  spl12_97),
  inference(avatar_split_clause,[],[f73,f626])).
fof(f73,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~type_(X1) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1] : ((X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))) | ~type_(X1) | ~type_(X0))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : ((X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))) | (~type_(X1) | ~type_(X0)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : ((type_(X1) & type_(X0)) => (X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7] : ((type_(X7) & type_(X6)) => (X6 = X7 <=> ! [X0] : (iof(X0,X6) <=> iof(X0,X7))))),
  file('/tmp/WCAvVxDflM/SOT_fqme_H',a9_equality)).
fof(f426,plain,(
  spl12_65),
  inference(avatar_split_clause,[],[f95,f424])).
fof(f95,plain,(
  ( ! [X0,X1] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(subsumption_resolution,[],[f94,f67])).
fof(f67,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/WCAvVxDflM/SOT_fqme_H',a3)).
fof(f94,plain,(
  ( ! [X0,X1] : (~type_(X1) | ~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(subsumption_resolution,[],[f74,f67])).
fof(f74,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~type_(X1) | ~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f46])).
fof(f226,plain,(
  spl12_28),
  inference(avatar_split_clause,[],[f88,f224])).
fof(f88,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : ((! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)) | ~isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) => (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) <=> (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : (isPowertypeOf(X6,X7) <=> (! [X8] : (iof(X8,X6) <=> specializes(X8,X7)) & type_(X6)))),
  file('/tmp/WCAvVxDflM/SOT_fqme_H',a10_isPowertypeOfDef)).
fof(f213,plain,(
  spl12_25),
  inference(avatar_split_clause,[],[f87,f211])).
fof(f87,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | ~specializes(X2,X1) | iof(X2,X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f162,plain,(
  spl12_15 | ~spl12_3 | ~spl12_11),
  inference(avatar_split_clause,[],[f152,f140,f108,f160])).
fof(f140,plain,(
  spl12_11 <=> ! [X1,X0] : (~isPowertypeOf(X0,X1) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_11])])).
fof(f152,plain,(
  type_(sK0) | (~spl12_3 | ~spl12_11)),
  inference(resolution,[],[f141,f109])).
fof(f141,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X0,X1) | type_(X0)) ) | ~spl12_11),
  inference(avatar_component_clause,[],[f140])).
fof(f156,plain,(
  spl12_14 | ~spl12_2 | ~spl12_11),
  inference(avatar_split_clause,[],[f151,f140,f104,f154])).
fof(f151,plain,(
  type_(sK2) | (~spl12_2 | ~spl12_11)),
  inference(resolution,[],[f141,f105])).
fof(f142,plain,(
  spl12_11),
  inference(avatar_split_clause,[],[f89,f140])).
fof(f89,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X0,X1) | type_(X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f110,plain,(
  spl12_3),
  inference(avatar_split_clause,[],[f55,f108])).
fof(f55,plain,(
  isPowertypeOf(sK0,sK1)),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ? [X0,X1] : (? [X2] : (isPowertypeOf(X2,X1) & X0 != X2) & isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1] : (isPowertypeOf(X0,X1) => ~? [X2] : (isPowertypeOf(X2,X1) & X0 != X2))),
  inference(rectify,[],[f22])).
fof(f22,negated_conjecture,(
  ~! [X10,X11] : (isPowertypeOf(X10,X11) => ~? [X12] : (isPowertypeOf(X12,X11) & X10 != X12))),
  inference(negated_conjecture,[],[f21])).
fof(f21,conjecture,(
  ! [X10,X11] : (isPowertypeOf(X10,X11) => ~? [X12] : (isPowertypeOf(X12,X11) & X10 != X12))),
  file('/tmp/WCAvVxDflM/SOT_fqme_H',t1_basetypeUnique)).
fof(f106,plain,(
  spl12_2),
  inference(avatar_split_clause,[],[f54,f104])).
fof(f54,plain,(
  isPowertypeOf(sK2,sK1)),
  inference(cnf_transformation,[],[f38])).
fof(f102,plain,(
  ~spl12_1),
  inference(avatar_split_clause,[],[f53,f100])).
fof(f53,plain,(
  sK0 != sK2),
  inference(cnf_transformation,[],[f38])).
% SZS output end Proof for SOT_fqme
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 7675
% Time elapsed: 0.079 s
% ------------------------------
% ------------------------------
% Success in time 0.129 s
% START OF SYSTEM OUTPUT
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_fqme
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_fqme
% SZS output start Proof for SOT_fqme
fof(f2433,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f102,f106,f110,f142,f156,f162,f213,f226,f426,f628,f995,f1673,f1679,f1692,f1713,f1721,f1725,f2422,f2429,f2432])).
fof(f2432,plain,(
  spl12_244 | ~spl12_3 | ~spl12_299),
  inference(avatar_split_clause,[],[f2424,f2420,f108,f1671])).
fof(f1671,plain,(
  spl12_244 <=> iof(sK8(sK2,sK0),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_244])])).
fof(f108,plain,(
  spl12_3 <=> isPowertypeOf(sK0,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_3])])).
fof(f2420,plain,(
  spl12_299 <=> ! [X0] : (~isPowertypeOf(X0,sK1) | iof(sK8(sK2,sK0),X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_299])])).
fof(f2424,plain,(
  iof(sK8(sK2,sK0),sK0) | (~spl12_3 | ~spl12_299)),
  inference(resolution,[],[f2421,f109])).
fof(f109,plain,(
  isPowertypeOf(sK0,sK1) | ~spl12_3),
  inference(avatar_component_clause,[],[f108])).
fof(f2421,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK1) | iof(sK8(sK2,sK0),X0)) ) | ~spl12_299),
  inference(avatar_component_clause,[],[f2420])).
fof(f2429,plain,(
  spl12_243 | ~spl12_2 | ~spl12_299),
  inference(avatar_split_clause,[],[f2423,f2420,f104,f1668])).
fof(f1668,plain,(
  spl12_243 <=> iof(sK8(sK2,sK0),sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_243])])).
fof(f104,plain,(
  spl12_2 <=> isPowertypeOf(sK2,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_2])])).
fof(f2423,plain,(
  iof(sK8(sK2,sK0),sK2) | (~spl12_2 | ~spl12_299)),
  inference(resolution,[],[f2421,f105])).
fof(f105,plain,(
  isPowertypeOf(sK2,sK1) | ~spl12_2),
  inference(avatar_component_clause,[],[f104])).
fof(f2422,plain,(
  spl12_299 | ~spl12_25 | ~spl12_251),
  inference(avatar_split_clause,[],[f1736,f1719,f211,f2420])).
fof(f211,plain,(
  spl12_25 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | ~specializes(X2,X1) | iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_25])])).
fof(f1719,plain,(
  spl12_251 <=> specializes(sK8(sK2,sK0),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_251])])).
fof(f1736,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK1) | iof(sK8(sK2,sK0),X0)) ) | (~spl12_25 | ~spl12_251)),
  inference(resolution,[],[f1720,f212])).
fof(f212,plain,(
  ( ! [X2,X0,X1] : (~specializes(X2,X1) | ~isPowertypeOf(X0,X1) | iof(X2,X0)) ) | ~spl12_25),
  inference(avatar_component_clause,[],[f211])).
fof(f1720,plain,(
  specializes(sK8(sK2,sK0),sK1) | ~spl12_251),
  inference(avatar_component_clause,[],[f1719])).
fof(f1725,plain,(
  spl12_251 | ~spl12_2 | ~spl12_250),
  inference(avatar_split_clause,[],[f1724,f1711,f104,f1719])).
fof(f1711,plain,(
  spl12_250 <=> ! [X0] : (specializes(sK8(sK2,sK0),X0) | ~isPowertypeOf(sK2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_250])])).
fof(f1724,plain,(
  specializes(sK8(sK2,sK0),sK1) | (~spl12_2 | ~spl12_250)),
  inference(resolution,[],[f1712,f105])).
fof(f1712,plain,(
  ( ! [X0] : (~isPowertypeOf(sK2,X0) | specializes(sK8(sK2,sK0),X0)) ) | ~spl12_250),
  inference(avatar_component_clause,[],[f1711])).
fof(f1721,plain,(
  spl12_251 | ~spl12_3 | ~spl12_247),
  inference(avatar_split_clause,[],[f1717,f1690,f108,f1719])).
fof(f1690,plain,(
  spl12_247 <=> ! [X0] : (specializes(sK8(sK2,sK0),X0) | ~isPowertypeOf(sK0,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_247])])).
fof(f1717,plain,(
  specializes(sK8(sK2,sK0),sK1) | (~spl12_3 | ~spl12_247)),
  inference(resolution,[],[f1691,f109])).
fof(f1691,plain,(
  ( ! [X0] : (~isPowertypeOf(sK0,X0) | specializes(sK8(sK2,sK0),X0)) ) | ~spl12_247),
  inference(avatar_component_clause,[],[f1690])).
fof(f1713,plain,(
  spl12_250 | ~spl12_28 | ~spl12_243),
  inference(avatar_split_clause,[],[f1693,f1668,f224,f1711])).
fof(f224,plain,(
  spl12_28 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_28])])).
fof(f1693,plain,(
  ( ! [X0] : (specializes(sK8(sK2,sK0),X0) | ~isPowertypeOf(sK2,X0)) ) | (~spl12_28 | ~spl12_243)),
  inference(resolution,[],[f1669,f225])).
fof(f225,plain,(
  ( ! [X2,X0,X1] : (~iof(X2,X0) | specializes(X2,X1) | ~isPowertypeOf(X0,X1)) ) | ~spl12_28),
  inference(avatar_component_clause,[],[f224])).
fof(f1669,plain,(
  iof(sK8(sK2,sK0),sK2) | ~spl12_243),
  inference(avatar_component_clause,[],[f1668])).
fof(f1692,plain,(
  spl12_247 | ~spl12_28 | ~spl12_244),
  inference(avatar_split_clause,[],[f1675,f1671,f224,f1690])).
fof(f1675,plain,(
  ( ! [X0] : (specializes(sK8(sK2,sK0),X0) | ~isPowertypeOf(sK0,X0)) ) | (~spl12_28 | ~spl12_244)),
  inference(resolution,[],[f1672,f225])).
fof(f1672,plain,(
  iof(sK8(sK2,sK0),sK0) | ~spl12_244),
  inference(avatar_component_clause,[],[f1671])).
fof(f1679,plain,(
  ~spl12_243 | spl12_1 | ~spl12_65 | ~spl12_244),
  inference(avatar_split_clause,[],[f1677,f1671,f424,f100,f1668])).
fof(f100,plain,(
  spl12_1 <=> sK0 = sK2),
  introduced(avatar_definition,[new_symbols(naming,[spl12_1])])).
fof(f424,plain,(
  spl12_65 <=> ! [X1,X0] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_65])])).
fof(f1677,plain,(
  ~iof(sK8(sK2,sK0),sK2) | (spl12_1 | ~spl12_65 | ~spl12_244)),
  inference(subsumption_resolution,[],[f1674,f101])).
fof(f101,plain,(
  sK0 != sK2 | spl12_1),
  inference(avatar_component_clause,[],[f100])).
fof(f1674,plain,(
  ~iof(sK8(sK2,sK0),sK2) | sK0 = sK2 | (~spl12_65 | ~spl12_244)),
  inference(resolution,[],[f1672,f425])).
fof(f425,plain,(
  ( ! [X0,X1] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) ) | ~spl12_65),
  inference(avatar_component_clause,[],[f424])).
fof(f1673,plain,(
  spl12_243 | spl12_244 | spl12_1 | ~spl12_14 | ~spl12_151),
  inference(avatar_split_clause,[],[f1014,f993,f154,f100,f1671,f1668])).
fof(f154,plain,(
  spl12_14 <=> type_(sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_14])])).
fof(f993,plain,(
  spl12_151 <=> ! [X0] : (~type_(X0) | iof(sK8(X0,sK0),sK0) | iof(sK8(X0,sK0),X0) | sK0 = X0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_151])])).
fof(f1014,plain,(
  iof(sK8(sK2,sK0),sK0) | iof(sK8(sK2,sK0),sK2) | (spl12_1 | ~spl12_14 | ~spl12_151)),
  inference(subsumption_resolution,[],[f1010,f101])).
fof(f1010,plain,(
  iof(sK8(sK2,sK0),sK0) | iof(sK8(sK2,sK0),sK2) | sK0 = sK2 | (~spl12_14 | ~spl12_151)),
  inference(resolution,[],[f994,f155])).
fof(f155,plain,(
  type_(sK2) | ~spl12_14),
  inference(avatar_component_clause,[],[f154])).
fof(f994,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK8(X0,sK0),sK0) | iof(sK8(X0,sK0),X0) | sK0 = X0) ) | ~spl12_151),
  inference(avatar_component_clause,[],[f993])).
fof(f995,plain,(
  spl12_151 | ~spl12_15 | ~spl12_97),
  inference(avatar_split_clause,[],[f652,f626,f160,f993])).
fof(f160,plain,(
  spl12_15 <=> type_(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_15])])).
fof(f626,plain,(
  spl12_97 <=> ! [X1,X0] : (~type_(X0) | ~type_(X1) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_97])])).
fof(f652,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK8(X0,sK0),sK0) | iof(sK8(X0,sK0),X0) | sK0 = X0) ) | (~spl12_15 | ~spl12_97)),
  inference(resolution,[],[f627,f161])).
fof(f161,plain,(
  type_(sK0) | ~spl12_15),
  inference(avatar_component_clause,[],[f160])).
fof(f627,plain,(
  ( ! [X0,X1] : (~type_(X1) | ~type_(X0) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1) ) | ~spl12_97),
  inference(avatar_component_clause,[],[f626])).
fof(f628,plain,(
  spl12_97),
  inference(avatar_split_clause,[],[f73,f626])).
fof(f73,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~type_(X1) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1] : ((X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))) | ~type_(X1) | ~type_(X0))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : ((X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))) | (~type_(X1) | ~type_(X0)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : ((type_(X1) & type_(X0)) => (X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7] : ((type_(X7) & type_(X6)) => (X6 = X7 <=> ! [X0] : (iof(X0,X6) <=> iof(X0,X7))))),
  file('/tmp/WCAvVxDflM/SOT_fqme_H',a9_equality)).
fof(f426,plain,(
  spl12_65),
  inference(avatar_split_clause,[],[f95,f424])).
fof(f95,plain,(
  ( ! [X0,X1] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(subsumption_resolution,[],[f94,f67])).
fof(f67,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/WCAvVxDflM/SOT_fqme_H',a3)).
fof(f94,plain,(
  ( ! [X0,X1] : (~type_(X1) | ~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(subsumption_resolution,[],[f74,f67])).
fof(f74,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~type_(X1) | ~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f46])).
fof(f226,plain,(
  spl12_28),
  inference(avatar_split_clause,[],[f88,f224])).
fof(f88,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : ((! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)) | ~isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) => (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) <=> (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : (isPowertypeOf(X6,X7) <=> (! [X8] : (iof(X8,X6) <=> specializes(X8,X7)) & type_(X6)))),
  file('/tmp/WCAvVxDflM/SOT_fqme_H',a10_isPowertypeOfDef)).
fof(f213,plain,(
  spl12_25),
  inference(avatar_split_clause,[],[f87,f211])).
fof(f87,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | ~specializes(X2,X1) | iof(X2,X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f162,plain,(
  spl12_15 | ~spl12_3 | ~spl12_11),
  inference(avatar_split_clause,[],[f152,f140,f108,f160])).
fof(f140,plain,(
  spl12_11 <=> ! [X1,X0] : (~isPowertypeOf(X0,X1) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_11])])).
fof(f152,plain,(
  type_(sK0) | (~spl12_3 | ~spl12_11)),
  inference(resolution,[],[f141,f109])).
fof(f141,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X0,X1) | type_(X0)) ) | ~spl12_11),
  inference(avatar_component_clause,[],[f140])).
fof(f156,plain,(
  spl12_14 | ~spl12_2 | ~spl12_11),
  inference(avatar_split_clause,[],[f151,f140,f104,f154])).
fof(f151,plain,(
  type_(sK2) | (~spl12_2 | ~spl12_11)),
  inference(resolution,[],[f141,f105])).
fof(f142,plain,(
  spl12_11),
  inference(avatar_split_clause,[],[f89,f140])).
fof(f89,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X0,X1) | type_(X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f110,plain,(
  spl12_3),
  inference(avatar_split_clause,[],[f55,f108])).
fof(f55,plain,(
  isPowertypeOf(sK0,sK1)),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ? [X0,X1] : (? [X2] : (isPowertypeOf(X2,X1) & X0 != X2) & isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1] : (isPowertypeOf(X0,X1) => ~? [X2] : (isPowertypeOf(X2,X1) & X0 != X2))),
  inference(rectify,[],[f22])).
fof(f22,negated_conjecture,(
  ~! [X10,X11] : (isPowertypeOf(X10,X11) => ~? [X12] : (isPowertypeOf(X12,X11) & X10 != X12))),
  inference(negated_conjecture,[],[f21])).
fof(f21,conjecture,(
  ! [X10,X11] : (isPowertypeOf(X10,X11) => ~? [X12] : (isPowertypeOf(X12,X11) & X10 != X12))),
  file('/tmp/WCAvVxDflM/SOT_fqme_H',t1_basetypeUnique)).
fof(f106,plain,(
  spl12_2),
  inference(avatar_split_clause,[],[f54,f104])).
fof(f54,plain,(
  isPowertypeOf(sK2,sK1)),
  inference(cnf_transformation,[],[f38])).
fof(f102,plain,(
  ~spl12_1),
  inference(avatar_split_clause,[],[f53,f100])).
fof(f53,plain,(
  sK0 != sK2),
  inference(cnf_transformation,[],[f38])).
% SZS output end Proof for SOT_fqme
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 7675
% Time elapsed: 0.089 s
% ------------------------------
% ------------------------------
% Success in time 0.161 s

% END OF SYSTEM OUTPUT

% END OF SYSTEM OUTPUT

% END OF SYSTEM OUTPUT
% RESULT: SOT_fqme_H - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.14 
% OUTPUT: SOT_fqme_H - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.14 
% RESULT: SOT_fqme_H - E---FNT-2.5 says Theorem - CPU = 0.00 WC = 0.16 
% OUTPUT: SOT_fqme_H - E---FNT-2.5 says CNFRefutation - CPU = 0.00 WC = 0.16 
% RESULT: SOT_fqme_H - E---FNT-2.5 says Theorem - CPU = 0.00 WC = 0.15 
% OUTPUT: SOT_fqme_H - E---FNT-2.5 says CNFRefutation - CPU = 0.00 WC = 0.15 

% END OF SYSTEM OUTPUT
% RESULT: SOT_fqme_H - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.14 
% OUTPUT: SOT_fqme_H - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.14 

% END OF SYSTEM OUTPUT

% END OF SYSTEM OUTPUT
% RESULT: SOT_fqme_H - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.17 
% OUTPUT: SOT_fqme_H - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.17 
% RESULT: SOT_fqme_H - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.13 
% OUTPUT: SOT_fqme_H - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.13 
Vampire---SAT-4.5 has ended (PID 97302)
Vampire---SAT-4.5 has ended (PID 97304)
Vampire---SAT-4.5 has ended (PID 97305)
Vampire---SAT-4.5 has ended (PID 97307)
E---FNT-2.5 has ended (PID 97309)
E---FNT-2.5 has ended (PID 97311)
RESULT: SOT_fqme_H - SSCPA---0.0 says Theorem - CPU = 4.1 WC = 6.7  SolvingSystem = Vampire---SAT-4.5
</PRE>

</body>
</html>