<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>System Results</title>
<meta name="ROBOTS" content="INDEX,NOFOLLOW" />
<script type="text/javascript">//<![CDATA[

  var tstp_dump;
  function openSoTSTP (dump) {   
    var tstp_url = 'http://www.tptp.org/cgi-bin/SystemOnTSTP';
    var tstp_browser = window.open(tstp_url, '_blank');
    tstp_dump = dump;
  }
  function getTSTPDump () {
    return tstp_dump;
  }

//]]></script>
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>

% SZS start RequiredInformation<BR>
% Congratulations - you have become a registered power user of SystemOnTPTP, at IP address 191.6.38.118.<BR>
% Please consider donating to the TPTP project - see www.tptp.org for details.<BR>
% When you donate this message will disappear.<BR>
% If you do not donate a random delay might be added to your processing time.<BR>
% SZS end RequiredInformation<BR>
<PRE>
Getting system recommendations for generic problem /tmp/i1lpDJS7Ag/SOT_M7drGb
This is a FOF_THM_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.95
Vampire---TF1-4.5    recommended  0.95
Vampire---4.5        recommended  0.95
Vampire---SAT-4.5    recommended  0.95
ET---2.0             recommended  0.80
E---2.5              recommended  0.78
E---FNT-2.5          recommended  0.78
CSE_E---1.0          recommended  0.76
E---FNT-2.5          recommended  0.74
Isabelle---2TH0      recommended  0.69
Isabelle---2TF0      recommended  0.69
Isabelle---2FOF      recommended  0.69
Isabelle---2016      recommended  0.69
Isabelle-HOT---2016  recommended  0.68
CVC4---1.7           recommended  0.68
CVC4---SAT-1.7       recommended  0.68
Leo-III---1.4        recommended  0.61
Vampire---THF-4.5    recommended  0.59
Vampire---TF1-4.5    recommended  0.59
Vampire---SAT-4.5    recommended  0.59
iProver---3.0        recommended  0.58
iProver---SAT-3.0    recommended  0.58
Fampire---1.3        recommended  0.57
SInE---0.4           recommended  0.53
iProver---SAT-3.0    recommended  0.52
ePrincess---1.0      recommended  0.50
Princess---170717    recommended  0.48
SPASS---3.9          recommended  0.48
Z3---4.4.1           recommended  0.46
Zipperpin---1.5      recommended  0.40
CVC4---SAT-1.7       recommended  0.40
Prover9---1109a      recommended  0.39
Equinox---6.0.1a     recommended  0.37
LEO-II---1.7.0       recommended  0.35
nanoCoP---1.1        recommended  0.34
Muscadet---4.5       recommended  0.25
Twee---2.2           recommended  0.23
Metis---2.4          recommended  0.23
Geo-III---2018C      recommended  0.18
iProverMo---2.5-0.1  recommended  0.14
Beagle---0.9.47      subsumed     0.41
leanCoP---2.2        subsumed     0.36
E-Darwin---1.5       subsumed     0.31
CSE---1.0            subsumed     0.28
Darwin---1.4.5       subsumed     0.28
Bliksem---1.12       subsumed     0.21
Zenon---0.7.1        subsumed     0.16
ZenonModulo---0.4.2  subsumed     0.15
SOS---2.0            subsumed     0.11
Otter---3.3          subsumed     0.04
This is a FOF_CSA_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.98
Vampire---TF1-4.5    recommended  0.98
Vampire---SAT-4.5    recommended  0.98
Paradox---4.0        recommended  0.60
Nitpick---2016       subsumed     0.55
DarwinFM---1.4.5     subsumed     0.43
CVC4---SAT-1.7       subsumed     0.40
Vampire---THF-4.5    subsumed     0.34
Vampire---TF1-4.5    subsumed     0.34
Vampire---4.5        subsumed     0.34
Vampire---SAT-4.5    subsumed     0.34
E---2.5              subsumed     0.34
E---FNT-2.5          subsumed     0.34
E---FNT-2.5          subsumed     0.34
iProver---SAT-3.0    subsumed     0.32
Metis---2.4          subsumed     0.32
SPASS---3.9          subsumed     0.28
Beagle---0.9.47      subsumed     0.26
Geo-III---2018C      subsumed     0.25
Z3---4.4.1           subsumed     0.25
Princess---170717    subsumed     0.25
Equinox---6.0.1a     subsumed     0.19
iProver---3.0        subsumed     0.11
iProver---SAT-3.0    subsumed     0.11
Darwin---1.4.5       subsumed     0.06
E-Darwin---1.5       subsumed     0.04
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use ET,2.0
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use CSE_E,1.0
CSE_E---1.0 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use Isabelle,2TH0
Isabelle---2016 can be used
Check if we can use Isabelle,2TF0
Isabelle---2016 can be used
Check if we can use Isabelle,2FOF
Isabelle---2016 can be used
Check if we can use Isabelle,2016
Isabelle---2016 can be used
Check if we can use Isabelle-HOT,2016
Isabelle-HOT---2016 can be used
Check if we can use CVC4,1.7
CVC4---SAT-1.7 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Leo-III,1.4
Leo-III---1.4 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Fampire,1.3
Check if we can use SInE,0.4
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use ePrincess,1.0
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Zipperpin,1.5
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Prover9,1109a
Prover9---1109a can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use LEO-II,1.7.0
LEO-II---1.7.0 can be used
Check if we can use nanoCoP,1.1
nanoCoP---1.1 can be used
Check if we can use Muscadet,4.5
Muscadet---4.5 can be used
Check if we can use Twee,2.2
Twee---2.2 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use iProverMo,2.5-0.1
iProverMo---2.5-0.1 can be used
Check if we can use Beagle,0.9.47
Check if we can use leanCoP,2.2
leanCoP---2.2 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Check if we can use CSE,1.0
CSE---1.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use Bliksem,1.12
Bliksem---1.12 can be used
Check if we can use Zenon,0.7.1
Zenon---0.7.1 can be used
Check if we can use ZenonModulo,0.4.2
Check if we can use SOS,2.0
SOS---2.0 can be used
Check if we can use Otter,3.3
Otter---3.3 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use Paradox,4.0
Paradox---4.0 can be used
Check if we can use Nitpick,2016
Nitpick---2016 can be used
Check if we can use DarwinFM,1.4.5
DarwinFM---1.4.5 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Beagle,0.9.47
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
E---FNT-2.5 will be used
E---FNT-2.5 will be used
Vampire---SAT-4.5 is given 499 seconds in -cE mode
Vampire---SAT-4.5 is given 250 seconds in -cE mode
Vampire---SAT-4.5 is given 125 seconds in -cE mode
Vampire---SAT-4.5 is given 62 seconds in -cE mode
E---FNT-2.5 is given 31 seconds in -cE mode
E---FNT-2.5 is given 32 seconds in -cE mode
Vampire---SAT-4.5 has started (PID 97253)
Vampire---SAT-4.5 has started (PID 97254)
Vampire---SAT-4.5 has started (PID 97256)
Vampire---SAT-4.5 has started (PID 97258)
E---FNT-2.5 has started (PID 97261)
E---FNT-2.5 has started (PID 97265)
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  499  /tmp/i1lpDJS7Ag/SOT_M7drGb
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  250  /tmp/i1lpDJS7Ag/SOT_M7drGb
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  125  /tmp/i1lpDJS7Ag/SOT_M7drGb
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  62  /tmp/i1lpDJS7Ag/SOT_M7drGb
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  31  /tmp/i1lpDJS7Ag/SOT_M7drGb
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  32  /tmp/i1lpDJS7Ag/SOT_M7drGb
% START OF SYSTEM OUTPUT
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_M7dr
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_M7dr
% SZS output start Proof for SOT_M7dr
fof(f275,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f104,f108,f116,f124,f140,f149,f157,f163,f189,f196,f205,f230,f254,f260,f267,f274])).
fof(f274,plain,(
  ~spl11_14 | ~spl11_23 | ~spl11_36),
  inference(avatar_contradiction_clause,[],[f273])).
fof(f273,plain,(
  $false | (~spl11_14 | ~spl11_23 | ~spl11_36)),
  inference(subsumption_resolution,[],[f270,f204])).
fof(f204,plain,(
  ( ! [X0] : (~iof(X0,sK6(sK0))) ) | ~spl11_23),
  inference(avatar_component_clause,[],[f203])).
fof(f203,plain,(
  spl11_23 <=> ! [X0] : ~iof(X0,sK6(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_23])])).
fof(f270,plain,(
  iof(sK6(sK6(sK0)),sK6(sK0)) | (~spl11_14 | ~spl11_36)),
  inference(resolution,[],[f266,f156])).
fof(f156,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK6(X0),X0)) ) | ~spl11_14),
  inference(avatar_component_clause,[],[f155])).
fof(f155,plain,(
  spl11_14 <=> ! [X0] : (iof(sK6(X0),X0) | ~type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_14])])).
fof(f266,plain,(
  type_(sK6(sK0)) | ~spl11_36),
  inference(avatar_component_clause,[],[f265])).
fof(f265,plain,(
  spl11_36 <=> type_(sK6(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_36])])).
fof(f267,plain,(
  spl11_36 | ~spl11_10 | ~spl11_35),
  inference(avatar_split_clause,[],[f263,f258,f138,f265])).
fof(f138,plain,(
  spl11_10 <=> ! [X1,X0] : (type_(X0) | ~specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_10])])).
fof(f258,plain,(
  spl11_35 <=> specializes(sK6(sK0),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_35])])).
fof(f263,plain,(
  type_(sK6(sK0)) | (~spl11_10 | ~spl11_35)),
  inference(resolution,[],[f259,f139])).
fof(f139,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) ) | ~spl11_10),
  inference(avatar_component_clause,[],[f138])).
fof(f259,plain,(
  specializes(sK6(sK0),sK1) | ~spl11_35),
  inference(avatar_component_clause,[],[f258])).
fof(f260,plain,(
  spl11_35 | ~spl11_20 | ~spl11_34),
  inference(avatar_split_clause,[],[f255,f252,f187,f258])).
fof(f187,plain,(
  spl11_20 <=> iof(sK6(sK0),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_20])])).
fof(f252,plain,(
  spl11_34 <=> ! [X0] : (specializes(X0,sK1) | ~iof(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_34])])).
fof(f255,plain,(
  specializes(sK6(sK0),sK1) | (~spl11_20 | ~spl11_34)),
  inference(resolution,[],[f253,f188])).
fof(f188,plain,(
  iof(sK6(sK0),sK0) | ~spl11_20),
  inference(avatar_component_clause,[],[f187])).
fof(f253,plain,(
  ( ! [X0] : (~iof(X0,sK0) | specializes(X0,sK1)) ) | ~spl11_34),
  inference(avatar_component_clause,[],[f252])).
fof(f254,plain,(
  spl11_34 | ~spl11_2 | ~spl11_29),
  inference(avatar_split_clause,[],[f250,f228,f106,f252])).
fof(f106,plain,(
  spl11_2 <=> isPowertypeOf(sK0,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_2])])).
fof(f228,plain,(
  spl11_29 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_29])])).
fof(f250,plain,(
  ( ! [X0] : (specializes(X0,sK1) | ~iof(X0,sK0)) ) | (~spl11_2 | ~spl11_29)),
  inference(resolution,[],[f229,f107])).
fof(f107,plain,(
  isPowertypeOf(sK0,sK1) | ~spl11_2),
  inference(avatar_component_clause,[],[f106])).
fof(f229,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) ) | ~spl11_29),
  inference(avatar_component_clause,[],[f228])).
fof(f230,plain,(
  spl11_29),
  inference(avatar_split_clause,[],[f90,f228])).
fof(f90,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : ((! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)) | ~isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) => (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) <=> (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : (isPowertypeOf(X6,X7) <=> (! [X8] : (iof(X8,X6) <=> specializes(X8,X7)) & type_(X6)))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a10_isPowertypeOfDef)).
fof(f205,plain,(
  spl11_23 | ~spl11_6 | ~spl11_21),
  inference(avatar_split_clause,[],[f197,f194,f122,f203])).
fof(f122,plain,(
  spl11_6 <=> ! [X1,X0] : (~individual(X0) | ~iof(X1,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_6])])).
fof(f194,plain,(
  spl11_21 <=> individual(sK6(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_21])])).
fof(f197,plain,(
  ( ! [X0] : (~iof(X0,sK6(sK0))) ) | (~spl11_6 | ~spl11_21)),
  inference(resolution,[],[f195,f123])).
fof(f123,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) ) | ~spl11_6),
  inference(avatar_component_clause,[],[f122])).
fof(f195,plain,(
  individual(sK6(sK0)) | ~spl11_21),
  inference(avatar_component_clause,[],[f194])).
fof(f196,plain,(
  spl11_21 | ~spl11_1 | ~spl11_15 | ~spl11_20),
  inference(avatar_split_clause,[],[f192,f187,f161,f102,f194])).
fof(f102,plain,(
  spl11_1 <=> firstordertype(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_1])])).
fof(f161,plain,(
  spl11_15 <=> ! [X1,X0] : (~firstordertype(X0) | ~iof(X1,X0) | individual(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_15])])).
fof(f192,plain,(
  individual(sK6(sK0)) | (~spl11_1 | ~spl11_15 | ~spl11_20)),
  inference(subsumption_resolution,[],[f190,f103])).
fof(f103,plain,(
  firstordertype(sK0) | ~spl11_1),
  inference(avatar_component_clause,[],[f102])).
fof(f190,plain,(
  ~firstordertype(sK0) | individual(sK6(sK0)) | (~spl11_15 | ~spl11_20)),
  inference(resolution,[],[f188,f162])).
fof(f162,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | ~firstordertype(X0) | individual(X1)) ) | ~spl11_15),
  inference(avatar_component_clause,[],[f161])).
fof(f189,plain,(
  spl11_20 | ~spl11_12 | ~spl11_14),
  inference(avatar_split_clause,[],[f169,f155,f147,f187])).
fof(f147,plain,(
  spl11_12 <=> type_(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_12])])).
fof(f169,plain,(
  iof(sK6(sK0),sK0) | (~spl11_12 | ~spl11_14)),
  inference(resolution,[],[f156,f148])).
fof(f148,plain,(
  type_(sK0) | ~spl11_12),
  inference(avatar_component_clause,[],[f147])).
fof(f163,plain,(
  spl11_15),
  inference(avatar_split_clause,[],[f66,f161])).
fof(f66,plain,(
  ( ! [X0,X1] : (~firstordertype(X0) | ~iof(X1,X0) | individual(X1)) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : ((! [X1] : (individual(X1) | ~iof(X1,X0)) & type_(X0)) | ~firstordertype(X0))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : (firstordertype(X0) => (! [X1] : (iof(X1,X0) => individual(X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f10])).
fof(f10,axiom,(
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (iof(X1,X0) => individual(X1)) & type_(X0)))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a5_firstOrderTypeDef)).
fof(f157,plain,(
  spl11_14),
  inference(avatar_split_clause,[],[f70,f155])).
fof(f70,plain,(
  ( ! [X0] : (iof(sK6(X0),X0) | ~type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a3)).
fof(f149,plain,(
  spl11_12 | ~spl11_1 | ~spl11_4),
  inference(avatar_split_clause,[],[f145,f114,f102,f147])).
fof(f114,plain,(
  spl11_4 <=> ! [X0] : (~firstordertype(X0) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_4])])).
fof(f145,plain,(
  type_(sK0) | (~spl11_1 | ~spl11_4)),
  inference(resolution,[],[f115,f103])).
fof(f115,plain,(
  ( ! [X0] : (~firstordertype(X0) | type_(X0)) ) | ~spl11_4),
  inference(avatar_component_clause,[],[f114])).
fof(f140,plain,(
  spl11_10),
  inference(avatar_split_clause,[],[f88,f138])).
fof(f88,plain,(
  ( ! [X0,X1] : (type_(X0) | ~specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X1) | ~iof(X2,X0)) & type_(X0)))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X0) => iof(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7] : (specializes(X6,X7) <=> (! [X3] : (iof(X3,X6) => iof(X3,X7)) & type_(X6)))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a7_specDef)).
fof(f124,plain,(
  spl11_6),
  inference(avatar_split_clause,[],[f68,f122])).
fof(f68,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (! [X1] : ~iof(X1,X0) | ~individual(X0))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : (individual(X0) => ~? [X1] : iof(X1,X0))),
  inference(unused_predicate_definition_removal,[],[f2])).
fof(f2,axiom,(
  ! [X0] : (individual(X0) <=> ~? [X1] : iof(X1,X0))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a2)).
fof(f116,plain,(
  spl11_4),
  inference(avatar_split_clause,[],[f67,f114])).
fof(f67,plain,(
  ( ! [X0] : (~firstordertype(X0) | type_(X0)) )),
  inference(cnf_transformation,[],[f43])).
fof(f108,plain,(
  spl11_2),
  inference(avatar_split_clause,[],[f54,f106])).
fof(f54,plain,(
  isPowertypeOf(sK0,sK1)),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ? [X0,X1] : (firstordertype(X0) & isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f22])).
fof(f22,negated_conjecture,(
  ~! [X0,X1] : (isPowertypeOf(X0,X1) => ~firstordertype(X0))),
  inference(negated_conjecture,[],[f21])).
fof(f21,conjecture,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) => ~firstordertype(X0))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',powertypeNotFirstOrder)).
fof(f104,plain,(
  spl11_1),
  inference(avatar_split_clause,[],[f55,f102])).
fof(f55,plain,(
  firstordertype(sK0)),
  inference(cnf_transformation,[],[f38])).
% SZS output end Proof for SOT_M7dr
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Success in time 0.048 s
% START OF SYSTEM OUTPUT
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 15 seconds
# AutoSched0-Mode selected heuristic G_E___208_C02CMA_F1_SE_CS_SP_PS_S5PRR_RG_S04AN
# and selection function SelectComplexExceptUniqMaxHorn.
#
# Preprocessing time       : 0.016 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(a6_secondOrderTypeDef, axiom, ![X1]:(secondordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>firstordertype(X2)))), file('/tmp/i1lpDJS7Ag/SOT_M7drGb', a6_secondOrderTypeDef)).
fof(a3, axiom, ![X1]:(type_(X1)<=>?[X2]:iof(X2,X1)), file('/tmp/i1lpDJS7Ag/SOT_M7drGb', a3)).
fof(a2, axiom, ![X1]:(individual(X1)<=>~(?[X2]:iof(X2,X1))), file('/tmp/i1lpDJS7Ag/SOT_M7drGb', a2)).
fof(a5_firstOrderTypeDef, axiom, ![X1]:(firstordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>individual(X2)))), file('/tmp/i1lpDJS7Ag/SOT_M7drGb', a5_firstOrderTypeDef)).
fof(powertypeNotFirstOrder, conjecture, ![X1, X2]:(isPowertypeOf(X1,X2)=>~(firstordertype(X1))), file('/tmp/i1lpDJS7Ag/SOT_M7drGb', powertypeNotFirstOrder)).
fof(a10_isPowertypeOfDef, axiom, ![X7, X8]:(isPowertypeOf(X7,X8)<=>(type_(X7)&![X9]:(iof(X9,X7)<=>specializes(X9,X8)))), file('/tmp/i1lpDJS7Ag/SOT_M7drGb', a10_isPowertypeOfDef)).
fof(a7_specDef, axiom, ![X7, X8]:(specializes(X7,X8)<=>(type_(X7)&![X4]:(iof(X4,X7)=>iof(X4,X8)))), file('/tmp/i1lpDJS7Ag/SOT_M7drGb', a7_specDef)).
fof(c_0_7, plain, ![X45, X46, X47]:(((type_(X45)|~secondordertype(X45))&(~iof(X46,X45)|firstordertype(X46)|~secondordertype(X45)))&((iof(esk10_1(X47),X47)|~type_(X47)|secondordertype(X47))&(~firstordertype(esk10_1(X47))|~type_(X47)|secondordertype(X47)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a6_secondOrderTypeDef])])])])])])).
fof(c_0_8, plain, ![X16, X18, X19]:((~type_(X16)|iof(esk2_1(X16),X16))&(~iof(X19,X18)|type_(X18))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a3])])])])])).
fof(c_0_9, plain, ![X12, X13, X14]:((~individual(X12)|~iof(X13,X12))&(iof(esk1_1(X14),X14)|individual(X14))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a2])])])])])).
fof(c_0_10, plain, ![X41, X42, X43]:(((type_(X41)|~firstordertype(X41))&(~iof(X42,X41)|individual(X42)|~firstordertype(X41)))&((iof(esk9_1(X43),X43)|~type_(X43)|firstordertype(X43))&(~individual(esk9_1(X43))|~type_(X43)|firstordertype(X43)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a5_firstOrderTypeDef])])])])])])).
cnf(c_0_11, plain, (firstordertype(X1)|~iof(X1,X2)|~secondordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (iof(esk2_1(X1),X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (type_(X1)|~secondordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14, plain, (~individual(X1)|~iof(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (type_(X1)|~firstordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (firstordertype(esk2_1(X1))|~secondordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])).
fof(c_0_17, negated_conjecture, ~(![X1, X2]:(isPowertypeOf(X1,X2)=>~(firstordertype(X1)))), inference(assume_negation,[status(cth)],[powertypeNotFirstOrder])).
cnf(c_0_18, plain, (~type_(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_14, c_0_12])).
cnf(c_0_19, plain, (type_(esk2_1(X1))|~secondordertype(X1)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20, plain, (individual(X1)|~iof(X1,X2)|~firstordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_21, plain, ![X62, X63, X64, X65, X66, X67]:(((type_(X62)|~isPowertypeOf(X62,X63))&((~iof(X64,X62)|specializes(X64,X63)|~isPowertypeOf(X62,X63))&(~specializes(X65,X63)|iof(X65,X62)|~isPowertypeOf(X62,X63))))&((~iof(esk13_2(X66,X67),X66)|~specializes(esk13_2(X66,X67),X67)|~type_(X66)|isPowertypeOf(X66,X67))&(iof(esk13_2(X66,X67),X66)|specializes(esk13_2(X66,X67),X67)|~type_(X66)|isPowertypeOf(X66,X67)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a10_isPowertypeOfDef])])])])])])).
fof(c_0_22, negated_conjecture, (isPowertypeOf(esk22_0,esk23_0)&firstordertype(esk22_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])).
cnf(c_0_23, plain, (~secondordertype(X1)|~individual(esk2_1(X1))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24, plain, (individual(esk2_1(X1))|~firstordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_12]), c_0_15])).
fof(c_0_25, plain, ![X49, X50, X51, X52, X53]:(((type_(X49)|~specializes(X49,X50))&(~iof(X51,X49)|iof(X51,X50)|~specializes(X49,X50)))&((iof(esk11_2(X52,X53),X52)|~type_(X52)|specializes(X52,X53))&(~iof(esk11_2(X52,X53),X53)|~type_(X52)|specializes(X52,X53)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a7_specDef])])])])])])).
cnf(c_0_26, plain, (specializes(X1,X3)|~iof(X1,X2)|~isPowertypeOf(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, negated_conjecture, (isPowertypeOf(esk22_0,esk23_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, negated_conjecture, (firstordertype(esk22_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (~secondordertype(X1)|~firstordertype(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_30, plain, (type_(X1)|~specializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, negated_conjecture, (specializes(X1,esk23_0)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32, plain, (iof(esk10_1(X1),X1)|secondordertype(X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_33, negated_conjecture, (type_(esk22_0)), inference(spm,[status(thm)],[c_0_15, c_0_28])).
cnf(c_0_34, negated_conjecture, (~secondordertype(esk22_0)), inference(spm,[status(thm)],[c_0_29, c_0_28])).
cnf(c_0_35, negated_conjecture, (type_(X1)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_36, negated_conjecture, (iof(esk10_1(esk22_0),esk22_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
cnf(c_0_37, negated_conjecture, (type_(esk10_1(esk22_0))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_38, negated_conjecture, (individual(esk10_1(esk22_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_36]), c_0_28])])).
cnf(c_0_39, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_37]), c_0_38])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 40
# Proof object clause steps            : 25
# Proof object formula steps           : 15
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 76
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 71
# Processed clauses                    : 147
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 144
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 62
# ...of the previous two non-trivial   : 48
# Contextual simplify-reflections      : 16
# Paramodulations                      : 60
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 71
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 61
# Current number of unprocessed clauses: 39
# ...number of literals in the above   : 124
# Current number of archived formulas  : 0
# Current number of archived clauses   : 73
# Clause-clause subsumption calls (NU) : 2174
# Rec. Clause-clause subsumption calls : 1467
# Non-unit clause-clause subsumptions  : 17
# Unit Clause-clause subsumption calls : 23
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5994

# -------------------------------------------------
# User time                : 0.019 s
# System time              : 0.005 s
# Total time               : 0.024 s
# Maximum resident set size: 1600 pages
% START OF SYSTEM OUTPUT
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_M7dr
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_M7dr
% SZS output start Proof for SOT_M7dr
fof(f275,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f104,f108,f116,f124,f140,f149,f157,f163,f189,f196,f205,f230,f254,f260,f267,f274])).
fof(f274,plain,(
  ~spl11_14 | ~spl11_23 | ~spl11_36),
  inference(avatar_contradiction_clause,[],[f273])).
fof(f273,plain,(
  $false | (~spl11_14 | ~spl11_23 | ~spl11_36)),
  inference(subsumption_resolution,[],[f270,f204])).
fof(f204,plain,(
  ( ! [X0] : (~iof(X0,sK6(sK0))) ) | ~spl11_23),
  inference(avatar_component_clause,[],[f203])).
fof(f203,plain,(
  spl11_23 <=> ! [X0] : ~iof(X0,sK6(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_23])])).
fof(f270,plain,(
  iof(sK6(sK6(sK0)),sK6(sK0)) | (~spl11_14 | ~spl11_36)),
  inference(resolution,[],[f266,f156])).
fof(f156,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK6(X0),X0)) ) | ~spl11_14),
  inference(avatar_component_clause,[],[f155])).
fof(f155,plain,(
  spl11_14 <=> ! [X0] : (iof(sK6(X0),X0) | ~type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_14])])).
fof(f266,plain,(
  type_(sK6(sK0)) | ~spl11_36),
  inference(avatar_component_clause,[],[f265])).
fof(f265,plain,(
  spl11_36 <=> type_(sK6(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_36])])).
fof(f267,plain,(
  spl11_36 | ~spl11_10 | ~spl11_35),
  inference(avatar_split_clause,[],[f263,f258,f138,f265])).
fof(f138,plain,(
  spl11_10 <=> ! [X1,X0] : (type_(X0) | ~specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_10])])).
fof(f258,plain,(
  spl11_35 <=> specializes(sK6(sK0),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_35])])).
fof(f263,plain,(
  type_(sK6(sK0)) | (~spl11_10 | ~spl11_35)),
  inference(resolution,[],[f259,f139])).
fof(f139,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) ) | ~spl11_10),
  inference(avatar_component_clause,[],[f138])).
fof(f259,plain,(
  specializes(sK6(sK0),sK1) | ~spl11_35),
  inference(avatar_component_clause,[],[f258])).
fof(f260,plain,(
  spl11_35 | ~spl11_20 | ~spl11_34),
  inference(avatar_split_clause,[],[f255,f252,f187,f258])).
fof(f187,plain,(
  spl11_20 <=> iof(sK6(sK0),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_20])])).
fof(f252,plain,(
  spl11_34 <=> ! [X0] : (specializes(X0,sK1) | ~iof(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_34])])).
fof(f255,plain,(
  specializes(sK6(sK0),sK1) | (~spl11_20 | ~spl11_34)),
  inference(resolution,[],[f253,f188])).
fof(f188,plain,(
  iof(sK6(sK0),sK0) | ~spl11_20),
  inference(avatar_component_clause,[],[f187])).
fof(f253,plain,(
  ( ! [X0] : (~iof(X0,sK0) | specializes(X0,sK1)) ) | ~spl11_34),
  inference(avatar_component_clause,[],[f252])).
fof(f254,plain,(
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_M7dr
  spl11_34 | ~spl11_2 | ~spl11_29),
  inference(avatar_split_clause,[],[f250,f228,f106,f252])).
fof(f106,plain,(
  spl11_2 <=> isPowertypeOf(sK0,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_2])])).
fof(f228,plain,(
  spl11_29 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_29])])).
fof(f250,plain,(
  ( ! [X0] : (specializes(X0,sK1) | ~iof(X0,sK0)) ) | (~spl11_2 | ~spl11_29)),
  inference(resolution,[],[f229,f107])).
fof(f107,plain,(
  isPowertypeOf(sK0,sK1) | ~spl11_2),
  inference(avatar_component_clause,[],[f106])).
fof(f229,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) ) | ~spl11_29),
  inference(avatar_component_clause,[],[f228])).
fof(f230,plain,(
  spl11_29),
  inference(avatar_split_clause,[],[f90,f228])).
fof(f90,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : ((! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)) | ~isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) => (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) <=> (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : (isPowertypeOf(X6,X7) <=> (! [X8] : (iof(X8,X6) <=> specializes(X8,X7)) & type_(X6)))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a10_isPowertypeOfDef)).
fof(f205,plain,(
  spl11_23 | ~spl11_6 | ~spl11_21),
  inference(avatar_split_clause,[],[f197,f194,f122,f203])).
fof(f122,plain,(
  spl11_6 <=> ! [X1,X0] : (~individual(X0) | ~iof(X1,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_6])])).
fof(f194,plain,(
  spl11_21 <=> individual(sK6(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_21])])).
fof(f197,plain,(
  ( ! [X0] : (~iof(X0,sK6(sK0))) ) | (~spl11_6 | ~spl11_21)),
  inference(resolution,[],[f195,f123])).
fof(f123,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) ) | ~spl11_6),
  inference(avatar_component_clause,[],[f122])).
fof(f195,plain,(
  individual(sK6(sK0)) | ~spl11_21),
  inference(avatar_component_clause,[],[f194])).
fof(f196,plain,(
  spl11_21 | ~spl11_1 | ~spl11_15 | ~spl11_20),
  inference(avatar_split_clause,[],[f192,f187,f161,f102,f194])).
fof(f102,plain,(
  spl11_1 <=> firstordertype(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_1])])).
fof(f161,plain,(
  spl11_15 <=> ! [X1,X0] : (~firstordertype(X0) | ~iof(X1,X0) | individual(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_15])])).
fof(f192,plain,(
  individual(sK6(sK0)) | (~spl11_1 | ~spl11_15 | ~spl11_20)),
  inference(subsumption_resolution,[],[f190,f103])).
fof(f103,plain,(
  firstordertype(sK0) | ~spl11_1),
  inference(avatar_component_clause,[],[f102])).
fof(f190,plain,(
  ~firstordertype(sK0) | individual(sK6(sK0)) | (~spl11_15 | ~spl11_20)),
  inference(resolution,[],[f188,f162])).
fof(f162,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | ~firstordertype(X0) | individual(X1)) ) | ~spl11_15),
  inference(avatar_component_clause,[],[f161])).
fof(f189,plain,(
  spl11_20 | ~spl11_12 | ~spl11_14),
  inference(avatar_split_clause,[],[f169,f155,f147,f187])).
fof(f147,plain,(
  spl11_12 <=> type_(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_12])])).
fof(f169,plain,(
  iof(sK6(sK0),sK0) | (~spl11_12 | ~spl11_14)),
  inference(resolution,[],[f156,f148])).
fof(f148,plain,(
  type_(sK0) | ~spl11_12),
  inference(avatar_component_clause,[],[f147])).
fof(f163,plain,(
  spl11_15),
  inference(avatar_split_clause,[],[f66,f161])).
fof(f66,plain,(
  ( ! [X0,X1] : (~firstordertype(X0) | ~iof(X1,X0) | individual(X1)) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : ((! [X1] : (individual(X1) | ~iof(X1,X0)) & type_(X0)) | ~firstordertype(X0))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : (firstordertype(X0) => (! [X1] : (iof(X1,X0) => individual(X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f10])).
fof(f10,axiom,(
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (iof(X1,X0) => individual(X1)) & type_(X0)))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a5_firstOrderTypeDef)).
fof(f157,plain,(
  spl11_14),
  inference(avatar_split_clause,[],[f70,f155])).
fof(f70,plain,(
  ( ! [X0] : (iof(sK6(X0),X0) | ~type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a3)).
fof(f149,plain,(
  spl11_12 | ~spl11_1 | ~spl11_4),
  inference(avatar_split_clause,[],[f145,f114,f102,f147])).
fof(f114,plain,(
  spl11_4 <=> ! [X0] : (~firstordertype(X0) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_4])])).
fof(f145,plain,(
  type_(sK0) | (~spl11_1 | ~spl11_4)),
  inference(resolution,[],[f115,f103])).
fof(f115,plain,(
  ( ! [X0] : (~firstordertype(X0) | type_(X0)) ) | ~spl11_4),
  inference(avatar_component_clause,[],[f114])).
fof(f140,plain,(
  spl11_10),
  inference(avatar_split_clause,[],[f88,f138])).
fof(f88,plain,(
  ( ! [X0,X1] : (type_(X0) | ~specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X1) | ~iof(X2,X0)) & type_(X0)))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X0) => iof(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7] : (specializes(X6,X7) <=> (! [X3] : (iof(X3,X6) => iof(X3,X7)) & type_(X6)))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a7_specDef)).
fof(f124,plain,(
  spl11_6),
  inference(avatar_split_clause,[],[f68,f122])).
fof(f68,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (! [X1] : ~iof(X1,X0) | ~individual(X0))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : (individual(X0) => ~? [X1] : iof(X1,X0))),
  inference(unused_predicate_definition_removal,[],[f2])).
fof(f2,axiom,(
  ! [X0] : (individual(X0) <=> ~? [X1] : iof(X1,X0))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a2)).
fof(f116,plain,(
  spl11_4),
  inference(avatar_split_clause,[],[f67,f114])).
fof(f67,plain,(
  ( ! [X0] : (~firstordertype(X0) | type_(X0)) )),
  inference(cnf_transformation,[],[f43])).
fof(f108,plain,(
  spl11_2),
  inference(avatar_split_clause,[],[f54,f106])).
fof(f54,plain,(
  isPowertypeOf(sK0,sK1)),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ? [X0,X1] : (firstordertype(X0) & isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f22])).
fof(f22,negated_conjecture,(
  ~! [X0,X1] : (isPowertypeOf(X0,X1) => ~firstordertype(X0))),
  inference(negated_conjecture,[],[f21])).
fof(f21,conjecture,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) => ~firstordertype(X0))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',powertypeNotFirstOrder)).
fof(f104,plain,(
  spl11_1),
  inference(avatar_split_clause,[],[f55,f102])).
fof(f55,plain,(
  firstordertype(sK0)),
  inference(cnf_transformation,[],[f38])).
% SZS output end Proof for SOT_M7dr
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Success in time 0.046 s
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_M7dr
% SZS output start Proof for SOT_M7dr
fof(f275,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f104,f108,f116,f124,f140,f149,f157,f163,f189,f196,f205,f230,f254,f260,f267,f274])).
fof(f274,plain,(
  ~spl11_14 | ~spl11_23 | ~spl11_36),
  inference(avatar_contradiction_clause,[],[f273])).
fof(f273,plain,(
  $false | (~spl11_14 | ~spl11_23 | ~spl11_36)),
  inference(subsumption_resolution,[],[f270,f204])).
fof(f204,plain,(
  ( ! [X0] : (~iof(X0,sK6(sK0))) ) | ~spl11_23),
  inference(avatar_component_clause,[],[f203])).
fof(f203,plain,(
  spl11_23 <=> ! [X0] : ~iof(X0,sK6(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_23])])).
fof(f270,plain,(
  iof(sK6(sK6(sK0)),sK6(sK0)) | (~spl11_14 | ~spl11_36)),
  inference(resolution,[],[f266,f156])).
fof(f156,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK6(X0),X0)) ) | ~spl11_14),
  inference(avatar_component_clause,[],[f155])).
fof(f155,plain,(
  spl11_14 <=> ! [X0] : (iof(sK6(X0),X0) | ~type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_14])])).
fof(f266,plain,(
  type_(sK6(sK0)) | ~spl11_36),
  inference(avatar_component_clause,[],[f265])).
fof(f265,plain,(
  spl11_36 <=> type_(sK6(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_36])])).
fof(f267,plain,(
  spl11_36 | ~spl11_10 | ~spl11_35),
  inference(avatar_split_clause,[],[f263,f258,f138,f265])).
fof(f138,plain,(
  spl11_10 <=> ! [X1,X0] : (type_(X0) | ~specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_10])])).
fof(f258,plain,(
  spl11_35 <=> specializes(sK6(sK0),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_35])])).
fof(f263,plain,(
  type_(sK6(sK0)) | (~spl11_10 | ~spl11_35)),
  inference(resolution,[],[f259,f139])).
fof(f139,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) ) | ~spl11_10),
  inference(avatar_component_clause,[],[f138])).
fof(f259,plain,(
  specializes(sK6(sK0),sK1) | ~spl11_35),
  inference(avatar_component_clause,[],[f258])).
fof(f260,plain,(
  spl11_35 | ~spl11_20 | ~spl11_34),
  inference(avatar_split_clause,[],[f255,f252,f187,f258])).
fof(f187,plain,(
  spl11_20 <=> iof(sK6(sK0),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_20])])).
fof(f252,plain,(
  spl11_34 <=> ! [X0] : (specializes(X0,sK1) | ~iof(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_34])])).
fof(f255,plain,(
  specializes(sK6(sK0),sK1) | (~spl11_20 | ~spl11_34)),
  inference(resolution,[],[f253,f188])).
fof(f188,plain,(
  iof(sK6(sK0),sK0) | ~spl11_20),
  inference(avatar_component_clause,[],[f187])).
fof(f253,plain,(
  ( ! [X0] : (~iof(X0,sK0) | specializes(X0,sK1)) ) | ~spl11_34),
  inference(avatar_component_clause,[],[f252])).
fof(f254,plain,(
  spl11_34 | ~spl11_2 | ~spl11_29),
  inference(avatar_split_clause,[],[f250,f228,f106,f252])).
fof(f106,plain,(
  spl11_2 <=> isPowertypeOf(sK0,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_2])])).
fof(f228,plain,(
  spl11_29 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_29])])).
fof(f250,plain,(
  ( ! [X0] : (specializes(X0,sK1) | ~iof(X0,sK0)) ) | (~spl11_2 | ~spl11_29)),
  inference(resolution,[],[f229,f107])).
fof(f107,plain,(
  isPowertypeOf(sK0,sK1) | ~spl11_2),
  inference(avatar_component_clause,[],[f106])).
fof(f229,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) ) | ~spl11_29),
  inference(avatar_component_clause,[],[f228])).
fof(f230,plain,(
  spl11_29),
  inference(avatar_split_clause,[],[f90,f228])).
fof(f90,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : ((! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)) | ~isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) => (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) <=> (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : (isPowertypeOf(X6,X7) <=> (! [X8] : (iof(X8,X6) <=> specializes(X8,X7)) & type_(X6)))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a10_isPowertypeOfDef)).
fof(f205,plain,(
  spl11_23 | ~spl11_6 | ~spl11_21),
  inference(avatar_split_clause,[],[f197,f194,f122,f203])).
fof(f122,plain,(
  spl11_6 <=> ! [X1,X0] : (~individual(X0) | ~iof(X1,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_6])])).
fof(f194,plain,(
  spl11_21 <=> individual(sK6(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_21])])).
fof(f197,plain,(
  ( ! [X0] : (~iof(X0,sK6(sK0))) ) | (~spl11_6 | ~spl11_21)),
  inference(resolution,[],[f195,f123])).
fof(f123,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) ) | ~spl11_6),
  inference(avatar_component_clause,[],[f122])).
fof(f195,plain,(
  individual(sK6(sK0)) | ~spl11_21),
  inference(avatar_component_clause,[],[f194])).
fof(f196,plain,(
  spl11_21 | ~spl11_1 | ~spl11_15 | ~spl11_20),
  inference(avatar_split_clause,[],[f192,f187,f161,f102,f194])).
fof(f102,plain,(
  spl11_1 <=> firstordertype(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_1])])).
fof(f161,plain,(
  spl11_15 <=> ! [X1,X0] : (~firstordertype(X0) | ~iof(X1,X0) | individual(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_15])])).
fof(f192,plain,(
  individual(sK6(sK0)) | (~spl11_1 | ~spl11_15 | ~spl11_20)),
  inference(subsumption_resolution,[],[f190,f103])).
fof(f103,plain,(
  firstordertype(sK0) | ~spl11_1),
  inference(avatar_component_clause,[],[f102])).
fof(f190,plain,(
  ~firstordertype(sK0) | individual(sK6(sK0)) | (~spl11_15 | ~spl11_20)),
  inference(resolution,[],[f188,f162])).
fof(f162,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | ~firstordertype(X0) | individual(X1)) ) | ~spl11_15),
  inference(avatar_component_clause,[],[f161])).
fof(f189,plain,(
  spl11_20 | ~spl11_12 | ~spl11_14),
  inference(avatar_split_clause,[],[f169,f155,f147,f187])).
fof(f147,plain,(
  spl11_12 <=> type_(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_12])])).
fof(f169,plain,(
  iof(sK6(sK0),sK0) | (~spl11_12 | ~spl11_14)),
  inference(resolution,[],[f156,f148])).
fof(f148,plain,(
  type_(sK0) | ~spl11_12),
  inference(avatar_component_clause,[],[f147])).
fof(f163,plain,(
  spl11_15),
  inference(avatar_split_clause,[],[f66,f161])).
fof(f66,plain,(
  ( ! [X0,X1] : (~firstordertype(X0) | ~iof(X1,X0) | individual(X1)) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : ((! [X1] : (individual(X1) | ~iof(X1,X0)) & type_(X0)) | ~firstordertype(X0))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : (firstordertype(X0) => (! [X1] : (iof(X1,X0) => individual(X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f10])).
fof(f10,axiom,(
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (iof(X1,X0) => individual(X1)) & type_(X0)))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a5_firstOrderTypeDef)).
fof(f157,plain,(
  spl11_14),
  inference(avatar_split_clause,[],[f70,f155])).
fof(f70,plain,(
  ( ! [X0] : (iof(sK6(X0),X0) | ~type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a3)).
fof(f149,plain,(
  spl11_12 | ~spl11_1 | ~spl11_4),
  inference(avatar_split_clause,[],[f145,f114,f102,f147])).
fof(f114,plain,(
  spl11_4 <=> ! [X0] : (~firstordertype(X0) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_4])])).
fof(f145,plain,(
  type_(sK0) | (~spl11_1 | ~spl11_4)),
  inference(resolution,[],[f115,f103])).
fof(f115,plain,(
  ( ! [X0] : (~firstordertype(X0) | type_(X0)) ) | ~spl11_4),
  inference(avatar_component_clause,[],[f114])).
fof(f140,plain,(
  spl11_10),
  inference(avatar_split_clause,[],[f88,f138])).
fof(f88,plain,(
  ( ! [X0,X1] : (type_(X0) | ~specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X1) | ~iof(X2,X0)) & type_(X0)))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X0) => iof(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7] : (specializes(X6,X7) <=> (! [X3] : (iof(X3,X6) => iof(X3,X7)) & type_(X6)))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a7_specDef)).
fof(f124,plain,(
  spl11_6),
  inference(avatar_split_clause,[],[f68,f122])).
fof(f68,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (! [X1] : ~iof(X1,X0) | ~individual(X0))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : (individual(X0) => ~? [X1] : iof(X1,X0))),
  inference(unused_predicate_definition_removal,[],[f2])).
fof(f2,axiom,(
  ! [X0] : (individual(X0) <=> ~? [X1] : iof(X1,X0))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a2)).
fof(f116,plain,(
  spl11_4),
  inference(avatar_split_clause,[],[f67,f114])).
fof(f67,plain,(
  ( ! [X0] : (~firstordertype(X0) | type_(X0)) )),
  inference(cnf_transformation,[],[f43])).
fof(f108,plain,(
  spl11_2),
  inference(avatar_split_clause,[],[f54,f106])).
fof(f54,plain,(
  isPowertypeOf(sK0,sK1)),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ? [X0,X1] : (firstordertype(X0) & isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f22])).
fof(f22,negated_conjecture,(
  ~! [X0,X1] : (isPowertypeOf(X0,X1) => ~firstordertype(X0))),
  inference(negated_conjecture,[],[f21])).
fof(f21,conjecture,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) => ~firstordertype(X0))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',powertypeNotFirstOrder)).
fof(f104,plain,(
  spl11_1),
  inference(avatar_split_clause,[],[f55,f102])).
fof(f55,plain,(
  firstordertype(sK0)),
  inference(cnf_transformation,[],[f38])).
% SZS output end Proof for SOT_M7dr
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.010 s
% ------------------------------
% ------------------------------
% Success in time 0.064 s
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 15 seconds
% START OF SYSTEM OUTPUT
# AutoSched0-Mode selected heuristic G_E___208_C02CMA_F1_SE_CS_SP_PS_S5PRR_RG_S04AN
# and selection function SelectComplexExceptUniqMaxHorn.
#
# Preprocessing time       : 0.026 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(a6_secondOrderTypeDef, axiom, ![X1]:(secondordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>firstordertype(X2)))), file('/tmp/i1lpDJS7Ag/SOT_M7drGb', a6_secondOrderTypeDef)).
fof(a3, axiom, ![X1]:(type_(X1)<=>?[X2]:iof(X2,X1)), file('/tmp/i1lpDJS7Ag/SOT_M7drGb', a3)).
fof(a2, axiom, ![X1]:(individual(X1)<=>~(?[X2]:iof(X2,X1))), file('/tmp/i1lpDJS7Ag/SOT_M7drGb', a2)).
fof(a5_firstOrderTypeDef, axiom, ![X1]:(firstordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>individual(X2)))), file('/tmp/i1lpDJS7Ag/SOT_M7drGb', a5_firstOrderTypeDef)).
fof(powertypeNotFirstOrder, conjecture, ![X1, X2]:(isPowertypeOf(X1,X2)=>~(firstordertype(X1))), file('/tmp/i1lpDJS7Ag/SOT_M7drGb', powertypeNotFirstOrder)).
fof(a10_isPowertypeOfDef, axiom, ![X7, X8]:(isPowertypeOf(X7,X8)<=>(type_(X7)&![X9]:(iof(X9,X7)<=>specializes(X9,X8)))), file('/tmp/i1lpDJS7Ag/SOT_M7drGb', a10_isPowertypeOfDef)).
fof(a7_specDef, axiom, ![X7, X8]:(specializes(X7,X8)<=>(type_(X7)&![X4]:(iof(X4,X7)=>iof(X4,X8)))), file('/tmp/i1lpDJS7Ag/SOT_M7drGb', a7_specDef)).
fof(c_0_7, plain, ![X45, X46, X47]:(((type_(X45)|~secondordertype(X45))&(~iof(X46,X45)|firstordertype(X46)|~secondordertype(X45)))&((iof(esk10_1(X47),X47)|~type_(X47)|secondordertype(X47))&(~firstordertype(esk10_1(X47))|~type_(X47)|secondordertype(X47)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a6_secondOrderTypeDef])])])])])])).
fof(c_0_8, plain, ![X16, X18, X19]:((~type_(X16)|iof(esk2_1(X16),X16))&(~iof(X19,X18)|type_(X18))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a3])])])])])).
fof(c_0_9, plain, ![X12, X13, X14]:((~individual(X12)|~iof(X13,X12))&(iof(esk1_1(X14),X14)|individual(X14))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a2])])])])])).
fof(c_0_10, plain, ![X41, X42, X43]:(((type_(X41)|~firstordertype(X41))&(~iof(X42,X41)|individual(X42)|~firstordertype(X41)))&((iof(esk9_1(X43),X43)|~type_(X43)|firstordertype(X43))&(~individual(esk9_1(X43))|~type_(X43)|firstordertype(X43)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a5_firstOrderTypeDef])])])])])])).
cnf(c_0_11, plain, (firstordertype(X1)|~iof(X1,X2)|~secondordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (iof(esk2_1(X1),X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (type_(X1)|~secondordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14, plain, (~individual(X1)|~iof(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (type_(X1)|~firstordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (firstordertype(esk2_1(X1))|~secondordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])).
fof(c_0_17, negated_conjecture, ~(![X1, X2]:(isPowertypeOf(X1,X2)=>~(firstordertype(X1)))), inference(assume_negation,[status(cth)],[powertypeNotFirstOrder])).
cnf(c_0_18, plain, (~type_(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_14, c_0_12])).
cnf(c_0_19, plain, (type_(esk2_1(X1))|~secondordertype(X1)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20, plain, (individual(X1)|~iof(X1,X2)|~firstordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_21, plain, ![X62, X63, X64, X65, X66, X67]:(((type_(X62)|~isPowertypeOf(X62,X63))&((~iof(X64,X62)|specializes(X64,X63)|~isPowertypeOf(X62,X63))&(~specializes(X65,X63)|iof(X65,X62)|~isPowertypeOf(X62,X63))))&((~iof(esk13_2(X66,X67),X66)|~specializes(esk13_2(X66,X67),X67)|~type_(X66)|isPowertypeOf(X66,X67))&(iof(esk13_2(X66,X67),X66)|specializes(esk13_2(X66,X67),X67)|~type_(X66)|isPowertypeOf(X66,X67)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a10_isPowertypeOfDef])])])])])])).
fof(c_0_22, negated_conjecture, (isPowertypeOf(esk22_0,esk23_0)&firstordertype(esk22_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])).
cnf(c_0_23, plain, (~secondordertype(X1)|~individual(esk2_1(X1))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24, plain, (individual(esk2_1(X1))|~firstordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_12]), c_0_15])).
fof(c_0_25, plain, ![X49, X50, X51, X52, X53]:(((type_(X49)|~specializes(X49,X50))&(~iof(X51,X49)|iof(X51,X50)|~specializes(X49,X50)))&((iof(esk11_2(X52,X53),X52)|~type_(X52)|specializes(X52,X53))&(~iof(esk11_2(X52,X53),X53)|~type_(X52)|specializes(X52,X53)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a7_specDef])])])])])])).
cnf(c_0_26, plain, (specializes(X1,X3)|~iof(X1,X2)|~isPowertypeOf(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, negated_conjecture, (isPowertypeOf(esk22_0,esk23_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, negated_conjecture, (firstordertype(esk22_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (~secondordertype(X1)|~firstordertype(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_30, plain, (type_(X1)|~specializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, negated_conjecture, (specializes(X1,esk23_0)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32, plain, (iof(esk10_1(X1),X1)|secondordertype(X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_33, negated_conjecture, (type_(esk22_0)), inference(spm,[status(thm)],[c_0_15, c_0_28])).
cnf(c_0_34, negated_conjecture, (~secondordertype(esk22_0)), inference(spm,[status(thm)],[c_0_29, c_0_28])).
cnf(c_0_35, negated_conjecture, (type_(X1)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_36, negated_conjecture, (iof(esk10_1(esk22_0),esk22_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
cnf(c_0_37, negated_conjecture, (type_(esk10_1(esk22_0))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_38, negated_conjecture, (individual(esk10_1(esk22_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_36]), c_0_28])])).
cnf(c_0_39, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_37]), c_0_38])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 40
# Proof object clause steps            : 25
# Proof object formula steps           : 15
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 76
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 71
# Processed clauses                    : 147
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 144
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 62
# ...of the previous two non-trivial   : 48
# Contextual simplify-reflections      : 16
# Paramodulations                      : 60
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 71
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 61
# Current number of unprocessed clauses: 39
# ...number of literals in the above   : 124
# Current number of archived formulas  : 0
# Current number of archived clauses   : 73
# Clause-clause subsumption calls (NU) : 2174
# Rec. Clause-clause subsumption calls : 1467
# Non-unit clause-clause subsumptions  : 17
# Unit Clause-clause subsumption calls : 23
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5994

# -------------------------------------------------
# User time                : 0.031 s
# System time              : 0.006 s
# Total time               : 0.038 s
# Maximum resident set size: 1604 pages
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_M7dr
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_M7dr
% SZS output start Proof for SOT_M7dr
fof(f275,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f104,f108,f116,f124,f140,f149,f157,f163,f189,f196,f205,f230,f254,f260,f267,f274])).
fof(f274,plain,(
  ~spl11_14 | ~spl11_23 | ~spl11_36),
  inference(avatar_contradiction_clause,[],[f273])).
fof(f273,plain,(
  $false | (~spl11_14 | ~spl11_23 | ~spl11_36)),
  inference(subsumption_resolution,[],[f270,f204])).
fof(f204,plain,(
  ( ! [X0] : (~iof(X0,sK6(sK0))) ) | ~spl11_23),
  inference(avatar_component_clause,[],[f203])).
fof(f203,plain,(
  spl11_23 <=> ! [X0] : ~iof(X0,sK6(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_23])])).
fof(f270,plain,(
  iof(sK6(sK6(sK0)),sK6(sK0)) | (~spl11_14 | ~spl11_36)),
  inference(resolution,[],[f266,f156])).
fof(f156,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK6(X0),X0)) ) | ~spl11_14),
  inference(avatar_component_clause,[],[f155])).
fof(f155,plain,(
  spl11_14 <=> ! [X0] : (iof(sK6(X0),X0) | ~type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_14])])).
fof(f266,plain,(
  type_(sK6(sK0)) | ~spl11_36),
  inference(avatar_component_clause,[],[f265])).
fof(f265,plain,(
  spl11_36 <=> type_(sK6(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_36])])).
fof(f267,plain,(
  spl11_36 | ~spl11_10 | ~spl11_35),
  inference(avatar_split_clause,[],[f263,f258,f138,f265])).
fof(f138,plain,(
  spl11_10 <=> ! [X1,X0] : (type_(X0) | ~specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_10])])).
fof(f258,plain,(
  spl11_35 <=> specializes(sK6(sK0),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_35])])).
fof(f263,plain,(
  type_(sK6(sK0)) | (~spl11_10 | ~spl11_35)),
  inference(resolution,[],[f259,f139])).
fof(f139,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) ) | ~spl11_10),
  inference(avatar_component_clause,[],[f138])).
fof(f259,plain,(
  specializes(sK6(sK0),sK1) | ~spl11_35),
  inference(avatar_component_clause,[],[f258])).
fof(f260,plain,(
  spl11_35 | ~spl11_20 | ~spl11_34),
  inference(avatar_split_clause,[],[f255,f252,f187,f258])).
fof(f187,plain,(
  spl11_20 <=> iof(sK6(sK0),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_20])])).
fof(f252,plain,(
  spl11_34 <=> ! [X0] : (specializes(X0,sK1) | ~iof(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_34])])).
fof(f255,plain,(
  specializes(sK6(sK0),sK1) | (~spl11_20 | ~spl11_34)),
  inference(resolution,[],[f253,f188])).
fof(f188,plain,(
  iof(sK6(sK0),sK0) | ~spl11_20),
  inference(avatar_component_clause,[],[f187])).
fof(f253,plain,(
  ( ! [X0] : (~iof(X0,sK0) | specializes(X0,sK1)) ) | ~spl11_34),
  inference(avatar_component_clause,[],[f252])).
fof(f254,plain,(
  spl11_34 | ~spl11_2 | ~spl11_29),
  inference(avatar_split_clause,[],[f250,f228,f106,f252])).
fof(f106,plain,(
  spl11_2 <=> isPowertypeOf(sK0,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_2])])).
fof(f228,plain,(
  spl11_29 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_29])])).
fof(f250,plain,(
  ( ! [X0] : (specializes(X0,sK1) | ~iof(X0,sK0)) ) | (~spl11_2 | ~spl11_29)),
  inference(resolution,[],[f229,f107])).
fof(f107,plain,(
  isPowertypeOf(sK0,sK1) | ~spl11_2),
  inference(avatar_component_clause,[],[f106])).
fof(f229,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) ) | ~spl11_29),
  inference(avatar_component_clause,[],[f228])).
fof(f230,plain,(
  spl11_29),
  inference(avatar_split_clause,[],[f90,f228])).
fof(f90,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : ((! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)) | ~isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) => (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) <=> (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : (isPowertypeOf(X6,X7) <=> (! [X8] : (iof(X8,X6) <=> specializes(X8,X7)) & type_(X6)))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a10_isPowertypeOfDef)).
fof(f205,plain,(
  spl11_23 | ~spl11_6 | ~spl11_21),
  inference(avatar_split_clause,[],[f197,f194,f122,f203])).
fof(f122,plain,(
  spl11_6 <=> ! [X1,X0] : (~individual(X0) | ~iof(X1,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_6])])).
fof(f194,plain,(
  spl11_21 <=> individual(sK6(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_21])])).
fof(f197,plain,(
  ( ! [X0] : (~iof(X0,sK6(sK0))) ) | (~spl11_6 | ~spl11_21)),
  inference(resolution,[],[f195,f123])).
fof(f123,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) ) | ~spl11_6),
  inference(avatar_component_clause,[],[f122])).
fof(f195,plain,(
  individual(sK6(sK0)) | ~spl11_21),
  inference(avatar_component_clause,[],[f194])).
fof(f196,plain,(
  spl11_21 | ~spl11_1 | ~spl11_15 | ~spl11_20),
  inference(avatar_split_clause,[],[f192,f187,f161,f102,f194])).
fof(f102,plain,(
  spl11_1 <=> firstordertype(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_1])])).
fof(f161,plain,(
  spl11_15 <=> ! [X1,X0] : (~firstordertype(X0) | ~iof(X1,X0) | individual(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_15])])).
fof(f192,plain,(
  individual(sK6(sK0)) | (~spl11_1 | ~spl11_15 | ~spl11_20)),
  inference(subsumption_resolution,[],[f190,f103])).
fof(f103,plain,(
  firstordertype(sK0) | ~spl11_1),
  inference(avatar_component_clause,[],[f102])).
fof(f190,plain,(
  ~firstordertype(sK0) | individual(sK6(sK0)) | (~spl11_15 | ~spl11_20)),
  inference(resolution,[],[f188,f162])).
fof(f162,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | ~firstordertype(X0) | individual(X1)) ) | ~spl11_15),
  inference(avatar_component_clause,[],[f161])).
fof(f189,plain,(
  spl11_20 | ~spl11_12 | ~spl11_14),
  inference(avatar_split_clause,[],[f169,f155,f147,f187])).
fof(f147,plain,(
  spl11_12 <=> type_(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_12])])).
fof(f169,plain,(
  iof(sK6(sK0),sK0) | (~spl11_12 | ~spl11_14)),
  inference(resolution,[],[f156,f148])).
fof(f148,plain,(
  type_(sK0) | ~spl11_12),
  inference(avatar_component_clause,[],[f147])).
fof(f163,plain,(
  spl11_15),
  inference(avatar_split_clause,[],[f66,f161])).
fof(f66,plain,(
  ( ! [X0,X1] : (~firstordertype(X0) | ~iof(X1,X0) | individual(X1)) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : ((! [X1] : (individual(X1) | ~iof(X1,X0)) & type_(X0)) | ~firstordertype(X0))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : (firstordertype(X0) => (! [X1] : (iof(X1,X0) => individual(X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f10])).
fof(f10,axiom,(
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (iof(X1,X0) => individual(X1)) & type_(X0)))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a5_firstOrderTypeDef)).
fof(f157,plain,(
  spl11_14),
  inference(avatar_split_clause,[],[f70,f155])).
fof(f70,plain,(
  ( ! [X0] : (iof(sK6(X0),X0) | ~type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a3)).
fof(f149,plain,(
  spl11_12 | ~spl11_1 | ~spl11_4),
  inference(avatar_split_clause,[],[f145,f114,f102,f147])).
fof(f114,plain,(
  spl11_4 <=> ! [X0] : (~firstordertype(X0) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_4])])).
fof(f145,plain,(
  type_(sK0) | (~spl11_1 | ~spl11_4)),
  inference(resolution,[],[f115,f103])).
fof(f115,plain,(
  ( ! [X0] : (~firstordertype(X0) | type_(X0)) ) | ~spl11_4),
  inference(avatar_component_clause,[],[f114])).
fof(f140,plain,(
  spl11_10),
  inference(avatar_split_clause,[],[f88,f138])).
fof(f88,plain,(
  ( ! [X0,X1] : (type_(X0) | ~specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X1) | ~iof(X2,X0)) & type_(X0)))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X0) => iof(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7] : (specializes(X6,X7) <=> (! [X3] : (iof(X3,X6) => iof(X3,X7)) & type_(X6)))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a7_specDef)).
fof(f124,plain,(
  spl11_6),
  inference(avatar_split_clause,[],[f68,f122])).
fof(f68,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (! [X1] : ~iof(X1,X0) | ~individual(X0))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : (individual(X0) => ~? [X1] : iof(X1,X0))),
  inference(unused_predicate_definition_removal,[],[f2])).
fof(f2,axiom,(
  ! [X0] : (individual(X0) <=> ~? [X1] : iof(X1,X0))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',a2)).
fof(f116,plain,(
  spl11_4),
  inference(avatar_split_clause,[],[f67,f114])).
fof(f67,plain,(
  ( ! [X0] : (~firstordertype(X0) | type_(X0)) )),
  inference(cnf_transformation,[],[f43])).
fof(f108,plain,(
  spl11_2),
  inference(avatar_split_clause,[],[f54,f106])).
fof(f54,plain,(
  isPowertypeOf(sK0,sK1)),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ? [X0,X1] : (firstordertype(X0) & isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f22])).
fof(f22,negated_conjecture,(
  ~! [X0,X1] : (isPowertypeOf(X0,X1) => ~firstordertype(X0))),
  inference(negated_conjecture,[],[f21])).
fof(f21,conjecture,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) => ~firstordertype(X0))),
  file('/tmp/i1lpDJS7Ag/SOT_M7drGb',powertypeNotFirstOrder)).
fof(f104,plain,(
  spl11_1),
  inference(avatar_split_clause,[],[f55,f102])).
fof(f55,plain,(
  firstordertype(sK0)),
  inference(cnf_transformation,[],[f38])).
% SZS output end Proof for SOT_M7dr
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.011 s
% ------------------------------
% ------------------------------
% Success in time 0.077 s

% END OF SYSTEM OUTPUT

% END OF SYSTEM OUTPUT
% RESULT: SOT_M7drGb - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.04 
% OUTPUT: SOT_M7drGb - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.05 
% RESULT: SOT_M7drGb - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.27 
% OUTPUT: SOT_M7drGb - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.27 

% END OF SYSTEM OUTPUT
% RESULT: SOT_M7drGb - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.07 
% OUTPUT: SOT_M7drGb - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.07 

% END OF SYSTEM OUTPUT
% RESULT: SOT_M7drGb - E---FNT-2.5 says Theorem - CPU = 0.00 WC = 0.04 
% OUTPUT: SOT_M7drGb - E---FNT-2.5 says CNFRefutation - CPU = 0.00 WC = 0.04 

% END OF SYSTEM OUTPUT

% END OF SYSTEM OUTPUT
% RESULT: SOT_M7drGb - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.08 
% OUTPUT: SOT_M7drGb - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.09 
% RESULT: SOT_M7drGb - E---FNT-2.5 says Theorem - CPU = 0.00 WC = 0.02 
% OUTPUT: SOT_M7drGb - E---FNT-2.5 says CNFRefutation - CPU = 0.00 WC = 0.02 
Vampire---SAT-4.5 has ended (PID 97253)
Vampire---SAT-4.5 has ended (PID 97254)
Vampire---SAT-4.5 has ended (PID 97256)
Vampire---SAT-4.5 has ended (PID 97258)
E---FNT-2.5 has ended (PID 97261)
E---FNT-2.5 has ended (PID 97265)
RESULT: SOT_M7drGb - SSCPA---0.0 says Theorem - CPU = 3.5 WC = 6.7  SolvingSystem = Vampire---SAT-4.5
</PRE>

</body>
</html>