<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>System Results</title>
<meta name="ROBOTS" content="INDEX,NOFOLLOW" />
<script type="text/javascript">//<![CDATA[

  var tstp_dump;
  function openSoTSTP (dump) {   
    var tstp_url = 'http://www.tptp.org/cgi-bin/SystemOnTSTP';
    var tstp_browser = window.open(tstp_url, '_blank');
    tstp_dump = dump;
  }
  function getTSTPDump () {
    return tstp_dump;
  }

//]]></script>
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>

% SZS start RequiredInformation<BR>
% Congratulations - you have become a registered power user of SystemOnTPTP, at IP address 191.6.38.118.<BR>
% Please consider donating to the TPTP project - see www.tptp.org for details.<BR>
% When you donate this message will disappear.<BR>
% If you do not donate a random delay might be added to your processing time.<BR>
% SZS end RequiredInformation<BR>
<PRE>
Getting system recommendations for generic problem /tmp/WT2askIkzH/SOT_qBPqo3
This is a FOF_THM_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.95
Vampire---TF1-4.5    recommended  0.95
Vampire---4.5        recommended  0.95
Vampire---SAT-4.5    recommended  0.95
ET---2.0             recommended  0.80
E---2.5              recommended  0.78
E---FNT-2.5          recommended  0.78
CSE_E---1.0          recommended  0.76
E---FNT-2.5          recommended  0.74
Isabelle---2TH0      recommended  0.69
Isabelle---2TF0      recommended  0.69
Isabelle---2FOF      recommended  0.69
Isabelle---2016      recommended  0.69
Isabelle-HOT---2016  recommended  0.68
CVC4---1.7           recommended  0.68
CVC4---SAT-1.7       recommended  0.68
Leo-III---1.4        recommended  0.61
Vampire---THF-4.5    recommended  0.59
Vampire---TF1-4.5    recommended  0.59
Vampire---SAT-4.5    recommended  0.59
iProver---3.0        recommended  0.58
iProver---SAT-3.0    recommended  0.58
Fampire---1.3        recommended  0.57
SInE---0.4           recommended  0.53
iProver---SAT-3.0    recommended  0.52
ePrincess---1.0      recommended  0.50
Princess---170717    recommended  0.48
SPASS---3.9          recommended  0.48
Z3---4.4.1           recommended  0.46
Zipperpin---1.5      recommended  0.40
CVC4---SAT-1.7       recommended  0.40
Prover9---1109a      recommended  0.39
Equinox---6.0.1a     recommended  0.37
LEO-II---1.7.0       recommended  0.35
nanoCoP---1.1        recommended  0.34
Muscadet---4.5       recommended  0.25
Twee---2.2           recommended  0.23
Metis---2.4          recommended  0.23
Geo-III---2018C      recommended  0.18
iProverMo---2.5-0.1  recommended  0.14
Beagle---0.9.47      subsumed     0.41
leanCoP---2.2        subsumed     0.36
E-Darwin---1.5       subsumed     0.31
CSE---1.0            subsumed     0.28
Darwin---1.4.5       subsumed     0.28
Bliksem---1.12       subsumed     0.21
Zenon---0.7.1        subsumed     0.16
ZenonModulo---0.4.2  subsumed     0.15
SOS---2.0            subsumed     0.11
Otter---3.3          subsumed     0.04
This is a FOF_CSA_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.98
Vampire---TF1-4.5    recommended  0.98
Vampire---SAT-4.5    recommended  0.98
Paradox---4.0        recommended  0.60
Nitpick---2016       subsumed     0.55
DarwinFM---1.4.5     subsumed     0.43
CVC4---SAT-1.7       subsumed     0.40
Vampire---THF-4.5    subsumed     0.34
Vampire---TF1-4.5    subsumed     0.34
Vampire---4.5        subsumed     0.34
Vampire---SAT-4.5    subsumed     0.34
E---2.5              subsumed     0.34
E---FNT-2.5          subsumed     0.34
E---FNT-2.5          subsumed     0.34
iProver---SAT-3.0    subsumed     0.32
Metis---2.4          subsumed     0.32
SPASS---3.9          subsumed     0.28
Beagle---0.9.47      subsumed     0.26
Geo-III---2018C      subsumed     0.25
Z3---4.4.1           subsumed     0.25
Princess---170717    subsumed     0.25
Equinox---6.0.1a     subsumed     0.19
iProver---3.0        subsumed     0.11
iProver---SAT-3.0    subsumed     0.11
Darwin---1.4.5       subsumed     0.06
E-Darwin---1.5       subsumed     0.04
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use ET,2.0
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use CSE_E,1.0
CSE_E---1.0 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use Isabelle,2TH0
Isabelle---2016 can be used
Check if we can use Isabelle,2TF0
Isabelle---2016 can be used
Check if we can use Isabelle,2FOF
Isabelle---2016 can be used
Check if we can use Isabelle,2016
Isabelle---2016 can be used
Check if we can use Isabelle-HOT,2016
Isabelle-HOT---2016 can be used
Check if we can use CVC4,1.7
CVC4---SAT-1.7 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Leo-III,1.4
Leo-III---1.4 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Fampire,1.3
Check if we can use SInE,0.4
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use ePrincess,1.0
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Zipperpin,1.5
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Prover9,1109a
Prover9---1109a can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use LEO-II,1.7.0
LEO-II---1.7.0 can be used
Check if we can use nanoCoP,1.1
nanoCoP---1.1 can be used
Check if we can use Muscadet,4.5
Muscadet---4.5 can be used
Check if we can use Twee,2.2
Twee---2.2 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use iProverMo,2.5-0.1
iProverMo---2.5-0.1 can be used
Check if we can use Beagle,0.9.47
Check if we can use leanCoP,2.2
leanCoP---2.2 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Check if we can use CSE,1.0
CSE---1.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use Bliksem,1.12
Bliksem---1.12 can be used
Check if we can use Zenon,0.7.1
Zenon---0.7.1 can be used
Check if we can use ZenonModulo,0.4.2
Check if we can use SOS,2.0
SOS---2.0 can be used
Check if we can use Otter,3.3
Otter---3.3 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use Paradox,4.0
Paradox---4.0 can be used
Check if we can use Nitpick,2016
Nitpick---2016 can be used
Check if we can use DarwinFM,1.4.5
DarwinFM---1.4.5 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Beagle,0.9.47
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
E---FNT-2.5 will be used
E---FNT-2.5 will be used
CSE_E---1.0 will be used
E---FNT-2.5 will be used
Isabelle---2016 will be used
Isabelle---2016 will be used
Vampire---SAT-4.5 is given 30 seconds in -cS mode
Vampire---SAT-4.5 is given 30 seconds in -cS mode
Vampire---SAT-4.5 is given 30 seconds in -cS mode
Vampire---SAT-4.5 is given 30 seconds in -cS mode
E---FNT-2.5 is given 30 seconds in -cS mode
E---FNT-2.5 is given 30 seconds in -cS mode
CSE_E---1.0 is given 30 seconds in -cS mode
E---FNT-2.5 is given 30 seconds in -cS mode
Isabelle---2016 is given 30 seconds in -cS mode
Isabelle---2016 is given 30 seconds in -cS mode
Vampire---SAT-4.5 has started (PID 105525)
Vampire---SAT-4.5 has started (PID 105526)
Vampire---SAT-4.5 has started (PID 105527)
Vampire---SAT-4.5 has started (PID 105528)
E---FNT-2.5 has started (PID 105529)
E---FNT-2.5 has started (PID 105530)
CSE_E---1.0 has started (PID 105531)
E---FNT-2.5 has started (PID 105532)
Isabelle---2016 has started (PID 105533)
Isabelle---2016 has started (PID 105534)
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  30  /tmp/WT2askIkzH/SOT_qBPqo3
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  30  /tmp/WT2askIkzH/SOT_qBPqo3
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  30  /tmp/WT2askIkzH/SOT_qBPqo3
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  30  /tmp/WT2askIkzH/SOT_qBPqo3
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  30  /tmp/WT2askIkzH/SOT_qBPqo3
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  30  /tmp/WT2askIkzH/SOT_qBPqo3
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 CSE_E---1.0  30  /tmp/WT2askIkzH/SOT_qBPqo3
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  30  /tmp/WT2askIkzH/SOT_qBPqo3
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Isabelle---2016  30  /tmp/WT2askIkzH/SOT_qBPqo3
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Isabelle---2016  30  /tmp/WT2askIkzH/SOT_qBPqo3
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 14 seconds
% START OF SYSTEM OUTPUT
% START OF SYSTEM OUTPUT
% START OF SYSTEM OUTPUT
% START OF SYSTEM OUTPUT
# AutoSched0-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S5PRR_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.017 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(a2, axiom, ![X1]:(individual(X1)<=>~(?[X2]:iof(X2,X1))), file('/tmp/WT2askIkzH/SOT_qBPqo3', a2)).
fof(a3, axiom, ![X1]:(type_(X1)<=>?[X2]:iof(X2,X1)), file('/tmp/WT2askIkzH/SOT_qBPqo3', a3)).
fof(a6_secondOrderTypeDef, axiom, ![X1]:(secondordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>firstordertype(X2)))), file('/tmp/WT2askIkzH/SOT_qBPqo3', a6_secondOrderTypeDef)).
fof(a5_firstOrderTypeDef, axiom, ![X1]:(firstordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>individual(X2)))), file('/tmp/WT2askIkzH/SOT_qBPqo3', a5_firstOrderTypeDef)).
fof(a15_sotConstantDef, axiom, ![X11]:(X11=c_sot<=>![X1]:(iof(X1,X11)<=>secondordertype(X1))), file('/tmp/WT2askIkzH/SOT_qBPqo3', a15_sotConstantDef)).
fof(t8_ap2, conjecture, ~(?[X7, X8]:((iof(X7,c_sot)&iof(X8,c_sot))&iof(X8,X7))), file('/tmp/WT2askIkzH/SOT_qBPqo3', t8_ap2)).
fof(c_0_6, plain, ![X13, X14, X15]:((~individual(X13)|~iof(X14,X13))&(iof(esk1_1(X15),X15)|individual(X15))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a2])])])])])).
fof(c_0_7, plain, ![X17, X19, X20]:((~type_(X17)|iof(esk2_1(X17),X17))&(~iof(X20,X19)|type_(X19))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a3])])])])])).
cnf(c_0_8, plain, (~individual(X1)|~iof(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (iof(esk2_1(X1),X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, plain, ![X46, X47, X48]:(((type_(X46)|~secondordertype(X46))&(~iof(X47,X46)|firstordertype(X47)|~secondordertype(X46)))&((iof(esk10_1(X48),X48)|~type_(X48)|secondordertype(X48))&(~firstordertype(esk10_1(X48))|~type_(X48)|secondordertype(X48)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a6_secondOrderTypeDef])])])])])])).
fof(c_0_11, plain, ![X42, X43, X44]:(((type_(X42)|~firstordertype(X42))&(~iof(X43,X42)|individual(X43)|~firstordertype(X42)))&((iof(esk9_1(X44),X44)|~type_(X44)|firstordertype(X44))&(~individual(esk9_1(X44))|~type_(X44)|firstordertype(X44)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a5_firstOrderTypeDef])])])])])])).
cnf(c_0_12, plain, (~type_(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_13, plain, (type_(X1)|~secondordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, plain, ![X114, X115, X116, X117]:(((~iof(X115,X114)|secondordertype(X115)|X114!=c_sot)&(~secondordertype(X116)|iof(X116,X114)|X114!=c_sot))&((~iof(esk24_1(X117),X117)|~secondordertype(esk24_1(X117))|X117=c_sot)&(iof(esk24_1(X117),X117)|secondordertype(esk24_1(X117))|X117=c_sot))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a15_sotConstantDef])])])])])])).
fof(c_0_15, negated_conjecture, ~(~(?[X7, X8]:((iof(X7,c_sot)&iof(X8,c_sot))&iof(X8,X7)))), inference(assume_negation,[status(cth)],[t8_ap2])).
cnf(c_0_16, plain, (type_(X1)|~firstordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, plain, (firstordertype(X1)|~iof(X1,X2)|~secondordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, plain, (iof(esk1_1(X1),X1)|individual(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_19, plain, (~secondordertype(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_20, plain, (secondordertype(X1)|~iof(X1,X2)|X2!=c_sot), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_21, negated_conjecture, ((iof(esk25_0,c_sot)&iof(esk26_0,c_sot))&iof(esk26_0,esk25_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
cnf(c_0_22, plain, (~firstordertype(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_12, c_0_16])).
cnf(c_0_23, plain, (firstordertype(esk1_1(X1))|~secondordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])).
cnf(c_0_24, plain, (individual(X1)|~iof(X1,X2)|~firstordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_25, plain, (secondordertype(X1)|~iof(X1,c_sot)), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_26, negated_conjecture, (iof(esk25_0,c_sot)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (~secondordertype(X1)|~individual(esk1_1(X1))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_28, plain, (individual(esk1_1(X1))|~firstordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_18]), c_0_22])).
cnf(c_0_29, negated_conjecture, (iof(esk26_0,c_sot)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, negated_conjecture, (iof(esk26_0,esk25_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, negated_conjecture, (secondordertype(esk25_0)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_32, plain, (~secondordertype(X1)|~firstordertype(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33, negated_conjecture, (secondordertype(esk26_0)), inference(spm,[status(thm)],[c_0_25, c_0_29])).
cnf(c_0_34, negated_conjecture, (firstordertype(esk26_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_30]), c_0_31])])).
cnf(c_0_35, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 36
# Proof object clause steps            : 23
# Proof object formula steps           : 13
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 24
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 89
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 84
# Processed clauses                    : 153
# ...of these trivial                  : 2
# ...subsumed                          : 16
# ...remaining for further processing  : 135
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 3
# Generated clauses                    : 256
# ...of the previous two non-trivial   : 223
# Contextual simplify-reflections      : 15
# Paramodulations                      : 249
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 124
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 103
# Current number of unprocessed clauses: 149
# ...number of literals in the above   : 523
# Current number of archived formulas  : 0
# Current number of archived clauses   : 5
# Clause-clause subsumption calls (NU) : 3206
# Rec. Clause-clause subsumption calls : 2282
# Non-unit clause-clause subsumptions  : 22
# Unit Clause-clause subsumption calls : 565
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8115

# -------------------------------------------------
# User time                : 0.023 s
# System time              : 0.004 s
# Total time               : 0.027 s
# Maximum resident set size: 1612 pages
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_qBPq
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_qBPq
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_qBPq
% START OF SYSTEM OUTPUT
% START OF SYSTEM OUTPUT
Moving to /home/tptp/Systems/CSE_E---1.0
java -jar /home/tptp/Systems/CSE_E---1.0/mcs_epr.jar 30 /tmp/WT2askIkzH/SOT_qBPqo3
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_qBPq
% SZS output start Proof for SOT_qBPq
fof(f1121,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f126,f130,f134,f142,f146,f154,f186,f202,f243,f249,f259,f268,f373,f455,f803,f820,f880,f890,f901,f1102,f1116,f1120])).
fof(f1120,plain,(
  ~spl16_8 | ~spl16_20 | ~spl16_100 | ~spl16_123),
  inference(avatar_contradiction_clause,[],[f1119])).
fof(f1119,plain,(
  $false | (~spl16_8 | ~spl16_20 | ~spl16_100 | ~spl16_123)),
  inference(subsumption_resolution,[],[f904,f1117])).
fof(f1117,plain,(
  ( ! [X0] : (~iof(X0,sK9(sK1))) ) | (~spl16_8 | ~spl16_123)),
  inference(resolution,[],[f1115,f153])).
fof(f153,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) ) | ~spl16_8),
  inference(avatar_component_clause,[],[f152])).
fof(f152,plain,(
  spl16_8 <=> ! [X1,X0] : (~iof(X1,X0) | ~individual(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_8])])).
fof(f1115,plain,(
  individual(sK9(sK1)) | ~spl16_123),
  inference(avatar_component_clause,[],[f1114])).
fof(f1114,plain,(
  spl16_123 <=> individual(sK9(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_123])])).
fof(f904,plain,(
  iof(sK9(sK9(sK1)),sK9(sK1)) | (~spl16_20 | ~spl16_100)),
  inference(resolution,[],[f900,f201])).
fof(f201,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK9(X0),X0)) ) | ~spl16_20),
  inference(avatar_component_clause,[],[f200])).
fof(f200,plain,(
  spl16_20 <=> ! [X0] : (iof(sK9(X0),X0) | ~type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_20])])).
fof(f900,plain,(
  type_(sK9(sK1)) | ~spl16_100),
  inference(avatar_component_clause,[],[f899])).
fof(f899,plain,(
  spl16_100 <=> type_(sK9(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_100])])).
fof(f1116,plain,(
  spl16_123 | ~spl16_54 | ~spl16_121),
  inference(avatar_split_clause,[],[f1107,f1100,f453,f1114])).
fof(f453,plain,(
  spl16_54 <=> iof(sK9(sK1),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_54])])).
fof(f1100,plain,(
  spl16_121 <=> ! [X0] : (individual(X0) | ~iof(X0,sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_121])])).
fof(f1107,plain,(
  individual(sK9(sK1)) | (~spl16_54 | ~spl16_121)),
  inference(resolution,[],[f1101,f454])).
fof(f454,plain,(
  iof(sK9(sK1),sK1) | ~spl16_54),
  inference(avatar_component_clause,[],[f453])).
fof(f1101,plain,(
  ( ! [X0] : (~iof(X0,sK1) | individual(X0)) ) | ~spl16_121),
  inference(avatar_component_clause,[],[f1100])).
fof(f1102,plain,(
  spl16_121 | ~spl16_40 | ~spl16_85),
  inference(avatar_split_clause,[],[f821,f818,f371,f1100])).
fof(f371,plain,(
  spl16_40 <=> ! [X1,X0] : (~iof(X1,X0) | individual(X1) | ~firstordertype(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_40])])).
fof(f818,plain,(
  spl16_85 <=> firstordertype(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_85])])).
fof(f821,plain,(
  ( ! [X0] : (individual(X0) | ~iof(X0,sK1)) ) | (~spl16_40 | ~spl16_85)),
  inference(resolution,[],[f819,f372])).
fof(f372,plain,(
  ( ! [X0,X1] : (~firstordertype(X0) | individual(X1) | ~iof(X1,X0)) ) | ~spl16_40),
  inference(avatar_component_clause,[],[f371])).
fof(f819,plain,(
  firstordertype(sK1) | ~spl16_85),
  inference(avatar_component_clause,[],[f818])).
fof(f901,plain,(
  spl16_100 | ~spl16_6 | ~spl16_98),
  inference(avatar_split_clause,[],[f893,f888,f144,f899])).
fof(f144,plain,(
  spl16_6 <=> ! [X0] : (type_(X0) | ~firstordertype(X0))),
% Refutation found. Thanks to Tanya!
  introduced(avatar_definition,[new_symbols(naming,[spl16_6])])).
fof(f888,plain,(
  spl16_98 <=> firstordertype(sK9(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_98])])).
% SZS status Theorem for SOT_qBPq
fof(f893,plain,(
  type_(sK9(sK1)) | (~spl16_6 | ~spl16_98)),
  inference(resolution,[],[f889,f145])).
% SZS output start Proof for SOT_qBPq
fof(f145,plain,(
  ( ! [X0] : (~firstordertype(X0) | type_(X0)) ) | ~spl16_6),
  inference(avatar_component_clause,[],[f144])).
fof(f889,plain,(
fof(f1121,plain,(
  firstordertype(sK9(sK1)) | ~spl16_98),
  $false),
  inference(avatar_component_clause,[],[f888])).
fof(f890,plain,(
  inference(avatar_sat_refutation,[],[f126,f130,f134,f142,f146,f154,f186,f202,f243,f249,f259,f268,f373,f455,f803,f820,f880,f890,f901,f1102,f1116,f1120])).
  spl16_98 | ~spl16_54 | ~spl16_96),
fof(f1120,plain,(
  inference(avatar_split_clause,[],[f885,f878,f453,f888])).
  ~spl16_8 | ~spl16_20 | ~spl16_100 | ~spl16_123),
fof(f878,plain,(
  inference(avatar_contradiction_clause,[],[f1119])).
  spl16_96 <=> ! [X1] : (firstordertype(X1) | ~iof(X1,sK1))),
fof(f1119,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl16_96])])).
  $false | (~spl16_8 | ~spl16_20 | ~spl16_100 | ~spl16_123)),
fof(f885,plain,(
  inference(subsumption_resolution,[],[f904,f1117])).
  firstordertype(sK9(sK1)) | (~spl16_54 | ~spl16_96)),
fof(f1117,plain,(
  inference(resolution,[],[f879,f454])).
  ( ! [X0] : (~iof(X0,sK9(sK1))) ) | (~spl16_8 | ~spl16_123)),
fof(f879,plain,(
  inference(resolution,[],[f1115,f153])).
  ( ! [X1] : (~iof(X1,sK1) | firstordertype(X1)) ) | ~spl16_96),
fof(f153,plain,(
  inference(avatar_component_clause,[],[f878])).
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) ) | ~spl16_8),
fof(f880,plain,(
  inference(avatar_component_clause,[],[f152])).
  spl16_96 | ~spl16_28 | ~spl16_32),
fof(f152,plain,(
  inference(avatar_split_clause,[],[f462,f266,f247,f878])).
  spl16_8 <=> ! [X1,X0] : (~iof(X1,X0) | ~individual(X0))),
fof(f247,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl16_8])])).
  spl16_28 <=> secondordertype(sK1)),
fof(f1115,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl16_28])])).
  individual(sK9(sK1)) | ~spl16_123),
fof(f266,plain,(
  inference(avatar_component_clause,[],[f1114])).
  spl16_32 <=> ! [X1,X0] : (~iof(X1,X0) | firstordertype(X1) | ~secondordertype(X0))),
fof(f1114,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl16_32])])).
  spl16_123 <=> individual(sK9(sK1))),
fof(f462,plain,(
  ( ! [X1] : (firstordertype(X1) | ~iof(X1,sK1)) ) | (~spl16_28 | ~spl16_32)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_123])])).
  inference(resolution,[],[f267,f248])).
fof(f904,plain,(
fof(f248,plain,(
  iof(sK9(sK9(sK1)),sK9(sK1)) | (~spl16_20 | ~spl16_100)),
  secondordertype(sK1) | ~spl16_28),
  inference(resolution,[],[f900,f201])).
  inference(avatar_component_clause,[],[f247])).
fof(f201,plain,(
fof(f267,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK9(X0),X0)) ) | ~spl16_20),
  ( ! [X0,X1] : (~secondordertype(X0) | firstordertype(X1) | ~iof(X1,X0)) ) | ~spl16_32),
  inference(avatar_component_clause,[],[f200])).
  inference(avatar_component_clause,[],[f266])).
fof(f200,plain,(
fof(f820,plain,(
  spl16_20 <=> ! [X0] : (iof(sK9(X0),X0) | ~type_(X0))),
  spl16_85 | ~spl16_3 | ~spl16_84),
  introduced(avatar_definition,[new_symbols(naming,[spl16_20])])).
  inference(avatar_split_clause,[],[f804,f801,f132,f818])).
fof(f900,plain,(
fof(f132,plain,(
  type_(sK9(sK1)) | ~spl16_100),
  spl16_3 <=> iof(sK1,sK0)),
  inference(avatar_component_clause,[],[f899])).
  introduced(avatar_definition,[new_symbols(naming,[spl16_3])])).
fof(f899,plain,(
fof(f801,plain,(
  spl16_100 <=> type_(sK9(sK1))),
  spl16_84 <=> ! [X0] : (firstordertype(X0) | ~iof(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_100])])).
  introduced(avatar_definition,[new_symbols(naming,[spl16_84])])).
fof(f1116,plain,(
fof(f804,plain,(
  spl16_123 | ~spl16_54 | ~spl16_121),
  firstordertype(sK1) | (~spl16_3 | ~spl16_84)),
  inference(avatar_split_clause,[],[f1107,f1100,f453,f1114])).
  inference(resolution,[],[f802,f133])).
fof(f453,plain,(
fof(f133,plain,(
  spl16_54 <=> iof(sK9(sK1),sK1)),
  iof(sK1,sK0) | ~spl16_3),
  introduced(avatar_definition,[new_symbols(naming,[spl16_54])])).
  inference(avatar_component_clause,[],[f132])).
fof(f1100,plain,(
fof(f802,plain,(
  spl16_121 <=> ! [X0] : (individual(X0) | ~iof(X0,sK1))),
  ( ! [X0] : (~iof(X0,sK0) | firstordertype(X0)) ) | ~spl16_84),
  introduced(avatar_definition,[new_symbols(naming,[spl16_121])])).
  inference(avatar_component_clause,[],[f801])).
fof(f1107,plain,(
fof(f803,plain,(
  individual(sK9(sK1)) | (~spl16_54 | ~spl16_121)),
  spl16_84 | ~spl16_27 | ~spl16_32),
  inference(resolution,[],[f1101,f454])).
  inference(avatar_split_clause,[],[f461,f266,f241,f801])).
fof(f454,plain,(
fof(f241,plain,(
  iof(sK9(sK1),sK1) | ~spl16_54),
  spl16_27 <=> secondordertype(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_27])])).
fof(f461,plain,(
  inference(avatar_component_clause,[],[f453])).
  ( ! [X0] : (firstordertype(X0) | ~iof(X0,sK0)) ) | (~spl16_27 | ~spl16_32)),
fof(f1101,plain,(
  inference(resolution,[],[f267,f242])).
  ( ! [X0] : (~iof(X0,sK1) | individual(X0)) ) | ~spl16_121),
fof(f242,plain,(
  inference(avatar_component_clause,[],[f1100])).
  secondordertype(sK0) | ~spl16_27),
fof(f1102,plain,(
  inference(avatar_component_clause,[],[f241])).
  spl16_121 | ~spl16_40 | ~spl16_85),
fof(f455,plain,(
  inference(avatar_split_clause,[],[f821,f818,f371,f1100])).
  spl16_54 | ~spl16_20 | ~spl16_30),
fof(f371,plain,(
  inference(avatar_split_clause,[],[f405,f257,f200,f453])).
  spl16_40 <=> ! [X1,X0] : (~iof(X1,X0) | individual(X1) | ~firstordertype(X0))),
fof(f257,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl16_40])])).
  spl16_30 <=> type_(sK1)),
fof(f818,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl16_30])])).
  spl16_85 <=> firstordertype(sK1)),
fof(f405,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl16_85])])).
  iof(sK9(sK1),sK1) | (~spl16_20 | ~spl16_30)),
fof(f821,plain,(
  inference(resolution,[],[f201,f258])).
  ( ! [X0] : (individual(X0) | ~iof(X0,sK1)) ) | (~spl16_40 | ~spl16_85)),
fof(f258,plain,(
  inference(resolution,[],[f819,f372])).
  type_(sK1) | ~spl16_30),
fof(f372,plain,(
  inference(avatar_component_clause,[],[f257])).
  ( ! [X0,X1] : (~firstordertype(X0) | individual(X1) | ~iof(X1,X0)) ) | ~spl16_40),
fof(f373,plain,(
  inference(avatar_component_clause,[],[f371])).
  spl16_40),
fof(f819,plain,(
  inference(avatar_split_clause,[],[f76,f371])).
  firstordertype(sK1) | ~spl16_85),
fof(f76,plain,(
  inference(avatar_component_clause,[],[f818])).
  ( ! [X0,X1] : (~iof(X1,X0) | individual(X1) | ~firstordertype(X0)) )),
fof(f901,plain,(
  inference(cnf_transformation,[],[f48])).
  spl16_100 | ~spl16_6 | ~spl16_98),
fof(f48,plain,(
  inference(avatar_split_clause,[],[f893,f888,f144,f899])).
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (individual(X1) | ~iof(X1,X0)) & type_(X0)))),
fof(f144,plain,(
  inference(ennf_transformation,[],[f10])).
  spl16_6 <=> ! [X0] : (type_(X0) | ~firstordertype(X0))),
fof(f10,axiom,(
  introduced(avatar_definition,[new_symbols(naming,[spl16_6])])).
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (iof(X1,X0) => individual(X1)) & type_(X0)))),
fof(f888,plain,(
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a5_firstOrderTypeDef)).
  spl16_98 <=> firstordertype(sK9(sK1))),
fof(f268,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl16_98])])).
  spl16_32),
fof(f893,plain,(
  inference(avatar_split_clause,[],[f72,f266])).
  type_(sK9(sK1)) | (~spl16_6 | ~spl16_98)),
fof(f72,plain,(
  inference(resolution,[],[f889,f145])).
  ( ! [X0,X1] : (~iof(X1,X0) | firstordertype(X1) | ~secondordertype(X0)) )),
fof(f145,plain,(
  inference(cnf_transformation,[],[f47])).
  ( ! [X0] : (~firstordertype(X0) | type_(X0)) ) | ~spl16_6),
fof(f47,plain,(
  inference(avatar_component_clause,[],[f144])).
  ! [X0] : (secondordertype(X0) <=> (! [X1] : (firstordertype(X1) | ~iof(X1,X0)) & type_(X0)))),
fof(f889,plain,(
  inference(ennf_transformation,[],[f11])).
  firstordertype(sK9(sK1)) | ~spl16_98),
  inference(avatar_component_clause,[],[f888])).
fof(f11,axiom,(
fof(f890,plain,(
  ! [X0] : (secondordertype(X0) <=> (! [X1] : (iof(X1,X0) => firstordertype(X1)) & type_(X0)))),
  spl16_98 | ~spl16_54 | ~spl16_96),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a6_secondOrderTypeDef)).
  inference(avatar_split_clause,[],[f885,f878,f453,f888])).
fof(f259,plain,(
fof(f878,plain,(
  spl16_30 | ~spl16_5 | ~spl16_28),
  spl16_96 <=> ! [X1] : (firstordertype(X1) | ~iof(X1,sK1))),
  inference(avatar_split_clause,[],[f251,f247,f140,f257])).
  introduced(avatar_definition,[new_symbols(naming,[spl16_96])])).
fof(f140,plain,(
fof(f885,plain,(
  spl16_5 <=> ! [X0] : (type_(X0) | ~secondordertype(X0))),
  firstordertype(sK9(sK1)) | (~spl16_54 | ~spl16_96)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_5])])).
  inference(resolution,[],[f879,f454])).
fof(f251,plain,(
fof(f879,plain,(
  type_(sK1) | (~spl16_5 | ~spl16_28)),
  ( ! [X1] : (~iof(X1,sK1) | firstordertype(X1)) ) | ~spl16_96),
  inference(resolution,[],[f248,f141])).
  inference(avatar_component_clause,[],[f878])).
fof(f141,plain,(
fof(f880,plain,(
  ( ! [X0] : (~secondordertype(X0) | type_(X0)) ) | ~spl16_5),
  spl16_96 | ~spl16_28 | ~spl16_32),
  inference(avatar_component_clause,[],[f140])).
  inference(avatar_split_clause,[],[f462,f266,f247,f878])).
fof(f249,plain,(
fof(f247,plain,(
  spl16_28 | ~spl16_2 | ~spl16_16),
  spl16_28 <=> secondordertype(sK1)),
  inference(avatar_split_clause,[],[f239,f184,f128,f247])).
  introduced(avatar_definition,[new_symbols(naming,[spl16_28])])).
fof(f128,plain,(
fof(f266,plain,(
  spl16_2 <=> iof(sK1,c_sot)),
  spl16_32 <=> ! [X1,X0] : (~iof(X1,X0) | firstordertype(X1) | ~secondordertype(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_2])])).
  introduced(avatar_definition,[new_symbols(naming,[spl16_32])])).
fof(f184,plain,(
fof(f462,plain,(
  spl16_16 <=> ! [X1] : (secondordertype(X1) | ~iof(X1,c_sot))),
  ( ! [X1] : (firstordertype(X1) | ~iof(X1,sK1)) ) | (~spl16_28 | ~spl16_32)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_16])])).
  inference(resolution,[],[f267,f248])).
fof(f239,plain,(
fof(f248,plain,(
  secondordertype(sK1) | (~spl16_2 | ~spl16_16)),
  secondordertype(sK1) | ~spl16_28),
  inference(resolution,[],[f185,f129])).
  inference(avatar_component_clause,[],[f247])).
fof(f129,plain,(
fof(f267,plain,(
  iof(sK1,c_sot) | ~spl16_2),
  ( ! [X0,X1] : (~secondordertype(X0) | firstordertype(X1) | ~iof(X1,X0)) ) | ~spl16_32),
  inference(avatar_component_clause,[],[f128])).
  inference(avatar_component_clause,[],[f266])).
fof(f185,plain,(
fof(f820,plain,(
  ( ! [X1] : (~iof(X1,c_sot) | secondordertype(X1)) ) | ~spl16_16),
  spl16_85 | ~spl16_3 | ~spl16_84),
  inference(avatar_component_clause,[],[f184])).
  inference(avatar_split_clause,[],[f804,f801,f132,f818])).
fof(f243,plain,(
fof(f132,plain,(
  spl16_27 | ~spl16_1 | ~spl16_16),
  spl16_3 <=> iof(sK1,sK0)),
  inference(avatar_split_clause,[],[f238,f184,f124,f241])).
  introduced(avatar_definition,[new_symbols(naming,[spl16_3])])).
fof(f124,plain,(
fof(f801,plain,(
  spl16_1 <=> iof(sK0,c_sot)),
  spl16_84 <=> ! [X0] : (firstordertype(X0) | ~iof(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_1])])).
  introduced(avatar_definition,[new_symbols(naming,[spl16_84])])).
fof(f238,plain,(
fof(f804,plain,(
  secondordertype(sK0) | (~spl16_1 | ~spl16_16)),
  firstordertype(sK1) | (~spl16_3 | ~spl16_84)),
  inference(resolution,[],[f185,f125])).
  inference(resolution,[],[f802,f133])).
fof(f125,plain,(
fof(f133,plain,(
  iof(sK0,c_sot) | ~spl16_1),
  iof(sK1,sK0) | ~spl16_3),
  inference(avatar_component_clause,[],[f124])).
  inference(avatar_component_clause,[],[f132])).
fof(f202,plain,(
fof(f802,plain,(
  spl16_20),
  ( ! [X0] : (~iof(X0,sK0) | firstordertype(X0)) ) | ~spl16_84),
  inference(avatar_split_clause,[],[f81,f200])).
  inference(avatar_component_clause,[],[f801])).
fof(f81,plain,(
fof(f803,plain,(
  ( ! [X0] : (iof(sK9(X0),X0) | ~type_(X0)) )),
  spl16_84 | ~spl16_27 | ~spl16_32),
  inference(cnf_transformation,[],[f3])).
  inference(avatar_split_clause,[],[f461,f266,f241,f801])).
fof(f3,axiom,(
fof(f241,plain,(
  spl16_27 <=> secondordertype(sK0)),
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_27])])).
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a3)).
fof(f461,plain,(
fof(f186,plain,(
  ( ! [X0] : (firstordertype(X0) | ~iof(X0,sK0)) ) | (~spl16_27 | ~spl16_32)),
  spl16_16),
  inference(resolution,[],[f267,f242])).
  inference(avatar_split_clause,[],[f112,f184])).
fof(f242,plain,(
fof(f112,plain,(
  secondordertype(sK0) | ~spl16_27),
  ( ! [X1] : (secondordertype(X1) | ~iof(X1,c_sot)) )),
  inference(avatar_component_clause,[],[f241])).
  inference(equality_resolution,[],[f91])).
fof(f455,plain,(
fof(f91,plain,(
  spl16_54 | ~spl16_20 | ~spl16_30),
  ( ! [X0,X1] : (secondordertype(X1) | ~iof(X1,X0) | c_sot != X0) )),
  inference(avatar_split_clause,[],[f405,f257,f200,f453])).
  inference(cnf_transformation,[],[f31])).
fof(f257,plain,(
fof(f31,plain,(
  spl16_30 <=> type_(sK1)),
  ! [X0] : (c_sot = X0 <=> ! [X1] : (iof(X1,X0) <=> secondordertype(X1)))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_30])])).
  inference(rectify,[],[f23])).
fof(f405,plain,(
fof(f23,axiom,(
  iof(sK9(sK1),sK1) | (~spl16_20 | ~spl16_30)),
  ! [X10] : (c_sot = X10 <=> ! [X0] : (iof(X0,X10) <=> secondordertype(X0)))),
  inference(resolution,[],[f201,f258])).
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a15_sotConstantDef)).
fof(f258,plain,(
fof(f154,plain,(
  type_(sK1) | ~spl16_30),
  spl16_8),
  inference(avatar_split_clause,[],[f79,f152])).
  inference(avatar_component_clause,[],[f257])).
fof(f79,plain,(
fof(f373,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | ~individual(X0)) )),
  spl16_40),
  inference(cnf_transformation,[],[f49])).
  inference(avatar_split_clause,[],[f76,f371])).
fof(f49,plain,(
fof(f76,plain,(
  ! [X0] : (individual(X0) <=> ! [X1] : ~iof(X1,X0))),
  ( ! [X0,X1] : (~iof(X1,X0) | individual(X1) | ~firstordertype(X0)) )),
  inference(cnf_transformation,[],[f48])).
  inference(ennf_transformation,[],[f2])).
fof(f48,plain,(
fof(f2,axiom,(
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (individual(X1) | ~iof(X1,X0)) & type_(X0)))),
  ! [X0] : (individual(X0) <=> ~? [X1] : iof(X1,X0))),
  inference(ennf_transformation,[],[f10])).
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a2)).
fof(f10,axiom,(
fof(f146,plain,(
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (iof(X1,X0) => individual(X1)) & type_(X0)))),
  spl16_6),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a5_firstOrderTypeDef)).
  inference(avatar_split_clause,[],[f77,f144])).
fof(f268,plain,(
fof(f77,plain,(
  spl16_32),
  ( ! [X0] : (type_(X0) | ~firstordertype(X0)) )),
  inference(avatar_split_clause,[],[f72,f266])).
  inference(cnf_transformation,[],[f48])).
fof(f72,plain,(
fof(f142,plain,(
  spl16_5),
  ( ! [X0,X1] : (~iof(X1,X0) | firstordertype(X1) | ~secondordertype(X0)) )),
  inference(avatar_split_clause,[],[f73,f140])).
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
fof(f73,plain,(
  ( ! [X0] : (type_(X0) | ~secondordertype(X0)) )),
  ! [X0] : (secondordertype(X0) <=> (! [X1] : (firstordertype(X1) | ~iof(X1,X0)) & type_(X0)))),
  inference(cnf_transformation,[],[f47])).
  inference(ennf_transformation,[],[f11])).
fof(f134,plain,(
fof(f11,axiom,(
  spl16_3),
  ! [X0] : (secondordertype(X0) <=> (! [X1] : (iof(X1,X0) => firstordertype(X1)) & type_(X0)))),
  inference(avatar_split_clause,[],[f59,f132])).
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a6_secondOrderTypeDef)).
fof(f59,plain,(
fof(f259,plain,(
  iof(sK1,sK0)),
  spl16_30 | ~spl16_5 | ~spl16_28),
  inference(avatar_split_clause,[],[f251,f247,f140,f257])).
  inference(cnf_transformation,[],[f27])).
fof(f140,plain,(
fof(f27,plain,(
  ? [X0,X1] : (iof(X1,X0) & iof(X1,c_sot) & iof(X0,c_sot))),
  spl16_5 <=> ! [X0] : (type_(X0) | ~secondordertype(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_5])])).
  inference(flattening,[],[f26])).
fof(f26,plain,(
fof(f251,plain,(
  ~~? [X0,X1] : (iof(X1,X0) & iof(X1,c_sot) & iof(X0,c_sot))),
  type_(sK1) | (~spl16_5 | ~spl16_28)),
  inference(rectify,[],[f25])).
  inference(resolution,[],[f248,f141])).
fof(f25,negated_conjecture,(
fof(f141,plain,(
  ~~? [X6,X7] : (iof(X7,X6) & iof(X7,c_sot) & iof(X6,c_sot))),
  ( ! [X0] : (~secondordertype(X0) | type_(X0)) ) | ~spl16_5),
  inference(negated_conjecture,[],[f24])).
  inference(avatar_component_clause,[],[f140])).
fof(f24,conjecture,(
fof(f249,plain,(
  ~? [X6,X7] : (iof(X7,X6) & iof(X7,c_sot) & iof(X6,c_sot))),
  spl16_28 | ~spl16_2 | ~spl16_16),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',t8_ap2)).
fof(f130,plain,(
  inference(avatar_split_clause,[],[f239,f184,f128,f247])).
  spl16_2),
fof(f128,plain,(
  inference(avatar_split_clause,[],[f58,f128])).
  spl16_2 <=> iof(sK1,c_sot)),
fof(f58,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl16_2])])).
  iof(sK1,c_sot)),
fof(f184,plain,(
  inference(cnf_transformation,[],[f27])).
  spl16_16 <=> ! [X1] : (secondordertype(X1) | ~iof(X1,c_sot))),
fof(f126,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl16_16])])).
  spl16_1),
fof(f239,plain,(
  inference(avatar_split_clause,[],[f57,f124])).
  secondordertype(sK1) | (~spl16_2 | ~spl16_16)),
fof(f57,plain,(
  inference(resolution,[],[f185,f129])).
  iof(sK0,c_sot)),
fof(f129,plain,(
  inference(cnf_transformation,[],[f27])).
  iof(sK1,c_sot) | ~spl16_2),
% SZS output end Proof for SOT_qBPq
  inference(avatar_component_clause,[],[f128])).
% ------------------------------
fof(f185,plain,(
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
  ( ! [X1] : (~iof(X1,c_sot) | secondordertype(X1)) ) | ~spl16_16),
% Termination reason: Refutation
  inference(avatar_component_clause,[],[f184])).

fof(f243,plain,(
% Memory used [KB]: 5884
  spl16_27 | ~spl16_1 | ~spl16_16),
% Time elapsed: 0.020 s
  inference(avatar_split_clause,[],[f238,f184,f124,f241])).
% ------------------------------
fof(f124,plain,(
% ------------------------------
  spl16_1 <=> iof(sK0,c_sot)),
% Success in time 0.061 s
  introduced(avatar_definition,[new_symbols(naming,[spl16_1])])).
fof(f238,plain,(
  secondordertype(sK0) | (~spl16_1 | ~spl16_16)),
  inference(resolution,[],[f185,f125])).
fof(f125,plain,(
  iof(sK0,c_sot) | ~spl16_1),
  inference(avatar_component_clause,[],[f124])).
fof(f202,plain,(
  spl16_20),
  inference(avatar_split_clause,[],[f81,f200])).
fof(f81,plain,(
  ( ! [X0] : (iof(sK9(X0),X0) | ~type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a3)).
fof(f186,plain,(
  spl16_16),
  inference(avatar_split_clause,[],[f112,f184])).
fof(f112,plain,(
  ( ! [X1] : (secondordertype(X1) | ~iof(X1,c_sot)) )),
  inference(equality_resolution,[],[f91])).
fof(f91,plain,(
  ( ! [X0,X1] : (secondordertype(X1) | ~iof(X1,X0) | c_sot != X0) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : (c_sot = X0 <=> ! [X1] : (iof(X1,X0) <=> secondordertype(X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10] : (c_sot = X10 <=> ! [X0] : (iof(X0,X10) <=> secondordertype(X0)))),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a15_sotConstantDef)).
fof(f154,plain,(
  spl16_8),
  inference(avatar_split_clause,[],[f79,f152])).
fof(f79,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | ~individual(X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (individual(X0) <=> ! [X1] : ~iof(X1,X0))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0] : (individual(X0) <=> ~? [X1] : iof(X1,X0))),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a2)).
fof(f146,plain,(
  spl16_6),
  inference(avatar_split_clause,[],[f77,f144])).
fof(f77,plain,(
  ( ! [X0] : (type_(X0) | ~firstordertype(X0)) )),
  inference(cnf_transformation,[],[f48])).
fof(f142,plain,(
  spl16_5),
  inference(avatar_split_clause,[],[f73,f140])).
fof(f73,plain,(
  ( ! [X0] : (type_(X0) | ~secondordertype(X0)) )),
  inference(cnf_transformation,[],[f47])).
fof(f134,plain,(
  spl16_3),
  inference(avatar_split_clause,[],[f59,f132])).
fof(f59,plain,(
  iof(sK1,sK0)),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ? [X0,X1] : (iof(X1,X0) & iof(X1,c_sot) & iof(X0,c_sot))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ~~? [X0,X1] : (iof(X1,X0) & iof(X1,c_sot) & iof(X0,c_sot))),
  inference(rectify,[],[f25])).
fof(f25,negated_conjecture,(
  ~~? [X6,X7] : (iof(X7,X6) & iof(X7,c_sot) & iof(X6,c_sot))),
  inference(negated_conjecture,[],[f24])).
fof(f24,conjecture,(
  ~? [X6,X7] : (iof(X7,X6) & iof(X7,c_sot) & iof(X6,c_sot))),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',t8_ap2)).
fof(f130,plain,(
  spl16_2),
  inference(avatar_split_clause,[],[f58,f128])).
fof(f58,plain,(
  iof(sK1,c_sot)),
  inference(cnf_transformation,[],[f27])).
fof(f126,plain,(
  spl16_1),
  inference(avatar_split_clause,[],[f57,f124])).
fof(f57,plain,(
  iof(sK0,c_sot)),
  inference(cnf_transformation,[],[f27])).
% SZS output end Proof for SOT_qBPq
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 5884
% Time elapsed: 0.021 s
% ------------------------------
% ------------------------------
% Success in time 0.065 s
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_qBPq
% SZS output start Proof for SOT_qBPq
fof(f1121,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f126,f130,f134,f142,f146,f154,f186,f202,f243,f249,f259,f268,f373,f455,f803,f820,f880,f890,f901,f1102,f1116,f1120])).
fof(f1120,plain,(
  ~spl16_8 | ~spl16_20 | ~spl16_100 | ~spl16_123),
  inference(avatar_contradiction_clause,[],[f1119])).
fof(f1119,plain,(
  $false | (~spl16_8 | ~spl16_20 | ~spl16_100 | ~spl16_123)),
  inference(subsumption_resolution,[],[f904,f1117])).
fof(f1117,plain,(
  ( ! [X0] : (~iof(X0,sK9(sK1))) ) | (~spl16_8 | ~spl16_123)),
  inference(resolution,[],[f1115,f153])).
fof(f153,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) ) | ~spl16_8),
  inference(avatar_component_clause,[],[f152])).
fof(f152,plain,(
  spl16_8 <=> ! [X1,X0] : (~iof(X1,X0) | ~individual(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_8])])).
fof(f1115,plain,(
  individual(sK9(sK1)) | ~spl16_123),
  inference(avatar_component_clause,[],[f1114])).
fof(f1114,plain,(
  spl16_123 <=> individual(sK9(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_123])])).
fof(f904,plain,(
  iof(sK9(sK9(sK1)),sK9(sK1)) | (~spl16_20 | ~spl16_100)),
  inference(resolution,[],[f900,f201])).
fof(f201,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK9(X0),X0)) ) | ~spl16_20),
  inference(avatar_component_clause,[],[f200])).
fof(f200,plain,(
  spl16_20 <=> ! [X0] : (iof(sK9(X0),X0) | ~type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_20])])).
fof(f900,plain,(
  type_(sK9(sK1)) | ~spl16_100),
  inference(avatar_component_clause,[],[f899])).
fof(f899,plain,(
  spl16_100 <=> type_(sK9(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_100])])).
fof(f1116,plain,(
  spl16_123 | ~spl16_54 | ~spl16_121),
  inference(avatar_split_clause,[],[f1107,f1100,f453,f1114])).
fof(f453,plain,(
  spl16_54 <=> iof(sK9(sK1),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_54])])).
fof(f1100,plain,(
  spl16_121 <=> ! [X0] : (individual(X0) | ~iof(X0,sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_121])])).
fof(f1107,plain,(
  individual(sK9(sK1)) | (~spl16_54 | ~spl16_121)),
  inference(resolution,[],[f1101,f454])).
fof(f454,plain,(
  iof(sK9(sK1),sK1) | ~spl16_54),
  inference(avatar_component_clause,[],[f453])).
fof(f1101,plain,(
  ( ! [X0] : (~iof(X0,sK1) | individual(X0)) ) | ~spl16_121),
  inference(avatar_component_clause,[],[f1100])).
fof(f1102,plain,(
  spl16_121 | ~spl16_40 | ~spl16_85),
  inference(avatar_split_clause,[],[f821,f818,f371,f1100])).
fof(f371,plain,(
  spl16_40 <=> ! [X1,X0] : (~iof(X1,X0) | individual(X1) | ~firstordertype(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_40])])).
fof(f818,plain,(
  spl16_85 <=> firstordertype(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_85])])).
fof(f821,plain,(
  ( ! [X0] : (individual(X0) | ~iof(X0,sK1)) ) | (~spl16_40 | ~spl16_85)),
  inference(resolution,[],[f819,f372])).
fof(f372,plain,(
  ( ! [X0,X1] : (~firstordertype(X0) | individual(X1) | ~iof(X1,X0)) ) | ~spl16_40),
  inference(avatar_component_clause,[],[f371])).
fof(f819,plain,(
  firstordertype(sK1) | ~spl16_85),
  inference(avatar_component_clause,[],[f818])).
fof(f901,plain,(
  spl16_100 | ~spl16_6 | ~spl16_98),
  inference(avatar_split_clause,[],[f893,f888,f144,f899])).
fof(f144,plain,(
  spl16_6 <=> ! [X0] : (type_(X0) | ~firstordertype(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_6])])).
fof(f888,plain,(
  spl16_98 <=> firstordertype(sK9(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_98])])).
fof(f893,plain,(
  type_(sK9(sK1)) | (~spl16_6 | ~spl16_98)),
  inference(resolution,[],[f889,f145])).
fof(f145,plain,(
  ( ! [X0] : (~firstordertype(X0) | type_(X0)) ) | ~spl16_6),
  inference(avatar_component_clause,[],[f144])).
fof(f889,plain,(
  firstordertype(sK9(sK1)) | ~spl16_98),
  inference(avatar_component_clause,[],[f888])).
fof(f890,plain,(
  spl16_98 | ~spl16_54 | ~spl16_96),
  inference(avatar_split_clause,[],[f885,f878,f453,f888])).
fof(f878,plain,(
  spl16_96 <=> ! [X1] : (firstordertype(X1) | ~iof(X1,sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_96])])).
fof(f885,plain,(
  firstordertype(sK9(sK1)) | (~spl16_54 | ~spl16_96)),
  inference(resolution,[],[f879,f454])).
fof(f879,plain,(
  ( ! [X1] : (~iof(X1,sK1) | firstordertype(X1)) ) | ~spl16_96),
  inference(avatar_component_clause,[],[f878])).
fof(f880,plain,(
  spl16_96 | ~spl16_28 | ~spl16_32),
  inference(avatar_split_clause,[],[f462,f266,f247,f878])).
fof(f247,plain,(
  spl16_28 <=> secondordertype(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_28])])).
fof(f266,plain,(
  spl16_32 <=> ! [X1,X0] : (~iof(X1,X0) | firstordertype(X1) | ~secondordertype(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_32])])).
fof(f462,plain,(
  ( ! [X1] : (firstordertype(X1) | ~iof(X1,sK1)) ) | (~spl16_28 | ~spl16_32)),
  inference(resolution,[],[f267,f248])).
fof(f248,plain,(
  secondordertype(sK1) | ~spl16_28),
  inference(avatar_component_clause,[],[f247])).
fof(f267,plain,(
  ( ! [X0,X1] : (~secondordertype(X0) | firstordertype(X1) | ~iof(X1,X0)) ) | ~spl16_32),
  inference(avatar_component_clause,[],[f266])).
fof(f820,plain,(
  spl16_85 | ~spl16_3 | ~spl16_84),
  inference(avatar_split_clause,[],[f804,f801,f132,f818])).
fof(f132,plain,(
  spl16_3 <=> iof(sK1,sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_3])])).
fof(f801,plain,(
  spl16_84 <=> ! [X0] : (firstordertype(X0) | ~iof(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_84])])).
fof(f804,plain,(
  firstordertype(sK1) | (~spl16_3 | ~spl16_84)),
  inference(resolution,[],[f802,f133])).
fof(f133,plain,(
  iof(sK1,sK0) | ~spl16_3),
  inference(avatar_component_clause,[],[f132])).
fof(f802,plain,(
  ( ! [X0] : (~iof(X0,sK0) | firstordertype(X0)) ) | ~spl16_84),
  inference(avatar_component_clause,[],[f801])).
fof(f803,plain,(
  spl16_84 | ~spl16_27 | ~spl16_32),
  inference(avatar_split_clause,[],[f461,f266,f241,f801])).
fof(f241,plain,(
  spl16_27 <=> secondordertype(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_27])])).
fof(f461,plain,(
  ( ! [X0] : (firstordertype(X0) | ~iof(X0,sK0)) ) | (~spl16_27 | ~spl16_32)),
  inference(resolution,[],[f267,f242])).
fof(f242,plain,(
  secondordertype(sK0) | ~spl16_27),
  inference(avatar_component_clause,[],[f241])).
fof(f455,plain,(
  spl16_54 | ~spl16_20 | ~spl16_30),
  inference(avatar_split_clause,[],[f405,f257,f200,f453])).
fof(f257,plain,(
  spl16_30 <=> type_(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_30])])).
fof(f405,plain,(
  iof(sK9(sK1),sK1) | (~spl16_20 | ~spl16_30)),
  inference(resolution,[],[f201,f258])).
fof(f258,plain,(
  type_(sK1) | ~spl16_30),
  inference(avatar_component_clause,[],[f257])).
fof(f373,plain,(
  spl16_40),
  inference(avatar_split_clause,[],[f76,f371])).
fof(f76,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | individual(X1) | ~firstordertype(X0)) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (individual(X1) | ~iof(X1,X0)) & type_(X0)))),
  inference(ennf_transformation,[],[f10])).
fof(f10,axiom,(
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (iof(X1,X0) => individual(X1)) & type_(X0)))),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a5_firstOrderTypeDef)).
fof(f268,plain,(
  spl16_32),
  inference(avatar_split_clause,[],[f72,f266])).
fof(f72,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | firstordertype(X1) | ~secondordertype(X0)) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (secondordertype(X0) <=> (! [X1] : (firstordertype(X1) | ~iof(X1,X0)) & type_(X0)))),
  inference(ennf_transformation,[],[f11])).
fof(f11,axiom,(
  ! [X0] : (secondordertype(X0) <=> (! [X1] : (iof(X1,X0) => firstordertype(X1)) & type_(X0)))),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a6_secondOrderTypeDef)).
fof(f259,plain,(
  spl16_30 | ~spl16_5 | ~spl16_28),
  inference(avatar_split_clause,[],[f251,f247,f140,f257])).
fof(f140,plain,(
  spl16_5 <=> ! [X0] : (type_(X0) | ~secondordertype(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_5])])).
fof(f251,plain,(
  type_(sK1) | (~spl16_5 | ~spl16_28)),
  inference(resolution,[],[f248,f141])).
fof(f141,plain,(
  ( ! [X0] : (~secondordertype(X0) | type_(X0)) ) | ~spl16_5),
  inference(avatar_component_clause,[],[f140])).
fof(f249,plain,(
  spl16_28 | ~spl16_2 | ~spl16_16),
  inference(avatar_split_clause,[],[f239,f184,f128,f247])).
fof(f128,plain,(
  spl16_2 <=> iof(sK1,c_sot)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_2])])).
fof(f184,plain,(
  spl16_16 <=> ! [X1] : (secondordertype(X1) | ~iof(X1,c_sot))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_16])])).
fof(f239,plain,(
  secondordertype(sK1) | (~spl16_2 | ~spl16_16)),
  inference(resolution,[],[f185,f129])).
fof(f129,plain,(
  iof(sK1,c_sot) | ~spl16_2),
  inference(avatar_component_clause,[],[f128])).
fof(f185,plain,(
  ( ! [X1] : (~iof(X1,c_sot) | secondordertype(X1)) ) | ~spl16_16),
  inference(avatar_component_clause,[],[f184])).
fof(f243,plain,(
  spl16_27 | ~spl16_1 | ~spl16_16),
  inference(avatar_split_clause,[],[f238,f184,f124,f241])).
fof(f124,plain,(
  spl16_1 <=> iof(sK0,c_sot)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_1])])).
fof(f238,plain,(
  secondordertype(sK0) | (~spl16_1 | ~spl16_16)),
  inference(resolution,[],[f185,f125])).
fof(f125,plain,(
  iof(sK0,c_sot) | ~spl16_1),
  inference(avatar_component_clause,[],[f124])).
fof(f202,plain,(
  spl16_20),
  inference(avatar_split_clause,[],[f81,f200])).
fof(f81,plain,(
  ( ! [X0] : (iof(sK9(X0),X0) | ~type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a3)).
fof(f186,plain,(
  spl16_16),
  inference(avatar_split_clause,[],[f112,f184])).
fof(f112,plain,(
  ( ! [X1] : (secondordertype(X1) | ~iof(X1,c_sot)) )),
  inference(equality_resolution,[],[f91])).
fof(f91,plain,(
  ( ! [X0,X1] : (secondordertype(X1) | ~iof(X1,X0) | c_sot != X0) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : (c_sot = X0 <=> ! [X1] : (iof(X1,X0) <=> secondordertype(X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10] : (c_sot = X10 <=> ! [X0] : (iof(X0,X10) <=> secondordertype(X0)))),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a15_sotConstantDef)).
fof(f154,plain,(
  spl16_8),
  inference(avatar_split_clause,[],[f79,f152])).
fof(f79,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | ~individual(X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (individual(X0) <=> ! [X1] : ~iof(X1,X0))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0] : (individual(X0) <=> ~? [X1] : iof(X1,X0))),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a2)).
fof(f146,plain,(
  spl16_6),
  inference(avatar_split_clause,[],[f77,f144])).
fof(f77,plain,(
  ( ! [X0] : (type_(X0) | ~firstordertype(X0)) )),
  inference(cnf_transformation,[],[f48])).
fof(f142,plain,(
  spl16_5),
  inference(avatar_split_clause,[],[f73,f140])).
fof(f73,plain,(
  ( ! [X0] : (type_(X0) | ~secondordertype(X0)) )),
  inference(cnf_transformation,[],[f47])).
fof(f134,plain,(
  spl16_3),
  inference(avatar_split_clause,[],[f59,f132])).
fof(f59,plain,(
  iof(sK1,sK0)),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ? [X0,X1] : (iof(X1,X0) & iof(X1,c_sot) & iof(X0,c_sot))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ~~? [X0,X1] : (iof(X1,X0) & iof(X1,c_sot) & iof(X0,c_sot))),
  inference(rectify,[],[f25])).
fof(f25,negated_conjecture,(
  ~~? [X6,X7] : (iof(X7,X6) & iof(X7,c_sot) & iof(X6,c_sot))),
  inference(negated_conjecture,[],[f24])).
fof(f24,conjecture,(
  ~? [X6,X7] : (iof(X7,X6) & iof(X7,c_sot) & iof(X6,c_sot))),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',t8_ap2)).
fof(f130,plain,(
  spl16_2),
  inference(avatar_split_clause,[],[f58,f128])).
fof(f58,plain,(
  iof(sK1,c_sot)),
  inference(cnf_transformation,[],[f27])).
fof(f126,plain,(
  spl16_1),
  inference(avatar_split_clause,[],[f57,f124])).
fof(f57,plain,(
  iof(sK0,c_sot)),
  inference(cnf_transformation,[],[f27])).
% SZS output end Proof for SOT_qBPq
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 5884
% Time elapsed: 0.026 s
% ------------------------------
% ------------------------------
% Success in time 0.078 s
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 14 seconds
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 14 seconds
% START OF SYSTEM OUTPUT
# AutoSched0-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S5PRR_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.017 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(a2, axiom, ![X1]:(individual(X1)<=>~(?[X2]:iof(X2,X1))), file('/tmp/WT2askIkzH/SOT_qBPqo3', a2)).
fof(a3, axiom, ![X1]:(type_(X1)<=>?[X2]:iof(X2,X1)), file('/tmp/WT2askIkzH/SOT_qBPqo3', a3)).
fof(a6_secondOrderTypeDef, axiom, ![X1]:(secondordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>firstordertype(X2)))), file('/tmp/WT2askIkzH/SOT_qBPqo3', a6_secondOrderTypeDef)).
fof(a5_firstOrderTypeDef, axiom, ![X1]:(firstordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>individual(X2)))), file('/tmp/WT2askIkzH/SOT_qBPqo3', a5_firstOrderTypeDef)).
fof(a15_sotConstantDef, axiom, ![X11]:(X11=c_sot<=>![X1]:(iof(X1,X11)<=>secondordertype(X1))), file('/tmp/WT2askIkzH/SOT_qBPqo3', a15_sotConstantDef)).
# AutoSched0-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S5PRR_S0Y
fof(t8_ap2, conjecture, ~(?[X7, X8]:((iof(X7,c_sot)&iof(X8,c_sot))&iof(X8,X7))), file('/tmp/WT2askIkzH/SOT_qBPqo3', t8_ap2)).
fof(c_0_6, plain, ![X13, X14, X15]:((~individual(X13)|~iof(X14,X13))&(iof(esk1_1(X15),X15)|individual(X15))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a2])])])])])).
fof(c_0_7, plain, ![X17, X19, X20]:((~type_(X17)|iof(esk2_1(X17),X17))&(~iof(X20,X19)|type_(X19))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a3])])])])])).
cnf(c_0_8, plain, (~individual(X1)|~iof(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
# and selection function SelectMaxLComplexAvoidPosPred.
cnf(c_0_9, plain, (iof(esk2_1(X1),X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, plain, ![X46, X47, X48]:(((type_(X46)|~secondordertype(X46))&(~iof(X47,X46)|firstordertype(X47)|~secondordertype(X46)))&((iof(esk10_1(X48),X48)|~type_(X48)|secondordertype(X48))&(~firstordertype(esk10_1(X48))|~type_(X48)|secondordertype(X48)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a6_secondOrderTypeDef])])])])])])).
fof(c_0_11, plain, ![X42, X43, X44]:(((type_(X42)|~firstordertype(X42))&(~iof(X43,X42)|individual(X43)|~firstordertype(X42)))&((iof(esk9_1(X44),X44)|~type_(X44)|firstordertype(X44))&(~individual(esk9_1(X44))|~type_(X44)|firstordertype(X44)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a5_firstOrderTypeDef])])])])])])).
#
cnf(c_0_12, plain, (~type_(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_13, plain, (type_(X1)|~secondordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
# Preprocessing time       : 0.016 s
fof(c_0_14, plain, ![X114, X115, X116, X117]:(((~iof(X115,X114)|secondordertype(X115)|X114!=c_sot)&(~secondordertype(X116)|iof(X116,X114)|X114!=c_sot))&((~iof(esk24_1(X117),X117)|~secondordertype(esk24_1(X117))|X117=c_sot)&(iof(esk24_1(X117),X117)|secondordertype(esk24_1(X117))|X117=c_sot))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a15_sotConstantDef])])])])])])).
fof(c_0_15, negated_conjecture, ~(~(?[X7, X8]:((iof(X7,c_sot)&iof(X8,c_sot))&iof(X8,X7)))), inference(assume_negation,[status(cth)],[t8_ap2])).
cnf(c_0_16, plain, (type_(X1)|~firstordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).

cnf(c_0_17, plain, (firstordertype(X1)|~iof(X1,X2)|~secondordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, plain, (iof(esk1_1(X1),X1)|individual(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_19, plain, (~secondordertype(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
# Proof found!
cnf(c_0_20, plain, (secondordertype(X1)|~iof(X1,X2)|X2!=c_sot), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_21, negated_conjecture, ((iof(esk25_0,c_sot)&iof(esk26_0,c_sot))&iof(esk26_0,esk25_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
cnf(c_0_22, plain, (~firstordertype(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_12, c_0_16])).
# SZS status Theorem
cnf(c_0_23, plain, (firstordertype(esk1_1(X1))|~secondordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])).
cnf(c_0_24, plain, (individual(X1)|~iof(X1,X2)|~firstordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_25, plain, (secondordertype(X1)|~iof(X1,c_sot)), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_26, negated_conjecture, (iof(esk25_0,c_sot)), inference(split_conjunct,[status(thm)],[c_0_21])).
# SZS output start CNFRefutation
cnf(c_0_27, plain, (~secondordertype(X1)|~individual(esk1_1(X1))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_28, plain, (individual(esk1_1(X1))|~firstordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_18]), c_0_22])).
fof(a2, axiom, ![X1]:(individual(X1)<=>~(?[X2]:iof(X2,X1))), file('/tmp/WT2askIkzH/SOT_qBPqo3', a2)).
cnf(c_0_29, negated_conjecture, (iof(esk26_0,c_sot)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(a3, axiom, ![X1]:(type_(X1)<=>?[X2]:iof(X2,X1)), file('/tmp/WT2askIkzH/SOT_qBPqo3', a3)).
cnf(c_0_30, negated_conjecture, (iof(esk26_0,esk25_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(a6_secondOrderTypeDef, axiom, ![X1]:(secondordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>firstordertype(X2)))), file('/tmp/WT2askIkzH/SOT_qBPqo3', a6_secondOrderTypeDef)).
cnf(c_0_31, negated_conjecture, (secondordertype(esk25_0)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
fof(a5_firstOrderTypeDef, axiom, ![X1]:(firstordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>individual(X2)))), file('/tmp/WT2askIkzH/SOT_qBPqo3', a5_firstOrderTypeDef)).
cnf(c_0_32, plain, (~secondordertype(X1)|~firstordertype(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
fof(a15_sotConstantDef, axiom, ![X11]:(X11=c_sot<=>![X1]:(iof(X1,X11)<=>secondordertype(X1))), file('/tmp/WT2askIkzH/SOT_qBPqo3', a15_sotConstantDef)).
cnf(c_0_33, negated_conjecture, (secondordertype(esk26_0)), inference(spm,[status(thm)],[c_0_25, c_0_29])).
fof(t8_ap2, conjecture, ~(?[X7, X8]:((iof(X7,c_sot)&iof(X8,c_sot))&iof(X8,X7))), file('/tmp/WT2askIkzH/SOT_qBPqo3', t8_ap2)).
cnf(c_0_34, negated_conjecture, (firstordertype(esk26_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_30]), c_0_31])])).
fof(c_0_6, plain, ![X13, X14, X15]:((~individual(X13)|~iof(X14,X13))&(iof(esk1_1(X15),X15)|individual(X15))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a2])])])])])).
cnf(c_0_35, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])]), ['proof']).
fof(c_0_7, plain, ![X17, X19, X20]:((~type_(X17)|iof(esk2_1(X17),X17))&(~iof(X20,X19)|type_(X19))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a3])])])])])).
# SZS output end CNFRefutation
cnf(c_0_8, plain, (~individual(X1)|~iof(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
# Proof object total steps             : 36
cnf(c_0_9, plain, (iof(esk2_1(X1),X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
# Proof object clause steps            : 23
# Proof object formula steps           : 13
fof(c_0_10, plain, ![X46, X47, X48]:(((type_(X46)|~secondordertype(X46))&(~iof(X47,X46)|firstordertype(X47)|~secondordertype(X46)))&((iof(esk10_1(X48),X48)|~type_(X48)|secondordertype(X48))&(~firstordertype(esk10_1(X48))|~type_(X48)|secondordertype(X48)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a6_secondOrderTypeDef])])])])])])).
# Proof object conjectures             : 10
fof(c_0_11, plain, ![X42, X43, X44]:(((type_(X42)|~firstordertype(X42))&(~iof(X43,X42)|individual(X43)|~firstordertype(X42)))&((iof(esk9_1(X44),X44)|~type_(X44)|firstordertype(X44))&(~individual(esk9_1(X44))|~type_(X44)|firstordertype(X44)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a5_firstOrderTypeDef])])])])])])).
# Proof object clause conjectures      : 7
cnf(c_0_12, plain, (~type_(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
# Proof object formula conjectures     : 3
cnf(c_0_13, plain, (type_(X1)|~secondordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
# Proof object initial clauses used    : 11
fof(c_0_14, plain, ![X114, X115, X116, X117]:(((~iof(X115,X114)|secondordertype(X115)|X114!=c_sot)&(~secondordertype(X116)|iof(X116,X114)|X114!=c_sot))&((~iof(esk24_1(X117),X117)|~secondordertype(esk24_1(X117))|X117=c_sot)&(iof(esk24_1(X117),X117)|secondordertype(esk24_1(X117))|X117=c_sot))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a15_sotConstantDef])])])])])])).
# Proof object initial formulas used   : 6
fof(c_0_15, negated_conjecture, ~(~(?[X7, X8]:((iof(X7,c_sot)&iof(X8,c_sot))&iof(X8,X7)))), inference(assume_negation,[status(cth)],[t8_ap2])).
# Proof object generating inferences   : 11
cnf(c_0_16, plain, (type_(X1)|~firstordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
# Proof object simplifying inferences  : 7
cnf(c_0_17, plain, (firstordertype(X1)|~iof(X1,X2)|~secondordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
# Training examples: 0 positive, 0 negative
cnf(c_0_18, plain, (iof(esk1_1(X1),X1)|individual(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
# Parsed axioms                        : 24
cnf(c_0_19, plain, (~secondordertype(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
# Removed by relevancy pruning/SinE    : 0
cnf(c_0_20, plain, (secondordertype(X1)|~iof(X1,X2)|X2!=c_sot), inference(split_conjunct,[status(thm)],[c_0_14])).
# Initial clauses                      : 89
fof(c_0_21, negated_conjecture, ((iof(esk25_0,c_sot)&iof(esk26_0,c_sot))&iof(esk26_0,esk25_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
# Removed in clause preprocessing      : 5
cnf(c_0_22, plain, (~firstordertype(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_12, c_0_16])).
# Initial clauses in saturation        : 84
cnf(c_0_23, plain, (firstordertype(esk1_1(X1))|~secondordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])).
# Processed clauses                    : 153
cnf(c_0_24, plain, (individual(X1)|~iof(X1,X2)|~firstordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
# ...of these trivial                  : 2
cnf(c_0_25, plain, (secondordertype(X1)|~iof(X1,c_sot)), inference(er,[status(thm)],[c_0_20])).
# ...subsumed                          : 16
cnf(c_0_26, negated_conjecture, (iof(esk25_0,c_sot)), inference(split_conjunct,[status(thm)],[c_0_21])).
# ...remaining for further processing  : 135
# Other redundant clauses eliminated   : 7
cnf(c_0_27, plain, (~secondordertype(X1)|~individual(esk1_1(X1))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
# Clauses deleted for lack of memory   : 0
cnf(c_0_28, plain, (individual(esk1_1(X1))|~firstordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_18]), c_0_22])).
# Backward-subsumed                    : 1
cnf(c_0_29, negated_conjecture, (iof(esk26_0,c_sot)), inference(split_conjunct,[status(thm)],[c_0_21])).
# Backward-rewritten                   : 3
cnf(c_0_30, negated_conjecture, (iof(esk26_0,esk25_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
# Generated clauses                    : 256
cnf(c_0_31, negated_conjecture, (secondordertype(esk25_0)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
# ...of the previous two non-trivial   : 223
cnf(c_0_32, plain, (~secondordertype(X1)|~firstordertype(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
# Contextual simplify-reflections      : 15
cnf(c_0_33, negated_conjecture, (secondordertype(esk26_0)), inference(spm,[status(thm)],[c_0_25, c_0_29])).
# Paramodulations                      : 249
cnf(c_0_34, negated_conjecture, (firstordertype(esk26_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_30]), c_0_31])])).
# Factorizations                       : 0
cnf(c_0_35, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])]), ['proof']).
# Equation resolutions                 : 7
# SZS output end CNFRefutation
# Propositional unsat checks           : 0
# Proof object total steps             : 36
#    Propositional check models        : 0
# Proof object clause steps            : 23
#    Propositional check unsatisfiable : 0
# Proof object formula steps           : 13
#    Propositional clauses             : 0
# Proof object conjectures             : 10
#    Propositional clauses after purity: 0
# Proof object clause conjectures      : 7
#    Propositional unsat core size     : 0
# Proof object formula conjectures     : 3
#    Propositional preprocessing time  : 0.000
# Proof object initial clauses used    : 11
#    Propositional encoding time       : 0.000
# Proof object initial formulas used   : 6
#    Propositional solver time         : 0.000
# Proof object generating inferences   : 11
#    Success case prop preproc time    : 0.000
# Proof object simplifying inferences  : 7
#    Success case prop encoding time   : 0.000
# Training examples: 0 positive, 0 negative
#    Success case prop solver time     : 0.000
# Parsed axioms                        : 24
# Current number of processed clauses  : 124
# Removed by relevancy pruning/SinE    : 0
#    Positive orientable unit clauses  : 12
# Initial clauses                      : 89
#    Positive unorientable unit clauses: 0
# Removed in clause preprocessing      : 5
#    Negative unit clauses             : 9
# Initial clauses in saturation        : 84
#    Non-unit-clauses                  : 103
# Processed clauses                    : 153
# Current number of unprocessed clauses: 149
# ...of these trivial                  : 2
# ...number of literals in the above   : 523
# ...subsumed                          : 16
# Current number of archived formulas  : 0
# ...remaining for further processing  : 135
# Current number of archived clauses   : 5
# Other redundant clauses eliminated   : 7
# Clause-clause subsumption calls (NU) : 3206
# Clauses deleted for lack of memory   : 0
# Rec. Clause-clause subsumption calls : 2282
# Backward-subsumed                    : 1
# Backward-rewritten                   : 3
# Non-unit clause-clause subsumptions  : 22
# Generated clauses                    : 256
# Unit Clause-clause subsumption calls : 565
# ...of the previous two non-trivial   : 223
# Rewrite failures with RHS unbound    : 0
# Contextual simplify-reflections      : 15
# BW rewrite match attempts            : 2
# Paramodulations                      : 249
# BW rewrite match successes           : 2
# Factorizations                       : 0
# Condensation attempts                : 0
# Equation resolutions                 : 7
# Condensation successes               : 0
# Propositional unsat checks           : 0
# Termbank termtop insertions          : 8115
#    Propositional check models        : 0

#    Propositional check unsatisfiable : 0
# -------------------------------------------------
#    Propositional clauses             : 0
# User time                : 0.022 s
#    Propositional clauses after purity: 0
# System time              : 0.006 s
#    Propositional unsat core size     : 0
# Total time               : 0.028 s
#    Propositional preprocessing time  : 0.000
# Maximum resident set size: 1608 pages
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 124
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 103
# Current number of unprocessed clauses: 149
# ...number of literals in the above   : 523
# Current number of archived formulas  : 0
# Current number of archived clauses   : 5
# Clause-clause subsumption calls (NU) : 3206
# Rec. Clause-clause subsumption calls : 2282
# Non-unit clause-clause subsumptions  : 22
# Unit Clause-clause subsumption calls : 565
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8115

# -------------------------------------------------
# User time                : 0.023 s
# System time              : 0.003 s
# Total time               : 0.025 s
# Maximum resident set size: 1612 pages
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_qBPq
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_qBPq
% SZS output start Proof for SOT_qBPq
fof(f1121,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f126,f130,f134,f142,f146,f154,f186,f202,f243,f249,f259,f268,f373,f455,f803,f820,f880,f890,f901,f1102,f1116,f1120])).
fof(f1120,plain,(
  ~spl16_8 | ~spl16_20 | ~spl16_100 | ~spl16_123),
  inference(avatar_contradiction_clause,[],[f1119])).
fof(f1119,plain,(
  $false | (~spl16_8 | ~spl16_20 | ~spl16_100 | ~spl16_123)),
  inference(subsumption_resolution,[],[f904,f1117])).
fof(f1117,plain,(
  ( ! [X0] : (~iof(X0,sK9(sK1))) ) | (~spl16_8 | ~spl16_123)),
  inference(resolution,[],[f1115,f153])).
fof(f153,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) ) | ~spl16_8),
  inference(avatar_component_clause,[],[f152])).
fof(f152,plain,(
  spl16_8 <=> ! [X1,X0] : (~iof(X1,X0) | ~individual(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_8])])).
fof(f1115,plain,(
  individual(sK9(sK1)) | ~spl16_123),
  inference(avatar_component_clause,[],[f1114])).
fof(f1114,plain,(
  spl16_123 <=> individual(sK9(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_123])])).
fof(f904,plain,(
  iof(sK9(sK9(sK1)),sK9(sK1)) | (~spl16_20 | ~spl16_100)),
  inference(resolution,[],[f900,f201])).
fof(f201,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK9(X0),X0)) ) | ~spl16_20),
  inference(avatar_component_clause,[],[f200])).
fof(f200,plain,(
  spl16_20 <=> ! [X0] : (iof(sK9(X0),X0) | ~type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_20])])).
fof(f900,plain,(
  type_(sK9(sK1)) | ~spl16_100),
  inference(avatar_component_clause,[],[f899])).
fof(f899,plain,(
  spl16_100 <=> type_(sK9(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_100])])).
fof(f1116,plain,(
  spl16_123 | ~spl16_54 | ~spl16_121),
  inference(avatar_split_clause,[],[f1107,f1100,f453,f1114])).
fof(f453,plain,(
  spl16_54 <=> iof(sK9(sK1),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_54])])).
fof(f1100,plain,(
  spl16_121 <=> ! [X0] : (individual(X0) | ~iof(X0,sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_121])])).
fof(f1107,plain,(
  individual(sK9(sK1)) | (~spl16_54 | ~spl16_121)),
  inference(resolution,[],[f1101,f454])).
fof(f454,plain,(
  iof(sK9(sK1),sK1) | ~spl16_54),
  inference(avatar_component_clause,[],[f453])).
fof(f1101,plain,(
  ( ! [X0] : (~iof(X0,sK1) | individual(X0)) ) | ~spl16_121),
  inference(avatar_component_clause,[],[f1100])).
fof(f1102,plain,(
  spl16_121 | ~spl16_40 | ~spl16_85),
  inference(avatar_split_clause,[],[f821,f818,f371,f1100])).
fof(f371,plain,(
  spl16_40 <=> ! [X1,X0] : (~iof(X1,X0) | individual(X1) | ~firstordertype(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_40])])).
fof(f818,plain,(
  spl16_85 <=> firstordertype(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_85])])).
fof(f821,plain,(
  ( ! [X0] : (individual(X0) | ~iof(X0,sK1)) ) | (~spl16_40 | ~spl16_85)),
  inference(resolution,[],[f819,f372])).
fof(f372,plain,(
  ( ! [X0,X1] : (~firstordertype(X0) | individual(X1) | ~iof(X1,X0)) ) | ~spl16_40),
  inference(avatar_component_clause,[],[f371])).
fof(f819,plain,(
  firstordertype(sK1) | ~spl16_85),
  inference(avatar_component_clause,[],[f818])).
fof(f901,plain,(
  spl16_100 | ~spl16_6 | ~spl16_98),
  inference(avatar_split_clause,[],[f893,f888,f144,f899])).
fof(f144,plain,(
  spl16_6 <=> ! [X0] : (type_(X0) | ~firstordertype(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_6])])).
fof(f888,plain,(
  spl16_98 <=> firstordertype(sK9(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_98])])).
fof(f893,plain,(
  type_(sK9(sK1)) | (~spl16_6 | ~spl16_98)),
  inference(resolution,[],[f889,f145])).
fof(f145,plain,(
  ( ! [X0] : (~firstordertype(X0) | type_(X0)) ) | ~spl16_6),
  inference(avatar_component_clause,[],[f144])).
fof(f889,plain,(
  firstordertype(sK9(sK1)) | ~spl16_98),
  inference(avatar_component_clause,[],[f888])).
fof(f890,plain,(
  spl16_98 | ~spl16_54 | ~spl16_96),
  inference(avatar_split_clause,[],[f885,f878,f453,f888])).
fof(f878,plain,(
  spl16_96 <=> ! [X1] : (firstordertype(X1) | ~iof(X1,sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_96])])).
fof(f885,plain,(
  firstordertype(sK9(sK1)) | (~spl16_54 | ~spl16_96)),
  inference(resolution,[],[f879,f454])).
fof(f879,plain,(
  ( ! [X1] : (~iof(X1,sK1) | firstordertype(X1)) ) | ~spl16_96),
  inference(avatar_component_clause,[],[f878])).
fof(f880,plain,(
  spl16_96 | ~spl16_28 | ~spl16_32),
  inference(avatar_split_clause,[],[f462,f266,f247,f878])).
fof(f247,plain,(
  spl16_28 <=> secondordertype(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_28])])).
fof(f266,plain,(
  spl16_32 <=> ! [X1,X0] : (~iof(X1,X0) | firstordertype(X1) | ~secondordertype(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_32])])).
fof(f462,plain,(
  ( ! [X1] : (firstordertype(X1) | ~iof(X1,sK1)) ) | (~spl16_28 | ~spl16_32)),
  inference(resolution,[],[f267,f248])).
fof(f248,plain,(
  secondordertype(sK1) | ~spl16_28),
  inference(avatar_component_clause,[],[f247])).
fof(f267,plain,(
  ( ! [X0,X1] : (~secondordertype(X0) | firstordertype(X1) | ~iof(X1,X0)) ) | ~spl16_32),
  inference(avatar_component_clause,[],[f266])).
fof(f820,plain,(
  spl16_85 | ~spl16_3 | ~spl16_84),
  inference(avatar_split_clause,[],[f804,f801,f132,f818])).
fof(f132,plain,(
  spl16_3 <=> iof(sK1,sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_3])])).
fof(f801,plain,(
  spl16_84 <=> ! [X0] : (firstordertype(X0) | ~iof(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_84])])).
fof(f804,plain,(
  firstordertype(sK1) | (~spl16_3 | ~spl16_84)),
  inference(resolution,[],[f802,f133])).
fof(f133,plain,(
  iof(sK1,sK0) | ~spl16_3),
  inference(avatar_component_clause,[],[f132])).
fof(f802,plain,(
  ( ! [X0] : (~iof(X0,sK0) | firstordertype(X0)) ) | ~spl16_84),
  inference(avatar_component_clause,[],[f801])).
fof(f803,plain,(
  spl16_84 | ~spl16_27 | ~spl16_32),
  inference(avatar_split_clause,[],[f461,f266,f241,f801])).
fof(f241,plain,(
  spl16_27 <=> secondordertype(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_27])])).
fof(f461,plain,(
  ( ! [X0] : (firstordertype(X0) | ~iof(X0,sK0)) ) | (~spl16_27 | ~spl16_32)),
  inference(resolution,[],[f267,f242])).
fof(f242,plain,(
  secondordertype(sK0) | ~spl16_27),
  inference(avatar_component_clause,[],[f241])).
fof(f455,plain,(
  spl16_54 | ~spl16_20 | ~spl16_30),
  inference(avatar_split_clause,[],[f405,f257,f200,f453])).
fof(f257,plain,(
  spl16_30 <=> type_(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_30])])).
fof(f405,plain,(
  iof(sK9(sK1),sK1) | (~spl16_20 | ~spl16_30)),
  inference(resolution,[],[f201,f258])).
fof(f258,plain,(
  type_(sK1) | ~spl16_30),
  inference(avatar_component_clause,[],[f257])).
fof(f373,plain,(
  spl16_40),
  inference(avatar_split_clause,[],[f76,f371])).
fof(f76,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | individual(X1) | ~firstordertype(X0)) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (individual(X1) | ~iof(X1,X0)) & type_(X0)))),
  inference(ennf_transformation,[],[f10])).
fof(f10,axiom,(
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (iof(X1,X0) => individual(X1)) & type_(X0)))),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a5_firstOrderTypeDef)).
fof(f268,plain,(
  spl16_32),
  inference(avatar_split_clause,[],[f72,f266])).
fof(f72,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | firstordertype(X1) | ~secondordertype(X0)) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (secondordertype(X0) <=> (! [X1] : (firstordertype(X1) | ~iof(X1,X0)) & type_(X0)))),
  inference(ennf_transformation,[],[f11])).
fof(f11,axiom,(
  ! [X0] : (secondordertype(X0) <=> (! [X1] : (iof(X1,X0) => firstordertype(X1)) & type_(X0)))),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a6_secondOrderTypeDef)).
fof(f259,plain,(
  spl16_30 | ~spl16_5 | ~spl16_28),
  inference(avatar_split_clause,[],[f251,f247,f140,f257])).
fof(f140,plain,(
  spl16_5 <=> ! [X0] : (type_(X0) | ~secondordertype(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_5])])).
fof(f251,plain,(
  type_(sK1) | (~spl16_5 | ~spl16_28)),
  inference(resolution,[],[f248,f141])).
fof(f141,plain,(
  ( ! [X0] : (~secondordertype(X0) | type_(X0)) ) | ~spl16_5),
  inference(avatar_component_clause,[],[f140])).
fof(f249,plain,(
  spl16_28 | ~spl16_2 | ~spl16_16),
  inference(avatar_split_clause,[],[f239,f184,f128,f247])).
fof(f128,plain,(
  spl16_2 <=> iof(sK1,c_sot)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_2])])).
fof(f184,plain,(
  spl16_16 <=> ! [X1] : (secondordertype(X1) | ~iof(X1,c_sot))),
  introduced(avatar_definition,[new_symbols(naming,[spl16_16])])).
fof(f239,plain,(
  secondordertype(sK1) | (~spl16_2 | ~spl16_16)),
  inference(resolution,[],[f185,f129])).
fof(f129,plain,(
  iof(sK1,c_sot) | ~spl16_2),
  inference(avatar_component_clause,[],[f128])).
fof(f185,plain,(
  ( ! [X1] : (~iof(X1,c_sot) | secondordertype(X1)) ) | ~spl16_16),
  inference(avatar_component_clause,[],[f184])).
fof(f243,plain,(
  spl16_27 | ~spl16_1 | ~spl16_16),
  inference(avatar_split_clause,[],[f238,f184,f124,f241])).
fof(f124,plain,(
  spl16_1 <=> iof(sK0,c_sot)),
  introduced(avatar_definition,[new_symbols(naming,[spl16_1])])).
fof(f238,plain,(
  secondordertype(sK0) | (~spl16_1 | ~spl16_16)),
  inference(resolution,[],[f185,f125])).
fof(f125,plain,(
  iof(sK0,c_sot) | ~spl16_1),
  inference(avatar_component_clause,[],[f124])).
fof(f202,plain,(
  spl16_20),
  inference(avatar_split_clause,[],[f81,f200])).
fof(f81,plain,(
  ( ! [X0] : (iof(sK9(X0),X0) | ~type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a3)).
fof(f186,plain,(
  spl16_16),
  inference(avatar_split_clause,[],[f112,f184])).
fof(f112,plain,(
  ( ! [X1] : (secondordertype(X1) | ~iof(X1,c_sot)) )),
  inference(equality_resolution,[],[f91])).
fof(f91,plain,(
  ( ! [X0,X1] : (secondordertype(X1) | ~iof(X1,X0) | c_sot != X0) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : (c_sot = X0 <=> ! [X1] : (iof(X1,X0) <=> secondordertype(X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10] : (c_sot = X10 <=> ! [X0] : (iof(X0,X10) <=> secondordertype(X0)))),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a15_sotConstantDef)).
fof(f154,plain,(
  spl16_8),
  inference(avatar_split_clause,[],[f79,f152])).
fof(f79,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | ~individual(X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (individual(X0) <=> ! [X1] : ~iof(X1,X0))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0] : (individual(X0) <=> ~? [X1] : iof(X1,X0))),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',a2)).
fof(f146,plain,(
  spl16_6),
  inference(avatar_split_clause,[],[f77,f144])).
fof(f77,plain,(
  ( ! [X0] : (type_(X0) | ~firstordertype(X0)) )),
  inference(cnf_transformation,[],[f48])).
fof(f142,plain,(
  spl16_5),
  inference(avatar_split_clause,[],[f73,f140])).
fof(f73,plain,(
  ( ! [X0] : (type_(X0) | ~secondordertype(X0)) )),
  inference(cnf_transformation,[],[f47])).
fof(f134,plain,(
  spl16_3),
  inference(avatar_split_clause,[],[f59,f132])).
fof(f59,plain,(
  iof(sK1,sK0)),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ? [X0,X1] : (iof(X1,X0) & iof(X1,c_sot) & iof(X0,c_sot))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ~~? [X0,X1] : (iof(X1,X0) & iof(X1,c_sot) & iof(X0,c_sot))),
  inference(rectify,[],[f25])).
fof(f25,negated_conjecture,(
  ~~? [X6,X7] : (iof(X7,X6) & iof(X7,c_sot) & iof(X6,c_sot))),
  inference(negated_conjecture,[],[f24])).
fof(f24,conjecture,(
  ~? [X6,X7] : (iof(X7,X6) & iof(X7,c_sot) & iof(X6,c_sot))),
  file('/tmp/WT2askIkzH/SOT_qBPqo3',t8_ap2)).
fof(f130,plain,(
  spl16_2),
  inference(avatar_split_clause,[],[f58,f128])).
fof(f58,plain,(
  iof(sK1,c_sot)),
  inference(cnf_transformation,[],[f27])).
fof(f126,plain,(
  spl16_1),
  inference(avatar_split_clause,[],[f57,f124])).
fof(f57,plain,(
  iof(sK0,c_sot)),
  inference(cnf_transformation,[],[f27])).
% SZS output end Proof for SOT_qBPq
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 5884
% Time elapsed: 0.019 s
% ------------------------------
% ------------------------------
% Success in time 0.061 s
start to proof:SOT_qBPqo3

% END OF SYSTEM OUTPUT

% END OF SYSTEM OUTPUT

% END OF SYSTEM OUTPUT

% END OF SYSTEM OUTPUT

% END OF SYSTEM OUTPUT
% RESULT: SOT_qBPqo3 - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.07 
% OUTPUT: SOT_qBPqo3 - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.08 
% RESULT: SOT_qBPqo3 - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.06 
% OUTPUT: SOT_qBPqo3 - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.06 
% RESULT: SOT_qBPqo3 - E---FNT-2.5 says Theorem - CPU = 0.00 WC = 0.02 
% OUTPUT: SOT_qBPqo3 - E---FNT-2.5 says CNFRefutation - CPU = 0.00 WC = 0.02 
% RESULT: SOT_qBPqo3 - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.06 
% OUTPUT: SOT_qBPqo3 - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.06 
% RESULT: SOT_qBPqo3 - E---FNT-2.5 says Theorem - CPU = 0.00 WC = 0.02 
% OUTPUT: SOT_qBPqo3 - E---FNT-2.5 says CNFRefutation - CPU = 0.00 WC = 0.03 

% END OF SYSTEM OUTPUT

% END OF SYSTEM OUTPUT
% RESULT: SOT_qBPqo3 - E---FNT-2.5 says Theorem - CPU = 0.00 WC = 0.02 
% OUTPUT: SOT_qBPqo3 - E---FNT-2.5 says CNFRefutation - CPU = 0.00 WC = 0.02 
% RESULT: SOT_qBPqo3 - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.06 
% OUTPUT: SOT_qBPqo3 - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.06 

% END OF SYSTEM OUTPUT
% RESULT: SOT_qBPqo3 - CSE_E---1.0 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_qBPqo3 - CSE_E---1.0 says None - CPU = 0 WC = 0 
% 10 wall clock seconds gone, waiting for system output
% 20 wall clock seconds gone, waiting for system output
% 30 wall clock seconds gone, waiting for system output
% 40 wall clock seconds gone, waiting for system output

% END OF SYSTEM OUTPUT
% RESULT: SOT_qBPqo3 - Isabelle---2016 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_qBPqo3 - Isabelle---2016 says None - CPU = 0 WC = 0 

% END OF SYSTEM OUTPUT
% RESULT: SOT_qBPqo3 - Isabelle---2016 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_qBPqo3 - Isabelle---2016 says None - CPU = 0 WC = 0 
Vampire---SAT-4.5 has ended (PID 105525)
Vampire---SAT-4.5 has ended (PID 105526)
Vampire---SAT-4.5 has ended (PID 105527)
Vampire---SAT-4.5 has ended (PID 105528)
E---FNT-2.5 has ended (PID 105529)
E---FNT-2.5 has ended (PID 105530)
CSE_E---1.0 has ended (PID 105531)
E---FNT-2.5 has ended (PID 105532)
Isabelle---2016 has ended (PID 105533)
Isabelle---2016 has ended (PID 105534)
RESULT: SOT_qBPqo3 - SSCPA---0.0 says Theorem - CPU = 213.8 WC = 45.3  SolvingSystem = Vampire---SAT-4.5
</PRE>

</body>
</html>