<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>System Results</title>
<meta name="ROBOTS" content="INDEX,NOFOLLOW" />
<script type="text/javascript">//<![CDATA[

  var tstp_dump;
  function openSoTSTP (dump) {   
    var tstp_url = 'http://www.tptp.org/cgi-bin/SystemOnTSTP';
    var tstp_browser = window.open(tstp_url, '_blank');
    tstp_dump = dump;
  }
  function getTSTPDump () {
    return tstp_dump;
  }

//]]></script>
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>

% SZS start RequiredInformation<BR>
% Congratulations - you have become a registered power user of SystemOnTPTP, at IP address 191.6.38.118.<BR>
% Please consider donating to the TPTP project - see www.tptp.org for details.<BR>
% When you donate this message will disappear.<BR>
% If you do not donate a random delay might be added to your processing time.<BR>
% SZS end RequiredInformation<BR>
<PRE>
Getting system recommendations for generic problem /tmp/wqEVxBSTg1/SOT_pdlXwd
This is a FOF_THM_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.95
Vampire---TF1-4.5    recommended  0.95
Vampire---4.5        recommended  0.95
Vampire---SAT-4.5    recommended  0.95
ET---2.0             recommended  0.80
E---2.5              recommended  0.78
E---FNT-2.5          recommended  0.78
CSE_E---1.0          recommended  0.76
E---FNT-2.5          recommended  0.74
Isabelle---2TH0      recommended  0.69
Isabelle---2TF0      recommended  0.69
Isabelle---2FOF      recommended  0.69
Isabelle---2016      recommended  0.69
Isabelle-HOT---2016  recommended  0.68
CVC4---1.7           recommended  0.68
CVC4---SAT-1.7       recommended  0.68
Leo-III---1.4        recommended  0.61
Vampire---THF-4.5    recommended  0.59
Vampire---TF1-4.5    recommended  0.59
Vampire---SAT-4.5    recommended  0.59
iProver---3.0        recommended  0.58
iProver---SAT-3.0    recommended  0.58
Fampire---1.3        recommended  0.57
SInE---0.4           recommended  0.53
iProver---SAT-3.0    recommended  0.52
ePrincess---1.0      recommended  0.50
Princess---170717    recommended  0.48
SPASS---3.9          recommended  0.48
Z3---4.4.1           recommended  0.46
Zipperpin---1.5      recommended  0.40
CVC4---SAT-1.7       recommended  0.40
Prover9---1109a      recommended  0.39
Equinox---6.0.1a     recommended  0.37
LEO-II---1.7.0       recommended  0.35
nanoCoP---1.1        recommended  0.34
Muscadet---4.5       recommended  0.25
Twee---2.2           recommended  0.23
Metis---2.4          recommended  0.23
Geo-III---2018C      recommended  0.18
iProverMo---2.5-0.1  recommended  0.14
Beagle---0.9.47      subsumed     0.41
leanCoP---2.2        subsumed     0.36
E-Darwin---1.5       subsumed     0.31
CSE---1.0            subsumed     0.28
Darwin---1.4.5       subsumed     0.28
Bliksem---1.12       subsumed     0.21
Zenon---0.7.1        subsumed     0.16
ZenonModulo---0.4.2  subsumed     0.15
SOS---2.0            subsumed     0.11
Otter---3.3          subsumed     0.04
This is a FOF_CSA_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.98
Vampire---TF1-4.5    recommended  0.98
Vampire---SAT-4.5    recommended  0.98
Paradox---4.0        recommended  0.60
Nitpick---2016       subsumed     0.55
DarwinFM---1.4.5     subsumed     0.43
CVC4---SAT-1.7       subsumed     0.40
Vampire---THF-4.5    subsumed     0.34
Vampire---TF1-4.5    subsumed     0.34
Vampire---4.5        subsumed     0.34
Vampire---SAT-4.5    subsumed     0.34
E---2.5              subsumed     0.34
E---FNT-2.5          subsumed     0.34
E---FNT-2.5          subsumed     0.34
iProver---SAT-3.0    subsumed     0.32
Metis---2.4          subsumed     0.32
SPASS---3.9          subsumed     0.28
Beagle---0.9.47      subsumed     0.26
Geo-III---2018C      subsumed     0.25
Z3---4.4.1           subsumed     0.25
Princess---170717    subsumed     0.25
Equinox---6.0.1a     subsumed     0.19
iProver---3.0        subsumed     0.11
iProver---SAT-3.0    subsumed     0.11
Darwin---1.4.5       subsumed     0.06
E-Darwin---1.5       subsumed     0.04
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use ET,2.0
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use CSE_E,1.0
CSE_E---1.0 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use Isabelle,2TH0
Isabelle---2016 can be used
Check if we can use Isabelle,2TF0
Isabelle---2016 can be used
Check if we can use Isabelle,2FOF
Isabelle---2016 can be used
Check if we can use Isabelle,2016
Isabelle---2016 can be used
Check if we can use Isabelle-HOT,2016
Isabelle-HOT---2016 can be used
Check if we can use CVC4,1.7
CVC4---SAT-1.7 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Leo-III,1.4
Leo-III---1.4 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Fampire,1.3
Check if we can use SInE,0.4
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use ePrincess,1.0
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Zipperpin,1.5
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Prover9,1109a
Prover9---1109a can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use LEO-II,1.7.0
LEO-II---1.7.0 can be used
Check if we can use nanoCoP,1.1
nanoCoP---1.1 can be used
Check if we can use Muscadet,4.5
Muscadet---4.5 can be used
Check if we can use Twee,2.2
Twee---2.2 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use iProverMo,2.5-0.1
iProverMo---2.5-0.1 can be used
Check if we can use Beagle,0.9.47
Check if we can use leanCoP,2.2
leanCoP---2.2 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Check if we can use CSE,1.0
CSE---1.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use Bliksem,1.12
Bliksem---1.12 can be used
Check if we can use Zenon,0.7.1
Zenon---0.7.1 can be used
Check if we can use ZenonModulo,0.4.2
Check if we can use SOS,2.0
SOS---2.0 can be used
Check if we can use Otter,3.3
Otter---3.3 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use Paradox,4.0
Paradox---4.0 can be used
Check if we can use Nitpick,2016
Nitpick---2016 can be used
Check if we can use DarwinFM,1.4.5
DarwinFM---1.4.5 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Beagle,0.9.47
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
E---FNT-2.5 will be used
E---FNT-2.5 will be used
Vampire---SAT-4.5 is given 499 seconds in -cE mode
Vampire---SAT-4.5 is given 250 seconds in -cE mode
Vampire---SAT-4.5 is given 125 seconds in -cE mode
Vampire---SAT-4.5 is given 62 seconds in -cE mode
E---FNT-2.5 is given 31 seconds in -cE mode
E---FNT-2.5 is given 32 seconds in -cE mode
Vampire---SAT-4.5 has started (PID 97259)
Vampire---SAT-4.5 has started (PID 97263)
Vampire---SAT-4.5 has started (PID 97266)
Vampire---SAT-4.5 has started (PID 97268)
E---FNT-2.5 has started (PID 97270)
E---FNT-2.5 has started (PID 97273)
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  499  /tmp/wqEVxBSTg1/SOT_pdlXwd
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  250  /tmp/wqEVxBSTg1/SOT_pdlXwd
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  125  /tmp/wqEVxBSTg1/SOT_pdlXwd
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  62  /tmp/wqEVxBSTg1/SOT_pdlXwd
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  31  /tmp/wqEVxBSTg1/SOT_pdlXwd
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  32  /tmp/wqEVxBSTg1/SOT_pdlXwd
% START OF SYSTEM OUTPUT
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_pdlX
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_pdlX
% SZS output start Proof for SOT_pdlX
fof(f7063,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f111,f115,f119,f123,f139,f151,f155,f159,f167,f171,f176,f182,f236,f240,f244,f248,f252,f257,f261,f272,f289,f320,f397,f403,f413,f433,f513,f602,f611,f634,f645,f676,f958,f2691,f3707,f3725,f4811,f4826,f4831,f4847,f5004,f5006,f5125,f5827,f6100,f6109,f6627])).
fof(f6627,plain,(
  spl12_3 | ~spl12_30 | spl12_198 | ~spl12_199 | ~spl12_321),
  inference(avatar_contradiction_clause,[],[f6626])).
fof(f6626,plain,(
  $false | (spl12_3 | ~spl12_30 | spl12_198 | ~spl12_199 | ~spl12_321)),
  inference(subsumption_resolution,[],[f6269,f6255])).
fof(f6255,plain,(
  iof(sK8(sK1,sK1),sK1) | (spl12_3 | ~spl12_30 | ~spl12_199 | ~spl12_321)),
  inference(backward_demodulation,[],[f2690,f6123])).
fof(f6123,plain,(
  sK1 = sK2 | (spl12_3 | ~spl12_30 | ~spl12_321)),
  inference(subsumption_resolution,[],[f6120,f118])).
fof(f118,plain,(
  ~properSpecializes(sK2,sK1) | spl12_3),
  inference(avatar_component_clause,[],[f117])).
fof(f117,plain,(
  spl12_3 <=> properSpecializes(sK2,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_3])])).
fof(f6120,plain,(
  sK1 = sK2 | properSpecializes(sK2,sK1) | (~spl12_30 | ~spl12_321)),
  inference(resolution,[],[f6108,f239])).
fof(f239,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | X0 = X1 | properSpecializes(X0,X1)) ) | ~spl12_30),
  inference(avatar_component_clause,[],[f238])).
fof(f238,plain,(
  spl12_30 <=> ! [X1,X0] : (~specializes(X0,X1) | X0 = X1 | properSpecializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_30])])).
fof(f6108,plain,(
  specializes(sK2,sK1) | ~spl12_321),
  inference(avatar_component_clause,[],[f6107])).
fof(f6107,plain,(
  spl12_321 <=> specializes(sK2,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_321])])).
fof(f2690,plain,(
  iof(sK8(sK2,sK1),sK1) | ~spl12_199),
  inference(avatar_component_clause,[],[f2689])).
fof(f2689,plain,(
  spl12_199 <=> iof(sK8(sK2,sK1),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_199])])).
fof(f6269,plain,(
  ~iof(sK8(sK1,sK1),sK1) | (spl12_3 | ~spl12_30 | spl12_198 | ~spl12_321)),
  inference(backward_demodulation,[],[f3706,f6123])).
fof(f3706,plain,(
  ~iof(sK8(sK2,sK1),sK2) | spl12_198),
  inference(avatar_component_clause,[],[f2686])).
fof(f2686,plain,(
  spl12_198 <=> iof(sK8(sK2,sK1),sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_198])])).
fof(f6109,plain,(
  spl12_321 | ~spl12_18 | ~spl12_35 | ~spl12_320),
  inference(avatar_split_clause,[],[f6105,f6098,f259,f180,f6107])).
fof(f180,plain,(
  spl12_18 <=> type_(sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_18])])).
fof(f259,plain,(
  spl12_35 <=> ! [X1,X0] : (~type_(X0) | ~iof(sK11(X0,X1),X1) | specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_35])])).
fof(f6098,plain,(
  spl12_320 <=> ! [X0] : (specializes(sK2,X0) | iof(sK11(sK2,X0),sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_320])])).
fof(f6105,plain,(
  specializes(sK2,sK1) | (~spl12_18 | ~spl12_35 | ~spl12_320)),
  inference(subsumption_resolution,[],[f6104,f181])).
fof(f181,plain,(
  type_(sK2) | ~spl12_18),
  inference(avatar_component_clause,[],[f180])).
fof(f6104,plain,(
  specializes(sK2,sK1) | ~type_(sK2) | (~spl12_35 | ~spl12_320)),
  inference(duplicate_literal_removal,[],[f6101])).
fof(f6101,plain,(
  specializes(sK2,sK1) | ~type_(sK2) | specializes(sK2,sK1) | (~spl12_35 | ~spl12_320)),
  inference(resolution,[],[f6099,f260])).
fof(f260,plain,(
  ( ! [X0,X1] : (~iof(sK11(X0,X1),X1) | ~type_(X0) | specializes(X0,X1)) ) | ~spl12_35),
  inference(avatar_component_clause,[],[f259])).
fof(f6099,plain,(
  ( ! [X0] : (iof(sK11(sK2,X0),sK1) | specializes(sK2,X0)) ) | ~spl12_320),
  inference(avatar_component_clause,[],[f6098])).
fof(f6100,plain,(
  spl12_320 | ~spl12_1 | ~spl12_289),
  inference(avatar_split_clause,[],[f5828,f5825,f109,f6098])).
fof(f109,plain,(
  spl12_1 <=> isPowertypeOf(sK1,sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_1])])).
fof(f5825,plain,(
  spl12_289 <=> ! [X1,X0] : (specializes(sK2,X0) | ~isPowertypeOf(X1,sK0) | iof(sK11(sK2,X0),X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_289])])).
fof(f5828,plain,(
  ( ! [X0] : (specializes(sK2,X0) | iof(sK11(sK2,X0),sK1)) ) | (~spl12_1 | ~spl12_289)),
  inference(resolution,[],[f5826,f110])).
fof(f110,plain,(
  isPowertypeOf(sK1,sK0) | ~spl12_1),
  inference(avatar_component_clause,[],[f109])).
fof(f5826,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X1,sK0) | specializes(sK2,X0) | iof(sK11(sK2,X0),X1)) ) | ~spl12_289),
  inference(avatar_component_clause,[],[f5825])).
fof(f5827,plain,(
  spl12_289 | ~spl12_31 | ~spl12_99),
  inference(avatar_split_clause,[],[f646,f643,f242,f5825])).
fof(f242,plain,(
  spl12_31 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | ~specializes(X2,X1) | iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_31])])).
fof(f643,plain,(
  spl12_99 <=> ! [X0] : (specializes(sK2,X0) | specializes(sK11(sK2,X0),sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_99])])).
fof(f646,plain,(
  ( ! [X0,X1] : (specializes(sK2,X0) | ~isPowertypeOf(X1,sK0) | iof(sK11(sK2,X0),X1)) ) | (~spl12_31 | ~spl12_99)),
  inference(resolution,[],[f644,f243])).
fof(f243,plain,(
  ( ! [X2,X0,X1] : (~specializes(X2,X1) | ~isPowertypeOf(X0,X1) | iof(X2,X0)) ) | ~spl12_31),
  inference(avatar_component_clause,[],[f242])).
fof(f644,plain,(
  ( ! [X0] : (specializes(sK11(sK2,X0),sK0) | specializes(sK2,X0)) ) | ~spl12_99),
  inference(avatar_component_clause,[],[f643])).
fof(f5125,plain,(
  spl12_199 | ~spl12_1 | ~spl12_221),
  inference(avatar_split_clause,[],[f5124,f3723,f109,f2689])).
fof(f3723,plain,(
  spl12_221 <=> ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK8(sK2,sK1),X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_221])])).
fof(f5124,plain,(
  iof(sK8(sK2,sK1),sK1) | (~spl12_1 | ~spl12_221)),
  inference(resolution,[],[f3724,f110])).
fof(f3724,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK8(sK2,sK1),X0)) ) | ~spl12_221),
  inference(avatar_component_clause,[],[f3723])).
fof(f5006,plain,(
  spl12_200 | ~spl12_16 | ~spl12_229),
  inference(avatar_split_clause,[],[f5005,f5002,f169,f3388])).
fof(f3388,plain,(
  spl12_200 <=> specializes(sK8(sK2,sK1),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_200])])).
fof(f169,plain,(
  spl12_16 <=> ! [X1,X0] : (specializes(X0,X1) | ~properSpecializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_16])])).
fof(f5002,plain,(
  spl12_229 <=> properSpecializes(sK8(sK2,sK1),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_229])])).
fof(f5005,plain,(
  specializes(sK8(sK2,sK1),sK0) | (~spl12_16 | ~spl12_229)),
  inference(resolution,[],[f5003,f170])).
fof(f170,plain,(
  ( ! [X0,X1] : (~properSpecializes(X0,X1) | specializes(X0,X1)) ) | ~spl12_16),
  inference(avatar_component_clause,[],[f169])).
fof(f5003,plain,(
  properSpecializes(sK8(sK2,sK1),sK0) | ~spl12_229),
  inference(avatar_component_clause,[],[f5002])).
fof(f5004,plain,(
  spl12_229 | ~spl12_45 | ~spl12_198),
  inference(avatar_split_clause,[],[f4943,f2686,f318,f5002])).
fof(f318,plain,(
  spl12_45 <=> ! [X0] : (~iof(X0,sK2) | properSpecializes(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_45])])).
fof(f4943,plain,(
  properSpecializes(sK8(sK2,sK1),sK0) | (~spl12_45 | ~spl12_198)),
  inference(resolution,[],[f2687,f319])).
fof(f319,plain,(
  ( ! [X0] : (~iof(X0,sK2) | properSpecializes(X0,sK0)) ) | ~spl12_45),
  inference(avatar_component_clause,[],[f318])).
fof(f2687,plain,(
  iof(sK8(sK2,sK1),sK2) | ~spl12_198),
  inference(avatar_component_clause,[],[f2686])).
fof(f4847,plain,(
  spl12_88 | ~spl12_35 | ~spl12_63 | ~spl12_103),
  inference(avatar_split_clause,[],[f4827,f674,f411,f259,f557])).
fof(f557,plain,(
  spl12_88 <=> specializes(sK0,sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_88])])).
fof(f411,plain,(
  spl12_63 <=> type_(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_63])])).
fof(f674,plain,(
  spl12_103 <=> ! [X0] : (iof(sK11(sK0,X0),sK0) | specializes(sK0,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_103])])).
fof(f4827,plain,(
  specializes(sK0,sK0) | (~spl12_35 | ~spl12_63 | ~spl12_103)),
  inference(subsumption_resolution,[],[f679,f412])).
fof(f412,plain,(
  type_(sK0) | ~spl12_63),
  inference(avatar_component_clause,[],[f411])).
fof(f679,plain,(
  specializes(sK0,sK0) | ~type_(sK0) | (~spl12_35 | ~spl12_103)),
  inference(duplicate_literal_removal,[],[f677])).
fof(f677,plain,(
  specializes(sK0,sK0) | ~type_(sK0) | specializes(sK0,sK0) | (~spl12_35 | ~spl12_103)),
  inference(resolution,[],[f675,f260])).
fof(f675,plain,(
  ( ! [X0] : (iof(sK11(sK0,X0),sK0) | specializes(sK0,X0)) ) | ~spl12_103),
  inference(avatar_component_clause,[],[f674])).
fof(f4831,plain,(
  spl12_38 | ~spl12_15 | ~spl12_17 | ~spl12_37),
  inference(avatar_split_clause,[],[f3251,f270,f174,f165,f277])).
fof(f277,plain,(
  spl12_38 <=> specializes(sK7(sK1),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_38])])).
fof(f165,plain,(
  spl12_15 <=> ! [X0] : (iof(sK7(X0),X0) | ~type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_15])])).
fof(f174,plain,(
  spl12_17 <=> type_(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_17])])).
fof(f270,plain,(
  spl12_37 <=> ! [X0] : (specializes(X0,sK0) | ~iof(X0,sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_37])])).
fof(f3251,plain,(
  specializes(sK7(sK1),sK0) | (~spl12_15 | ~spl12_17 | ~spl12_37)),
  inference(subsumption_resolution,[],[f3250,f175])).
fof(f175,plain,(
  type_(sK1) | ~spl12_17),
  inference(avatar_component_clause,[],[f174])).
fof(f3250,plain,(
  specializes(sK7(sK1),sK0) | ~type_(sK1) | (~spl12_15 | ~spl12_37)),
  inference(resolution,[],[f271,f166])).
fof(f166,plain,(
  ( ! [X0] : (iof(sK7(X0),X0) | ~type_(X0)) ) | ~spl12_15),
  inference(avatar_component_clause,[],[f165])).
fof(f271,plain,(
  ( ! [X0] : (~iof(X0,sK1) | specializes(X0,sK0)) ) | ~spl12_37),
  inference(avatar_component_clause,[],[f270])).
fof(f4826,plain,(
  ~spl12_1 | ~spl12_4 | ~spl12_93 | ~spl12_227),
  inference(avatar_contradiction_clause,[],[f4825])).
fof(f4825,plain,(
  $false | (~spl12_1 | ~spl12_4 | ~spl12_93 | ~spl12_227)),
  inference(subsumption_resolution,[],[f612,f4815])).
fof(f4815,plain,(
  ~iof(sK0,sK1) | (~spl12_4 | ~spl12_227)),
  inference(resolution,[],[f4810,f122])).
fof(f122,plain,(
  ( ! [X1] : (~properSpecializes(X1,X1)) ) | ~spl12_4),
  inference(avatar_component_clause,[],[f121])).
fof(f121,plain,(
  spl12_4 <=> ! [X1] : ~properSpecializes(X1,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_4])])).
fof(f4810,plain,(
  ( ! [X0] : (properSpecializes(X0,sK0) | ~iof(X0,sK1)) ) | ~spl12_227),
  inference(avatar_component_clause,[],[f4809])).
fof(f4809,plain,(
  spl12_227 <=> ! [X0] : (~iof(X0,sK1) | properSpecializes(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_227])])).
fof(f612,plain,(
  iof(sK0,sK1) | (~spl12_1 | ~spl12_93)),
  inference(resolution,[],[f601,f110])).
fof(f601,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK0,X0)) ) | ~spl12_93),
  inference(avatar_component_clause,[],[f600])).
fof(f600,plain,(
  spl12_93 <=> ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK0,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_93])])).
fof(f4811,plain,(
  spl12_227 | ~spl12_45 | ~spl12_197),
  inference(avatar_split_clause,[],[f4411,f2683,f318,f4809])).
fof(f2683,plain,(
  spl12_197 <=> sK1 = sK2),
  introduced(avatar_definition,[new_symbols(naming,[spl12_197])])).
fof(f4411,plain,(
  ( ! [X0] : (~iof(X0,sK1) | properSpecializes(X0,sK0)) ) | (~spl12_45 | ~spl12_197)),
  inference(backward_demodulation,[],[f319,f2684])).
fof(f2684,plain,(
  sK1 = sK2 | ~spl12_197),
  inference(avatar_component_clause,[],[f2683])).
fof(f3725,plain,(
  spl12_221 | ~spl12_31 | ~spl12_200),
  inference(avatar_split_clause,[],[f3396,f3388,f242,f3723])).
fof(f3396,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK8(sK2,sK1),X0)) ) | (~spl12_31 | ~spl12_200)),
  inference(resolution,[],[f3389,f243])).
fof(f3389,plain,(
  specializes(sK8(sK2,sK1),sK0) | ~spl12_200),
  inference(avatar_component_clause,[],[f3388])).
fof(f3707,plain,(
  spl12_197 | ~spl12_198 | ~spl12_67 | ~spl12_199),
  inference(avatar_split_clause,[],[f2692,f2689,f431,f2686,f2683])).
fof(f431,plain,(
  spl12_67 <=> ! [X1,X0] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_67])])).
fof(f2692,plain,(
  ~iof(sK8(sK2,sK1),sK2) | sK1 = sK2 | (~spl12_67 | ~spl12_199)),
  inference(resolution,[],[f2690,f432])).
fof(f432,plain,(
  ( ! [X0,X1] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) ) | ~spl12_67),
  inference(avatar_component_clause,[],[f431])).
fof(f2691,plain,(
  spl12_197 | spl12_198 | spl12_199 | ~spl12_18 | ~spl12_138),
  inference(avatar_split_clause,[],[f966,f956,f180,f2689,f2686,f2683])).
fof(f956,plain,(
  spl12_138 <=> ! [X1] : (~type_(X1) | iof(sK8(X1,sK1),sK1) | iof(sK8(X1,sK1),X1) | sK1 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_138])])).
fof(f966,plain,(
  iof(sK8(sK2,sK1),sK1) | iof(sK8(sK2,sK1),sK2) | sK1 = sK2 | (~spl12_18 | ~spl12_138)),
  inference(resolution,[],[f957,f181])).
fof(f957,plain,(
  ( ! [X1] : (~type_(X1) | iof(sK8(X1,sK1),sK1) | iof(sK8(X1,sK1),X1) | sK1 = X1) ) | ~spl12_138),
  inference(avatar_component_clause,[],[f956])).
fof(f958,plain,(
  spl12_138 | ~spl12_17 | ~spl12_94),
  inference(avatar_split_clause,[],[f636,f609,f174,f956])).
fof(f609,plain,(
  spl12_94 <=> ! [X1,X0] : (~type_(X0) | ~type_(X1) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_94])])).
fof(f636,plain,(
  ( ! [X1] : (~type_(X1) | iof(sK8(X1,sK1),sK1) | iof(sK8(X1,sK1),X1) | sK1 = X1) ) | (~spl12_17 | ~spl12_94)),
  inference(resolution,[],[f610,f175])).
fof(f610,plain,(
  ( ! [X0,X1] : (~type_(X1) | ~type_(X0) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1) ) | ~spl12_94),
  inference(avatar_component_clause,[],[f609])).
fof(f676,plain,(
  spl12_103 | ~spl12_34 | ~spl12_63),
  inference(avatar_split_clause,[],[f414,f411,f255,f674])).
fof(f255,plain,(
  spl12_34 <=> ! [X1,X0] : (~type_(X0) | iof(sK11(X0,X1),X0) | specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_34])])).
fof(f414,plain,(
  ( ! [X0] : (iof(sK11(sK0,X0),sK0) | specializes(sK0,X0)) ) | (~spl12_34 | ~spl12_63)),
  inference(resolution,[],[f412,f256])).
fof(f256,plain,(
  ( ! [X0,X1] : (~type_(X0) | iof(sK11(X0,X1),X0) | specializes(X0,X1)) ) | ~spl12_34),
  inference(avatar_component_clause,[],[f255])).
fof(f645,plain,(
  spl12_99 | ~spl12_16 | ~spl12_98),
  inference(avatar_split_clause,[],[f641,f632,f169,f643])).
fof(f632,plain,(
  spl12_98 <=> ! [X0] : (specializes(sK2,X0) | properSpecializes(sK11(sK2,X0),sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_98])])).
fof(f641,plain,(
  ( ! [X0] : (specializes(sK2,X0) | specializes(sK11(sK2,X0),sK0)) ) | (~spl12_16 | ~spl12_98)),
  inference(resolution,[],[f633,f170])).
fof(f633,plain,(
  ( ! [X0] : (properSpecializes(sK11(sK2,X0),sK0) | specializes(sK2,X0)) ) | ~spl12_98),
  inference(avatar_component_clause,[],[f632])).
fof(f634,plain,(
  spl12_98 | ~spl12_45 | ~spl12_82),
  inference(avatar_split_clause,[],[f515,f511,f318,f632])).
fof(f511,plain,(
  spl12_82 <=> ! [X1] : (iof(sK11(sK2,X1),sK2) | specializes(sK2,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_82])])).
fof(f515,plain,(
  ( ! [X0] : (specializes(sK2,X0) | properSpecializes(sK11(sK2,X0),sK0)) ) | (~spl12_45 | ~spl12_82)),
  inference(resolution,[],[f512,f319])).
fof(f512,plain,(
  ( ! [X1] : (iof(sK11(sK2,X1),sK2) | specializes(sK2,X1)) ) | ~spl12_82),
  inference(avatar_component_clause,[],[f511])).
fof(f611,plain,(
  spl12_94),
  inference(avatar_split_clause,[],[f76,f609])).
fof(f76,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~type_(X1) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ((X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))) | ~type_(X1) | ~type_(X0))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : ((X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))) | (~type_(X1) | ~type_(X0)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : ((type_(X1) & type_(X0)) => (X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7] : ((type_(X7) & type_(X6)) => (X6 = X7 <=> ! [X0] : (iof(X0,X6) <=> iof(X0,X7))))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a9_equality)).
fof(f602,plain,(
  spl12_93 | ~spl12_31 | ~spl12_88),
  inference(avatar_split_clause,[],[f560,f557,f242,f600])).
fof(f560,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK0,X0)) ) | (~spl12_31 | ~spl12_88)),
  inference(resolution,[],[f558,f243])).
fof(f558,plain,(
  specializes(sK0,sK0) | ~spl12_88),
  inference(avatar_component_clause,[],[f557])).
fof(f513,plain,(
  spl12_82 | ~spl12_18 | ~spl12_34),
  inference(avatar_split_clause,[],[f285,f255,f180,f511])).
fof(f285,plain,(
  ( ! [X1] : (iof(sK11(sK2,X1),sK2) | specializes(sK2,X1)) ) | (~spl12_18 | ~spl12_34)),
  inference(resolution,[],[f256,f181])).
fof(f433,plain,(
  spl12_67),
  inference(avatar_split_clause,[],[f104,f431])).
fof(f104,plain,(
  ( ! [X0,X1] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(subsumption_resolution,[],[f103,f70])).
fof(f70,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a3)).
fof(f103,plain,(
  ( ! [X0,X1] : (~type_(X1) | ~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(subsumption_resolution,[],[f77,f70])).
fof(f77,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~type_(X1) | ~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f48])).
fof(f413,plain,(
  spl12_63 | ~spl12_8 | ~spl12_61),
  inference(avatar_split_clause,[],[f409,f401,f137,f411])).
fof(f137,plain,(
  spl12_8 <=> ! [X1,X0] : (~iof(X1,X0) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_8])])).
fof(f401,plain,(
  spl12_61 <=> iof(sK7(sK7(sK1)),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_61])])).
fof(f409,plain,(
  type_(sK0) | (~spl12_8 | ~spl12_61)),
  inference(resolution,[],[f402,f138])).
fof(f138,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) ) | ~spl12_8),
  inference(avatar_component_clause,[],[f137])).
fof(f402,plain,(
  iof(sK7(sK7(sK1)),sK0) | ~spl12_61),
  inference(avatar_component_clause,[],[f401])).
fof(f403,plain,(
  spl12_61 | ~spl12_15 | ~spl12_39 | ~spl12_60),
  inference(avatar_split_clause,[],[f399,f395,f287,f165,f401])).
fof(f287,plain,(
  spl12_39 <=> type_(sK7(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_39])])).
fof(f395,plain,(
  spl12_60 <=> ! [X1] : (iof(X1,sK0) | ~iof(X1,sK7(sK1)))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_60])])).
fof(f399,plain,(
  iof(sK7(sK7(sK1)),sK0) | (~spl12_15 | ~spl12_39 | ~spl12_60)),
  inference(subsumption_resolution,[],[f398,f288])).
fof(f288,plain,(
  type_(sK7(sK1)) | ~spl12_39),
  inference(avatar_component_clause,[],[f287])).
fof(f398,plain,(
  iof(sK7(sK7(sK1)),sK0) | ~type_(sK7(sK1)) | (~spl12_15 | ~spl12_60)),
  inference(resolution,[],[f396,f166])).
fof(f396,plain,(
  ( ! [X1] : (~iof(X1,sK7(sK1)) | iof(X1,sK0)) ) | ~spl12_60),
  inference(avatar_component_clause,[],[f395])).
fof(f397,plain,(
  spl12_60 | ~spl12_29 | ~spl12_38),
  inference(avatar_split_clause,[],[f282,f277,f234,f395])).
fof(f234,plain,(
  spl12_29 <=> ! [X1,X0,X2] : (~iof(X2,X0) | iof(X2,X1) | ~specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_29])])).
fof(f282,plain,(
  ( ! [X1] : (iof(X1,sK0) | ~iof(X1,sK7(sK1))) ) | (~spl12_29 | ~spl12_38)),
  inference(resolution,[],[f278,f235])).
fof(f235,plain,(
  ( ! [X2,X0,X1] : (~specializes(X0,X1) | iof(X2,X1) | ~iof(X2,X0)) ) | ~spl12_29),
  inference(avatar_component_clause,[],[f234])).
fof(f278,plain,(
  specializes(sK7(sK1),sK0) | ~spl12_38),
  inference(avatar_component_clause,[],[f277])).
fof(f320,plain,(
  spl12_45 | ~spl12_2 | ~spl12_33),
  inference(avatar_split_clause,[],[f273,f250,f113,f318])).
fof(f113,plain,(
  spl12_2 <=> categorizes(sK2,sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_2])])).
fof(f250,plain,(
  spl12_33 <=> ! [X1,X0,X2] : (~categorizes(X0,X1) | ~iof(X2,X0) | properSpecializes(X2,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_33])])).
fof(f273,plain,(
  ( ! [X0] : (~iof(X0,sK2) | properSpecializes(X0,sK0)) ) | (~spl12_2 | ~spl12_33)),
  inference(resolution,[],[f251,f114])).
fof(f114,plain,(
  categorizes(sK2,sK0) | ~spl12_2),
  inference(avatar_component_clause,[],[f113])).
fof(f251,plain,(
  ( ! [X2,X0,X1] : (~categorizes(X0,X1) | ~iof(X2,X0) | properSpecializes(X2,X1)) ) | ~spl12_33),
  inference(avatar_component_clause,[],[f250])).
fof(f289,plain,(
  spl12_39 | ~spl12_11 | ~spl12_38),
  inference(avatar_split_clause,[],[f283,f277,f149,f287])).
fof(f149,plain,(
  spl12_11 <=> ! [X1,X0] : (type_(X0) | ~specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_11])])).
fof(f283,plain,(
  type_(sK7(sK1)) | (~spl12_11 | ~spl12_38)),
  inference(resolution,[],[f278,f150])).
fof(f150,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) ) | ~spl12_11),
  inference(avatar_component_clause,[],[f149])).
fof(f272,plain,(
  spl12_37 | ~spl12_1 | ~spl12_32),
  inference(avatar_split_clause,[],[f268,f246,f109,f270])).
fof(f246,plain,(
  spl12_32 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_32])])).
fof(f268,plain,(
  ( ! [X0] : (specializes(X0,sK0) | ~iof(X0,sK1)) ) | (~spl12_1 | ~spl12_32)),
  inference(resolution,[],[f247,f110])).
fof(f247,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) ) | ~spl12_32),
  inference(avatar_component_clause,[],[f246])).
fof(f261,plain,(
  spl12_35),
  inference(avatar_split_clause,[],[f87,f259])).
fof(f87,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~iof(sK11(X0,X1),X1) | specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X1) | ~iof(X2,X0)) & type_(X0)))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X0) => iof(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7] : (specializes(X6,X7) <=> (! [X3] : (iof(X3,X6) => iof(X3,X7)) & type_(X6)))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a7_specDef)).
fof(f257,plain,(
  spl12_34),
  inference(avatar_split_clause,[],[f86,f255])).
fof(f86,plain,(
  ( ! [X0,X1] : (~type_(X0) | iof(sK11(X0,X1),X0) | specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f53])).
fof(f252,plain,(
  spl12_33),
  inference(avatar_split_clause,[],[f96,f250])).
fof(f96,plain,(
  ( ! [X2,X0,X1] : (~categorizes(X0,X1) | ~iof(X2,X0) | properSpecializes(X2,X1)) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : ((! [X2] : (properSpecializes(X2,X1) | ~iof(X2,X0)) & type_(X0)) | ~categorizes(X0,X1))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : (categorizes(X0,X1) => (! [X2] : (iof(X2,X0) => properSpecializes(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (categorizes(X0,X1) <=> (! [X2] : (iof(X2,X0) => properSpecializes(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6,X7] : (categorizes(X6,X7) <=> (! [X8] : (iof(X8,X6) => properSpecializes(X8,X7)) & type_(X6)))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a11_categorizationDef)).
fof(f248,plain,(
  spl12_32),
  inference(avatar_split_clause,[],[f94,f246])).
fof(f94,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : ((! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)) | ~isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) => (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) <=> (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : (isPowertypeOf(X6,X7) <=> (! [X8] : (iof(X8,X6) <=> specializes(X8,X7)) & type_(X6)))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a10_isPowertypeOfDef)).
fof(f244,plain,(
  spl12_31),
  inference(avatar_split_clause,[],[f93,f242])).
fof(f93,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | ~specializes(X2,X1) | iof(X2,X0)) )),
  inference(cnf_transformation,[],[f54])).
fof(f240,plain,(
  spl12_30),
  inference(avatar_split_clause,[],[f92,f238])).
fof(f92,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | X0 = X1 | properSpecializes(X0,X1)) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1] : (properSpecializes(X0,X1) <=> (X0 != X1 & specializes(X0,X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X7] : (properSpecializes(X6,X7) <=> (X6 != X7 & specializes(X6,X7)))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a8_properSpecDef)).
fof(f236,plain,(
  spl12_29),
  inference(avatar_split_clause,[],[f88,f234])).
fof(f88,plain,(
  ( ! [X2,X0,X1] : (~iof(X2,X0) | iof(X2,X1) | ~specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f53])).
fof(f182,plain,(
  spl12_18 | ~spl12_2 | ~spl12_13),
  inference(avatar_split_clause,[],[f177,f157,f113,f180])).
fof(f157,plain,(
  spl12_13 <=> ! [X1,X0] : (~categorizes(X0,X1) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_13])])).
fof(f177,plain,(
  type_(sK2) | (~spl12_2 | ~spl12_13)),
  inference(resolution,[],[f158,f114])).
fof(f158,plain,(
  ( ! [X0,X1] : (~categorizes(X0,X1) | type_(X0)) ) | ~spl12_13),
  inference(avatar_component_clause,[],[f157])).
fof(f176,plain,(
  spl12_17 | ~spl12_1 | ~spl12_12),
  inference(avatar_split_clause,[],[f172,f153,f109,f174])).
fof(f153,plain,(
  spl12_12 <=> ! [X1,X0] : (~isPowertypeOf(X0,X1) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_12])])).
fof(f172,plain,(
  type_(sK1) | (~spl12_1 | ~spl12_12)),
  inference(resolution,[],[f154,f110])).
fof(f154,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X0,X1) | type_(X0)) ) | ~spl12_12),
  inference(avatar_component_clause,[],[f153])).
fof(f171,plain,(
  spl12_16),
  inference(avatar_split_clause,[],[f90,f169])).
fof(f90,plain,(
  ( ! [X0,X1] : (specializes(X0,X1) | ~properSpecializes(X0,X1)) )),
  inference(cnf_transformation,[],[f32])).
fof(f167,plain,(
  spl12_15),
  inference(avatar_split_clause,[],[f71,f165])).
fof(f71,plain,(
  ( ! [X0] : (iof(sK7(X0),X0) | ~type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f159,plain,(
  spl12_13),
  inference(avatar_split_clause,[],[f97,f157])).
fof(f97,plain,(
  ( ! [X0,X1] : (~categorizes(X0,X1) | type_(X0)) )),
  inference(cnf_transformation,[],[f55])).
fof(f155,plain,(
  spl12_12),
  inference(avatar_split_clause,[],[f95,f153])).
fof(f95,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X0,X1) | type_(X0)) )),
  inference(cnf_transformation,[],[f54])).
fof(f151,plain,(
  spl12_11),
  inference(avatar_split_clause,[],[f89,f149])).
fof(f89,plain,(
  ( ! [X0,X1] : (type_(X0) | ~specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f53])).
fof(f139,plain,(
  spl12_8),
  inference(avatar_split_clause,[],[f70,f137])).
fof(f123,plain,(
  spl12_4),
  inference(avatar_split_clause,[],[f102,f121])).
fof(f102,plain,(
  ( ! [X1] : (~properSpecializes(X1,X1)) )),
  inference(equality_resolution,[],[f91])).
fof(f91,plain,(
  ( ! [X0,X1] : (X0 != X1 | ~properSpecializes(X0,X1)) )),
  inference(cnf_transformation,[],[f32])).
fof(f119,plain,(
  ~spl12_3),
  inference(avatar_split_clause,[],[f58,f117])).
fof(f58,plain,(
  ~properSpecializes(sK2,sK1)),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ? [X0,X1,X2] : (~properSpecializes(X2,X1) & categorizes(X2,X0) & isPowertypeOf(X1,X0))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ? [X0,X1,X2] : (~properSpecializes(X2,X1) & (categorizes(X2,X0) & isPowertypeOf(X1,X0)))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1,X2] : ((categorizes(X2,X0) & isPowertypeOf(X1,X0)) => properSpecializes(X2,X1))),
  inference(rectify,[],[f22])).
fof(f22,negated_conjecture,(
  ~! [X6,X7,X8] : ((categorizes(X8,X6) & isPowertypeOf(X7,X6)) => properSpecializes(X8,X7))),
  inference(negated_conjecture,[],[f21])).
fof(f21,conjecture,(
  ! [X6,X7,X8] : ((categorizes(X8,X6) & isPowertypeOf(X7,X6)) => properSpecializes(X8,X7))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',t4)).
fof(f115,plain,(
  spl12_2),
  inference(avatar_split_clause,[],[f57,f113])).
fof(f57,plain,(
  categorizes(sK2,sK0)),
  inference(cnf_transformation,[],[f40])).
fof(f111,plain,(
  spl12_1),
  inference(avatar_split_clause,[],[f56,f109])).
fof(f56,plain,(
  isPowertypeOf(sK1,sK0)),
  inference(cnf_transformation,[],[f40])).
% SZS output end Proof for SOT_pdlX
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 9850
% Time elapsed: 0.231 s
% ------------------------------
% ------------------------------
% Success in time 0.268 s
% START OF SYSTEM OUTPUT
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_pdlX
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 15 seconds
# AutoSched0-Mode selected heuristic G_E___208_C02CMA_F1_SE_CS_SP_PS_S5PRR_RG_S04AN
# and selection function SelectComplexExceptUniqMaxHorn.
#
# Preprocessing time       : 0.024 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4, conjecture, ![X7, X8, X9]:((isPowertypeOf(X8,X7)&categorizes(X9,X7))=>properSpecializes(X9,X8)), file('/tmp/wqEVxBSTg1/SOT_pdlXwd', t4)).
fof(a11_categorizationDef, axiom, ![X7, X8]:(categorizes(X7,X8)<=>(type_(X7)&![X9]:(iof(X9,X7)=>properSpecializes(X9,X8)))), file('/tmp/wqEVxBSTg1/SOT_pdlXwd', a11_categorizationDef)).
fof(a8_properSpecDef, axiom, ![X7, X8]:(properSpecializes(X7,X8)<=>(specializes(X7,X8)&~(X7=X8))), file('/tmp/wqEVxBSTg1/SOT_pdlXwd', a8_properSpecDef)).
fof(a7_specDef, axiom, ![X7, X8]:(specializes(X7,X8)<=>(type_(X7)&![X4]:(iof(X4,X7)=>iof(X4,X8)))), file('/tmp/wqEVxBSTg1/SOT_pdlXwd', a7_specDef)).
fof(a2, axiom, ![X1]:(individual(X1)<=>~(?[X2]:iof(X2,X1))), file('/tmp/wqEVxBSTg1/SOT_pdlXwd', a2)).
fof(a3, axiom, ![X1]:(type_(X1)<=>?[X2]:iof(X2,X1)), file('/tmp/wqEVxBSTg1/SOT_pdlXwd', a3)).
fof(a5_firstOrderTypeDef, axiom, ![X1]:(firstordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>individual(X2)))), file('/tmp/wqEVxBSTg1/SOT_pdlXwd', a5_firstOrderTypeDef)).
fof(a10_isPowertypeOfDef, axiom, ![X7, X8]:(isPowertypeOf(X7,X8)<=>(type_(X7)&![X9]:(iof(X9,X7)<=>specializes(X9,X8)))), file('/tmp/wqEVxBSTg1/SOT_pdlXwd', a10_isPowertypeOfDef)).
fof(c_0_8, negated_conjecture, ~(![X7, X8, X9]:((isPowertypeOf(X8,X7)&categorizes(X9,X7))=>properSpecializes(X9,X8))), inference(assume_negation,[status(cth)],[t4])).
fof(c_0_9, plain, ![X69, X70, X71, X72, X73]:(((type_(X69)|~categorizes(X69,X70))&(~iof(X71,X69)|properSpecializes(X71,X70)|~categorizes(X69,X70)))&((iof(esk14_2(X72,X73),X72)|~type_(X72)|categorizes(X72,X73))&(~properSpecializes(esk14_2(X72,X73),X73)|~type_(X72)|categorizes(X72,X73)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a11_categorizationDef])])])])])])).
fof(c_0_10, negated_conjecture, ((isPowertypeOf(esk23_0,esk22_0)&categorizes(esk24_0,esk22_0))&~properSpecializes(esk24_0,esk23_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
fof(c_0_11, plain, ![X55, X56]:(((specializes(X55,X56)|~properSpecializes(X55,X56))&(X55!=X56|~properSpecializes(X55,X56)))&(~specializes(X55,X56)|X55=X56|properSpecializes(X55,X56))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[a8_properSpecDef])])])])).
cnf(c_0_12, plain, (properSpecializes(X1,X3)|~iof(X1,X2)|~categorizes(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, negated_conjecture, (categorizes(esk24_0,esk22_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, plain, ![X49, X50, X51, X52, X53]:(((type_(X49)|~specializes(X49,X50))&(~iof(X51,X49)|iof(X51,X50)|~specializes(X49,X50)))&((iof(esk11_2(X52,X53),X52)|~type_(X52)|specializes(X52,X53))&(~iof(esk11_2(X52,X53),X53)|~type_(X52)|specializes(X52,X53)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a7_specDef])])])])])])).
cnf(c_0_15, plain, (specializes(X1,X2)|~properSpecializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, negated_conjecture, (properSpecializes(X1,esk22_0)|~iof(X1,esk24_0)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
fof(c_0_17, plain, ![X12, X13, X14]:((~individual(X12)|~iof(X13,X12))&(iof(esk1_1(X14),X14)|individual(X14))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a2])])])])])).
fof(c_0_18, plain, ![X16, X18, X19]:((~type_(X16)|iof(esk2_1(X16),X16))&(~iof(X19,X18)|type_(X18))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a3])])])])])).
cnf(c_0_19, plain, (type_(X1)|~specializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, negated_conjecture, (specializes(X1,esk22_0)|~iof(X1,esk24_0)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_21, plain, (type_(X1)|~categorizes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_22, plain, (~individual(X1)|~iof(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (iof(esk2_1(X1),X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, negated_conjecture, (type_(X1)|~iof(X1,esk24_0)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_25, negated_conjecture, (type_(esk24_0)), inference(spm,[status(thm)],[c_0_21, c_0_13])).
fof(c_0_26, plain, ![X41, X42, X43]:(((type_(X41)|~firstordertype(X41))&(~iof(X42,X41)|individual(X42)|~firstordertype(X41)))&((iof(esk9_1(X43),X43)|~type_(X43)|firstordertype(X43))&(~individual(esk9_1(X43))|~type_(X43)|firstordertype(X43)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a5_firstOrderTypeDef])])])])])])).
fof(c_0_27, plain, ![X62, X63, X64, X65, X66, X67]:(((type_(X62)|~isPowertypeOf(X62,X63))&((~iof(X64,X62)|specializes(X64,X63)|~isPowertypeOf(X62,X63))&(~specializes(X65,X63)|iof(X65,X62)|~isPowertypeOf(X62,X63))))&((~iof(esk13_2(X66,X67),X66)|~specializes(esk13_2(X66,X67),X67)|~type_(X66)|isPowertypeOf(X66,X67))&(iof(esk13_2(X66,X67),X66)|specializes(esk13_2(X66,X67),X67)|~type_(X66)|isPowertypeOf(X66,X67)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a10_isPowertypeOfDef])])])])])])).
cnf(c_0_28, plain, (~type_(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_29, negated_conjecture, (type_(esk2_1(esk24_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_23]), c_0_25])])).
cnf(c_0_30, plain, (individual(X1)|~iof(X1,X2)|~firstordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_31, plain, (type_(X1)|~firstordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, plain, (iof(X1,X3)|~specializes(X1,X2)|~isPowertypeOf(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, negated_conjecture, (isPowertypeOf(esk23_0,esk22_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_34, plain, (iof(esk9_1(X1),X1)|firstordertype(X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35, negated_conjecture, (~individual(esk2_1(esk24_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_36, plain, (individual(esk2_1(X1))|~firstordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_23]), c_0_31])).
cnf(c_0_37, plain, (specializes(X1,X3)|~iof(X1,X2)|~isPowertypeOf(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_38, negated_conjecture, (iof(X1,esk23_0)|~specializes(X1,esk22_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_39, negated_conjecture, (firstordertype(esk24_0)|iof(esk9_1(esk24_0),esk24_0)), inference(spm,[status(thm)],[c_0_34, c_0_25])).
cnf(c_0_40, negated_conjecture, (~firstordertype(esk24_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_41, plain, (X1=X2|properSpecializes(X1,X2)|~specializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_42, negated_conjecture, (specializes(X1,esk22_0)|~iof(X1,esk23_0)), inference(spm,[status(thm)],[c_0_37, c_0_33])).
cnf(c_0_43, negated_conjecture, (iof(X1,esk23_0)|~iof(X1,esk24_0)), inference(spm,[status(thm)],[c_0_38, c_0_20])).
cnf(c_0_44, negated_conjecture, (iof(esk9_1(esk24_0),esk24_0)), inference(sr,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_45, negated_conjecture, (X1=esk22_0|properSpecializes(X1,esk22_0)|~iof(X1,esk23_0)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_46, negated_conjecture, (iof(esk9_1(esk24_0),esk23_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_47, negated_conjecture, (esk9_1(esk24_0)=esk22_0|properSpecializes(esk9_1(esk24_0),esk22_0)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_48, plain, (iof(X1,X3)|~iof(X1,X2)|~specializes(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_49, negated_conjecture, (esk9_1(esk24_0)=esk22_0|specializes(esk9_1(esk24_0),esk22_0)), inference(spm,[status(thm)],[c_0_15, c_0_47])).
cnf(c_0_50, negated_conjecture, (esk9_1(esk24_0)=esk22_0|iof(X1,esk22_0)|~iof(X1,esk9_1(esk24_0))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_51, negated_conjecture, (type_(esk9_1(esk24_0))), inference(spm,[status(thm)],[c_0_24, c_0_44])).
cnf(c_0_52, plain, (X1!=X2|~properSpecializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_53, plain, (type_(X2)|~iof(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_54, plain, (iof(esk1_1(X1),X1)|individual(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_55, negated_conjecture, (esk9_1(esk24_0)=esk22_0|iof(esk2_1(esk9_1(esk24_0)),esk22_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_23]), c_0_51])])).
cnf(c_0_56, plain, (~properSpecializes(X1,X1)), inference(er,[status(thm)],[c_0_52])).
cnf(c_0_57, plain, (iof(esk11_2(X1,X2),X1)|specializes(X1,X2)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_58, plain, (type_(X1)|individual(X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_59, negated_conjecture, (esk9_1(esk24_0)=esk22_0|type_(esk22_0)), inference(spm,[status(thm)],[c_0_53, c_0_55])).
cnf(c_0_60, negated_conjecture, (~iof(esk22_0,esk24_0)), inference(spm,[status(thm)],[c_0_56, c_0_16])).
cnf(c_0_61, plain, (specializes(X1,X2)|iof(esk11_2(X1,X2),X1)|individual(X1)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_62, negated_conjecture, (~individual(esk24_0)), inference(spm,[status(thm)],[c_0_28, c_0_25])).
cnf(c_0_63, negated_conjecture, (type_(esk22_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_59]), c_0_60])).
cnf(c_0_64, plain, (specializes(X1,X2)|~iof(esk11_2(X1,X2),X2)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_65, negated_conjecture, (specializes(esk24_0,X1)|iof(esk11_2(esk24_0,X1),esk23_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_61]), c_0_62])).
cnf(c_0_66, negated_conjecture, (specializes(esk22_0,X1)|iof(esk11_2(esk22_0,X1),esk22_0)), inference(spm,[status(thm)],[c_0_57, c_0_63])).
cnf(c_0_67, negated_conjecture, (specializes(esk24_0,esk23_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_25])])).
cnf(c_0_68, negated_conjecture, (~properSpecializes(esk24_0,esk23_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_69, negated_conjecture, (specializes(esk22_0,esk22_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_66]), c_0_63])])).
cnf(c_0_70, negated_conjecture, (esk24_0=esk23_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_67]), c_0_68])).
cnf(c_0_71, negated_conjecture, (iof(esk22_0,esk23_0)), inference(spm,[status(thm)],[c_0_38, c_0_69])).
cnf(c_0_72, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_70]), c_0_71])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 73
# Proof object clause steps            : 56
# Proof object formula steps           : 17
# Proof object conjectures             : 36
# Proof object clause conjectures      : 33
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 32
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 77
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 72
# Processed clauses                    : 805
# ...of these trivial                  : 1
# ...subsumed                          : 251
# ...remaining for further processing  : 553
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 31
# Backward-rewritten                   : 198
# Generated clauses                    : 1607
# ...of the previous two non-trivial   : 1630
# Contextual simplify-reflections      : 40
# Paramodulations                      : 1604
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 249
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 225
# Current number of unprocessed clauses: 778
# ...number of literals in the above   : 2543
# Current number of archived formulas  : 0
# Current number of archived clauses   : 304
# Clause-clause subsumption calls (NU) : 21319
# Rec. Clause-clause subsumption calls : 17565
# Non-unit clause-clause subsumptions  : 172
# Unit Clause-clause subsumption calls : 697
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 22
# BW rewrite match successes           : 12
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 26574

# -------------------------------------------------
# User time                : 0.066 s
# System time              : 0.002 s
# Total time               : 0.068 s
# Maximum resident set size: 1604 pages
% START OF SYSTEM OUTPUT
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 15 seconds
% START OF SYSTEM OUTPUT
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_pdlX
% SZS output start Proof for SOT_pdlX
fof(f7063,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f111,f115,f119,f123,f139,f151,f155,f159,f167,f171,f176,f182,f236,f240,f244,f248,f252,f257,f261,f272,f289,f320,f397,f403,f413,f433,f513,f602,f611,f634,f645,f676,f958,f2691,f3707,f3725,f4811,f4826,f4831,f4847,f5004,f5006,f5125,f5827,f6100,f6109,f6627])).
fof(f6627,plain,(
  spl12_3 | ~spl12_30 | spl12_198 | ~spl12_199 | ~spl12_321),
  inference(avatar_contradiction_clause,[],[f6626])).
fof(f6626,plain,(
  $false | (spl12_3 | ~spl12_30 | spl12_198 | ~spl12_199 | ~spl12_321)),
  inference(subsumption_resolution,[],[f6269,f6255])).
fof(f6255,plain,(
  iof(sK8(sK1,sK1),sK1) | (spl12_3 | ~spl12_30 | ~spl12_199 | ~spl12_321)),
  inference(backward_demodulation,[],[f2690,f6123])).
fof(f6123,plain,(
  sK1 = sK2 | (spl12_3 | ~spl12_30 | ~spl12_321)),
  inference(subsumption_resolution,[],[f6120,f118])).
fof(f118,plain,(
  ~properSpecializes(sK2,sK1) | spl12_3),
  inference(avatar_component_clause,[],[f117])).
fof(f117,plain,(
  spl12_3 <=> properSpecializes(sK2,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_3])])).
fof(f6120,plain,(
  sK1 = sK2 | properSpecializes(sK2,sK1) | (~spl12_30 | ~spl12_321)),
  inference(resolution,[],[f6108,f239])).
fof(f239,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | X0 = X1 | properSpecializes(X0,X1)) ) | ~spl12_30),
  inference(avatar_component_clause,[],[f238])).
fof(f238,plain,(
  spl12_30 <=> ! [X1,X0] : (~specializes(X0,X1) | X0 = X1 | properSpecializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_30])])).
fof(f6108,plain,(
  specializes(sK2,sK1) | ~spl12_321),
  inference(avatar_component_clause,[],[f6107])).
fof(f6107,plain,(
  spl12_321 <=> specializes(sK2,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_321])])).
fof(f2690,plain,(
  iof(sK8(sK2,sK1),sK1) | ~spl12_199),
  inference(avatar_component_clause,[],[f2689])).
fof(f2689,plain,(
  spl12_199 <=> iof(sK8(sK2,sK1),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_199])])).
fof(f6269,plain,(
  ~iof(sK8(sK1,sK1),sK1) | (spl12_3 | ~spl12_30 | spl12_198 | ~spl12_321)),
  inference(backward_demodulation,[],[f3706,f6123])).
fof(f3706,plain,(
  ~iof(sK8(sK2,sK1),sK2) | spl12_198),
  inference(avatar_component_clause,[],[f2686])).
fof(f2686,plain,(
  spl12_198 <=> iof(sK8(sK2,sK1),sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_198])])).
fof(f6109,plain,(
  spl12_321 | ~spl12_18 | ~spl12_35 | ~spl12_320),
  inference(avatar_split_clause,[],[f6105,f6098,f259,f180,f6107])).
fof(f180,plain,(
  spl12_18 <=> type_(sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_18])])).
fof(f259,plain,(
  spl12_35 <=> ! [X1,X0] : (~type_(X0) | ~iof(sK11(X0,X1),X1) | specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_35])])).
fof(f6098,plain,(
  spl12_320 <=> ! [X0] : (specializes(sK2,X0) | iof(sK11(sK2,X0),sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_320])])).
fof(f6105,plain,(
  specializes(sK2,sK1) | (~spl12_18 | ~spl12_35 | ~spl12_320)),
  inference(subsumption_resolution,[],[f6104,f181])).
fof(f181,plain,(
  type_(sK2) | ~spl12_18),
  inference(avatar_component_clause,[],[f180])).
fof(f6104,plain,(
  specializes(sK2,sK1) | ~type_(sK2) | (~spl12_35 | ~spl12_320)),
  inference(duplicate_literal_removal,[],[f6101])).
fof(f6101,plain,(
  specializes(sK2,sK1) | ~type_(sK2) | specializes(sK2,sK1) | (~spl12_35 | ~spl12_320)),
  inference(resolution,[],[f6099,f260])).
fof(f260,plain,(
  ( ! [X0,X1] : (~iof(sK11(X0,X1),X1) | ~type_(X0) | specializes(X0,X1)) ) | ~spl12_35),
  inference(avatar_component_clause,[],[f259])).
fof(f6099,plain,(
  ( ! [X0] : (iof(sK11(sK2,X0),sK1) | specializes(sK2,X0)) ) | ~spl12_320),
  inference(avatar_component_clause,[],[f6098])).
fof(f6100,plain,(
  spl12_320 | ~spl12_1 | ~spl12_289),
  inference(avatar_split_clause,[],[f5828,f5825,f109,f6098])).
fof(f109,plain,(
  spl12_1 <=> isPowertypeOf(sK1,sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_1])])).
fof(f5825,plain,(
  spl12_289 <=> ! [X1,X0] : (specializes(sK2,X0) | ~isPowertypeOf(X1,sK0) | iof(sK11(sK2,X0),X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_289])])).
fof(f5828,plain,(
  ( ! [X0] : (specializes(sK2,X0) | iof(sK11(sK2,X0),sK1)) ) | (~spl12_1 | ~spl12_289)),
  inference(resolution,[],[f5826,f110])).
fof(f110,plain,(
  isPowertypeOf(sK1,sK0) | ~spl12_1),
  inference(avatar_component_clause,[],[f109])).
fof(f5826,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X1,sK0) | specializes(sK2,X0) | iof(sK11(sK2,X0),X1)) ) | ~spl12_289),
  inference(avatar_component_clause,[],[f5825])).
fof(f5827,plain,(
  spl12_289 | ~spl12_31 | ~spl12_99),
  inference(avatar_split_clause,[],[f646,f643,f242,f5825])).
fof(f242,plain,(
  spl12_31 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | ~specializes(X2,X1) | iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_31])])).
fof(f643,plain,(
  spl12_99 <=> ! [X0] : (specializes(sK2,X0) | specializes(sK11(sK2,X0),sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_99])])).
fof(f646,plain,(
  ( ! [X0,X1] : (specializes(sK2,X0) | ~isPowertypeOf(X1,sK0) | iof(sK11(sK2,X0),X1)) ) | (~spl12_31 | ~spl12_99)),
  inference(resolution,[],[f644,f243])).
fof(f243,plain,(
  ( ! [X2,X0,X1] : (~specializes(X2,X1) | ~isPowertypeOf(X0,X1) | iof(X2,X0)) ) | ~spl12_31),
  inference(avatar_component_clause,[],[f242])).
fof(f644,plain,(
  ( ! [X0] : (specializes(sK11(sK2,X0),sK0) | specializes(sK2,X0)) ) | ~spl12_99),
  inference(avatar_component_clause,[],[f643])).
fof(f5125,plain,(
  spl12_199 | ~spl12_1 | ~spl12_221),
  inference(avatar_split_clause,[],[f5124,f3723,f109,f2689])).
fof(f3723,plain,(
  spl12_221 <=> ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK8(sK2,sK1),X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_221])])).
fof(f5124,plain,(
  iof(sK8(sK2,sK1),sK1) | (~spl12_1 | ~spl12_221)),
  inference(resolution,[],[f3724,f110])).
fof(f3724,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK8(sK2,sK1),X0)) ) | ~spl12_221),
  inference(avatar_component_clause,[],[f3723])).
fof(f5006,plain,(
  spl12_200 | ~spl12_16 | ~spl12_229),
  inference(avatar_split_clause,[],[f5005,f5002,f169,f3388])).
fof(f3388,plain,(
  spl12_200 <=> specializes(sK8(sK2,sK1),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_200])])).
fof(f169,plain,(
  spl12_16 <=> ! [X1,X0] : (specializes(X0,X1) | ~properSpecializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_16])])).
fof(f5002,plain,(
  spl12_229 <=> properSpecializes(sK8(sK2,sK1),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_229])])).
fof(f5005,plain,(
  specializes(sK8(sK2,sK1),sK0) | (~spl12_16 | ~spl12_229)),
  inference(resolution,[],[f5003,f170])).
fof(f170,plain,(
  ( ! [X0,X1] : (~properSpecializes(X0,X1) | specializes(X0,X1)) ) | ~spl12_16),
  inference(avatar_component_clause,[],[f169])).
fof(f5003,plain,(
  properSpecializes(sK8(sK2,sK1),sK0) | ~spl12_229),
  inference(avatar_component_clause,[],[f5002])).
fof(f5004,plain,(
  spl12_229 | ~spl12_45 | ~spl12_198),
  inference(avatar_split_clause,[],[f4943,f2686,f318,f5002])).
fof(f318,plain,(
  spl12_45 <=> ! [X0] : (~iof(X0,sK2) | properSpecializes(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_45])])).
fof(f4943,plain,(
  properSpecializes(sK8(sK2,sK1),sK0) | (~spl12_45 | ~spl12_198)),
  inference(resolution,[],[f2687,f319])).
fof(f319,plain,(
  ( ! [X0] : (~iof(X0,sK2) | properSpecializes(X0,sK0)) ) | ~spl12_45),
  inference(avatar_component_clause,[],[f318])).
fof(f2687,plain,(
  iof(sK8(sK2,sK1),sK2) | ~spl12_198),
  inference(avatar_component_clause,[],[f2686])).
fof(f4847,plain,(
  spl12_88 | ~spl12_35 | ~spl12_63 | ~spl12_103),
  inference(avatar_split_clause,[],[f4827,f674,f411,f259,f557])).
fof(f557,plain,(
  spl12_88 <=> specializes(sK0,sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_88])])).
fof(f411,plain,(
  spl12_63 <=> type_(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_63])])).
fof(f674,plain,(
  spl12_103 <=> ! [X0] : (iof(sK11(sK0,X0),sK0) | specializes(sK0,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_103])])).
fof(f4827,plain,(
  specializes(sK0,sK0) | (~spl12_35 | ~spl12_63 | ~spl12_103)),
  inference(subsumption_resolution,[],[f679,f412])).
fof(f412,plain,(
  type_(sK0) | ~spl12_63),
  inference(avatar_component_clause,[],[f411])).
fof(f679,plain,(
  specializes(sK0,sK0) | ~type_(sK0) | (~spl12_35 | ~spl12_103)),
  inference(duplicate_literal_removal,[],[f677])).
fof(f677,plain,(
  specializes(sK0,sK0) | ~type_(sK0) | specializes(sK0,sK0) | (~spl12_35 | ~spl12_103)),
  inference(resolution,[],[f675,f260])).
fof(f675,plain,(
  ( ! [X0] : (iof(sK11(sK0,X0),sK0) | specializes(sK0,X0)) ) | ~spl12_103),
  inference(avatar_component_clause,[],[f674])).
fof(f4831,plain,(
  spl12_38 | ~spl12_15 | ~spl12_17 | ~spl12_37),
  inference(avatar_split_clause,[],[f3251,f270,f174,f165,f277])).
fof(f277,plain,(
  spl12_38 <=> specializes(sK7(sK1),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_38])])).
fof(f165,plain,(
  spl12_15 <=> ! [X0] : (iof(sK7(X0),X0) | ~type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_15])])).
fof(f174,plain,(
  spl12_17 <=> type_(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_17])])).
fof(f270,plain,(
  spl12_37 <=> ! [X0] : (specializes(X0,sK0) | ~iof(X0,sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_37])])).
fof(f3251,plain,(
  specializes(sK7(sK1),sK0) | (~spl12_15 | ~spl12_17 | ~spl12_37)),
  inference(subsumption_resolution,[],[f3250,f175])).
fof(f175,plain,(
  type_(sK1) | ~spl12_17),
  inference(avatar_component_clause,[],[f174])).
fof(f3250,plain,(
  specializes(sK7(sK1),sK0) | ~type_(sK1) | (~spl12_15 | ~spl12_37)),
  inference(resolution,[],[f271,f166])).
fof(f166,plain,(
  ( ! [X0] : (iof(sK7(X0),X0) | ~type_(X0)) ) | ~spl12_15),
  inference(avatar_component_clause,[],[f165])).
fof(f271,plain,(
  ( ! [X0] : (~iof(X0,sK1) | specializes(X0,sK0)) ) | ~spl12_37),
  inference(avatar_component_clause,[],[f270])).
fof(f4826,plain,(
  ~spl12_1 | ~spl12_4 | ~spl12_93 | ~spl12_227),
  inference(avatar_contradiction_clause,[],[f4825])).
fof(f4825,plain,(
  $false | (~spl12_1 | ~spl12_4 | ~spl12_93 | ~spl12_227)),
  inference(subsumption_resolution,[],[f612,f4815])).
fof(f4815,plain,(
  ~iof(sK0,sK1) | (~spl12_4 | ~spl12_227)),
  inference(resolution,[],[f4810,f122])).
fof(f122,plain,(
  ( ! [X1] : (~properSpecializes(X1,X1)) ) | ~spl12_4),
  inference(avatar_component_clause,[],[f121])).
fof(f121,plain,(
  spl12_4 <=> ! [X1] : ~properSpecializes(X1,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_4])])).
fof(f4810,plain,(
  ( ! [X0] : (properSpecializes(X0,sK0) | ~iof(X0,sK1)) ) | ~spl12_227),
  inference(avatar_component_clause,[],[f4809])).
fof(f4809,plain,(
  spl12_227 <=> ! [X0] : (~iof(X0,sK1) | properSpecializes(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_227])])).
fof(f612,plain,(
  iof(sK0,sK1) | (~spl12_1 | ~spl12_93)),
  inference(resolution,[],[f601,f110])).
fof(f601,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK0,X0)) ) | ~spl12_93),
  inference(avatar_component_clause,[],[f600])).
fof(f600,plain,(
  spl12_93 <=> ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK0,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_93])])).
fof(f4811,plain,(
  spl12_227 | ~spl12_45 | ~spl12_197),
  inference(avatar_split_clause,[],[f4411,f2683,f318,f4809])).
fof(f2683,plain,(
  spl12_197 <=> sK1 = sK2),
  introduced(avatar_definition,[new_symbols(naming,[spl12_197])])).
fof(f4411,plain,(
  ( ! [X0] : (~iof(X0,sK1) | properSpecializes(X0,sK0)) ) | (~spl12_45 | ~spl12_197)),
  inference(backward_demodulation,[],[f319,f2684])).
fof(f2684,plain,(
  sK1 = sK2 | ~spl12_197),
  inference(avatar_component_clause,[],[f2683])).
fof(f3725,plain,(
  spl12_221 | ~spl12_31 | ~spl12_200),
  inference(avatar_split_clause,[],[f3396,f3388,f242,f3723])).
fof(f3396,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK8(sK2,sK1),X0)) ) | (~spl12_31 | ~spl12_200)),
  inference(resolution,[],[f3389,f243])).
fof(f3389,plain,(
  specializes(sK8(sK2,sK1),sK0) | ~spl12_200),
  inference(avatar_component_clause,[],[f3388])).
fof(f3707,plain,(
  spl12_197 | ~spl12_198 | ~spl12_67 | ~spl12_199),
  inference(avatar_split_clause,[],[f2692,f2689,f431,f2686,f2683])).
fof(f431,plain,(
  spl12_67 <=> ! [X1,X0] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_67])])).
fof(f2692,plain,(
  ~iof(sK8(sK2,sK1),sK2) | sK1 = sK2 | (~spl12_67 | ~spl12_199)),
  inference(resolution,[],[f2690,f432])).
fof(f432,plain,(
  ( ! [X0,X1] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) ) | ~spl12_67),
  inference(avatar_component_clause,[],[f431])).
fof(f2691,plain,(
  spl12_197 | spl12_198 | spl12_199 | ~spl12_18 | ~spl12_138),
  inference(avatar_split_clause,[],[f966,f956,f180,f2689,f2686,f2683])).
fof(f956,plain,(
  spl12_138 <=> ! [X1] : (~type_(X1) | iof(sK8(X1,sK1),sK1) | iof(sK8(X1,sK1),X1) | sK1 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_138])])).
fof(f966,plain,(
  iof(sK8(sK2,sK1),sK1) | iof(sK8(sK2,sK1),sK2) | sK1 = sK2 | (~spl12_18 | ~spl12_138)),
  inference(resolution,[],[f957,f181])).
fof(f957,plain,(
  ( ! [X1] : (~type_(X1) | iof(sK8(X1,sK1),sK1) | iof(sK8(X1,sK1),X1) | sK1 = X1) ) | ~spl12_138),
  inference(avatar_component_clause,[],[f956])).
fof(f958,plain,(
  spl12_138 | ~spl12_17 | ~spl12_94),
  inference(avatar_split_clause,[],[f636,f609,f174,f956])).
fof(f609,plain,(
  spl12_94 <=> ! [X1,X0] : (~type_(X0) | ~type_(X1) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_94])])).
fof(f636,plain,(
  ( ! [X1] : (~type_(X1) | iof(sK8(X1,sK1),sK1) | iof(sK8(X1,sK1),X1) | sK1 = X1) ) | (~spl12_17 | ~spl12_94)),
  inference(resolution,[],[f610,f175])).
fof(f610,plain,(
  ( ! [X0,X1] : (~type_(X1) | ~type_(X0) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1) ) | ~spl12_94),
  inference(avatar_component_clause,[],[f609])).
fof(f676,plain,(
  spl12_103 | ~spl12_34 | ~spl12_63),
  inference(avatar_split_clause,[],[f414,f411,f255,f674])).
fof(f255,plain,(
  spl12_34 <=> ! [X1,X0] : (~type_(X0) | iof(sK11(X0,X1),X0) | specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_34])])).
fof(f414,plain,(
  ( ! [X0] : (iof(sK11(sK0,X0),sK0) | specializes(sK0,X0)) ) | (~spl12_34 | ~spl12_63)),
  inference(resolution,[],[f412,f256])).
fof(f256,plain,(
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_pdlX
  ( ! [X0,X1] : (~type_(X0) | iof(sK11(X0,X1),X0) | specializes(X0,X1)) ) | ~spl12_34),
  inference(avatar_component_clause,[],[f255])).
fof(f645,plain,(
  spl12_99 | ~spl12_16 | ~spl12_98),
  inference(avatar_split_clause,[],[f641,f632,f169,f643])).
fof(f632,plain,(
  spl12_98 <=> ! [X0] : (specializes(sK2,X0) | properSpecializes(sK11(sK2,X0),sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_98])])).
fof(f641,plain,(
  ( ! [X0] : (specializes(sK2,X0) | specializes(sK11(sK2,X0),sK0)) ) | (~spl12_16 | ~spl12_98)),
  inference(resolution,[],[f633,f170])).
fof(f633,plain,(
  ( ! [X0] : (properSpecializes(sK11(sK2,X0),sK0) | specializes(sK2,X0)) ) | ~spl12_98),
  inference(avatar_component_clause,[],[f632])).
fof(f634,plain,(
  spl12_98 | ~spl12_45 | ~spl12_82),
  inference(avatar_split_clause,[],[f515,f511,f318,f632])).
fof(f511,plain,(
  spl12_82 <=> ! [X1] : (iof(sK11(sK2,X1),sK2) | specializes(sK2,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_82])])).
fof(f515,plain,(
  ( ! [X0] : (specializes(sK2,X0) | properSpecializes(sK11(sK2,X0),sK0)) ) | (~spl12_45 | ~spl12_82)),
  inference(resolution,[],[f512,f319])).
fof(f512,plain,(
  ( ! [X1] : (iof(sK11(sK2,X1),sK2) | specializes(sK2,X1)) ) | ~spl12_82),
  inference(avatar_component_clause,[],[f511])).
fof(f611,plain,(
  spl12_94),
  inference(avatar_split_clause,[],[f76,f609])).
fof(f76,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~type_(X1) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ((X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))) | ~type_(X1) | ~type_(X0))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : ((X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))) | (~type_(X1) | ~type_(X0)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : ((type_(X1) & type_(X0)) => (X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7] : ((type_(X7) & type_(X6)) => (X6 = X7 <=> ! [X0] : (iof(X0,X6) <=> iof(X0,X7))))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a9_equality)).
fof(f602,plain,(
  spl12_93 | ~spl12_31 | ~spl12_88),
  inference(avatar_split_clause,[],[f560,f557,f242,f600])).
fof(f560,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK0,X0)) ) | (~spl12_31 | ~spl12_88)),
  inference(resolution,[],[f558,f243])).
fof(f558,plain,(
  specializes(sK0,sK0) | ~spl12_88),
  inference(avatar_component_clause,[],[f557])).
fof(f513,plain,(
  spl12_82 | ~spl12_18 | ~spl12_34),
  inference(avatar_split_clause,[],[f285,f255,f180,f511])).
fof(f285,plain,(
  ( ! [X1] : (iof(sK11(sK2,X1),sK2) | specializes(sK2,X1)) ) | (~spl12_18 | ~spl12_34)),
  inference(resolution,[],[f256,f181])).
fof(f433,plain,(
  spl12_67),
  inference(avatar_split_clause,[],[f104,f431])).
fof(f104,plain,(
  ( ! [X0,X1] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(subsumption_resolution,[],[f103,f70])).
fof(f70,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a3)).
fof(f103,plain,(
  ( ! [X0,X1] : (~type_(X1) | ~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(subsumption_resolution,[],[f77,f70])).
fof(f77,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~type_(X1) | ~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f48])).
fof(f413,plain,(
  spl12_63 | ~spl12_8 | ~spl12_61),
  inference(avatar_split_clause,[],[f409,f401,f137,f411])).
fof(f137,plain,(
  spl12_8 <=> ! [X1,X0] : (~iof(X1,X0) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_8])])).
fof(f401,plain,(
  spl12_61 <=> iof(sK7(sK7(sK1)),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_61])])).
fof(f409,plain,(
  type_(sK0) | (~spl12_8 | ~spl12_61)),
  inference(resolution,[],[f402,f138])).
fof(f138,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) ) | ~spl12_8),
  inference(avatar_component_clause,[],[f137])).
fof(f402,plain,(
  iof(sK7(sK7(sK1)),sK0) | ~spl12_61),
  inference(avatar_component_clause,[],[f401])).
fof(f403,plain,(
  spl12_61 | ~spl12_15 | ~spl12_39 | ~spl12_60),
  inference(avatar_split_clause,[],[f399,f395,f287,f165,f401])).
fof(f287,plain,(
  spl12_39 <=> type_(sK7(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_39])])).
fof(f395,plain,(
  spl12_60 <=> ! [X1] : (iof(X1,sK0) | ~iof(X1,sK7(sK1)))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_60])])).
fof(f399,plain,(
  iof(sK7(sK7(sK1)),sK0) | (~spl12_15 | ~spl12_39 | ~spl12_60)),
  inference(subsumption_resolution,[],[f398,f288])).
fof(f288,plain,(
  type_(sK7(sK1)) | ~spl12_39),
  inference(avatar_component_clause,[],[f287])).
fof(f398,plain,(
  iof(sK7(sK7(sK1)),sK0) | ~type_(sK7(sK1)) | (~spl12_15 | ~spl12_60)),
  inference(resolution,[],[f396,f166])).
fof(f396,plain,(
  ( ! [X1] : (~iof(X1,sK7(sK1)) | iof(X1,sK0)) ) | ~spl12_60),
  inference(avatar_component_clause,[],[f395])).
fof(f397,plain,(
  spl12_60 | ~spl12_29 | ~spl12_38),
  inference(avatar_split_clause,[],[f282,f277,f234,f395])).
fof(f234,plain,(
  spl12_29 <=> ! [X1,X0,X2] : (~iof(X2,X0) | iof(X2,X1) | ~specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_29])])).
fof(f282,plain,(
  ( ! [X1] : (iof(X1,sK0) | ~iof(X1,sK7(sK1))) ) | (~spl12_29 | ~spl12_38)),
  inference(resolution,[],[f278,f235])).
fof(f235,plain,(
  ( ! [X2,X0,X1] : (~specializes(X0,X1) | iof(X2,X1) | ~iof(X2,X0)) ) | ~spl12_29),
  inference(avatar_component_clause,[],[f234])).
fof(f278,plain,(
  specializes(sK7(sK1),sK0) | ~spl12_38),
  inference(avatar_component_clause,[],[f277])).
fof(f320,plain,(
  spl12_45 | ~spl12_2 | ~spl12_33),
  inference(avatar_split_clause,[],[f273,f250,f113,f318])).
fof(f113,plain,(
  spl12_2 <=> categorizes(sK2,sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_2])])).
fof(f250,plain,(
  spl12_33 <=> ! [X1,X0,X2] : (~categorizes(X0,X1) | ~iof(X2,X0) | properSpecializes(X2,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_33])])).
fof(f273,plain,(
  ( ! [X0] : (~iof(X0,sK2) | properSpecializes(X0,sK0)) ) | (~spl12_2 | ~spl12_33)),
  inference(resolution,[],[f251,f114])).
fof(f114,plain,(
  categorizes(sK2,sK0) | ~spl12_2),
  inference(avatar_component_clause,[],[f113])).
fof(f251,plain,(
  ( ! [X2,X0,X1] : (~categorizes(X0,X1) | ~iof(X2,X0) | properSpecializes(X2,X1)) ) | ~spl12_33),
  inference(avatar_component_clause,[],[f250])).
fof(f289,plain,(
  spl12_39 | ~spl12_11 | ~spl12_38),
  inference(avatar_split_clause,[],[f283,f277,f149,f287])).
fof(f149,plain,(
  spl12_11 <=> ! [X1,X0] : (type_(X0) | ~specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_11])])).
fof(f283,plain,(
  type_(sK7(sK1)) | (~spl12_11 | ~spl12_38)),
  inference(resolution,[],[f278,f150])).
fof(f150,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) ) | ~spl12_11),
  inference(avatar_component_clause,[],[f149])).
fof(f272,plain,(
  spl12_37 | ~spl12_1 | ~spl12_32),
  inference(avatar_split_clause,[],[f268,f246,f109,f270])).
fof(f246,plain,(
  spl12_32 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_32])])).
fof(f268,plain,(
  ( ! [X0] : (specializes(X0,sK0) | ~iof(X0,sK1)) ) | (~spl12_1 | ~spl12_32)),
  inference(resolution,[],[f247,f110])).
fof(f247,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) ) | ~spl12_32),
  inference(avatar_component_clause,[],[f246])).
fof(f261,plain,(
  spl12_35),
  inference(avatar_split_clause,[],[f87,f259])).
fof(f87,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~iof(sK11(X0,X1),X1) | specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X1) | ~iof(X2,X0)) & type_(X0)))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X0) => iof(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7] : (specializes(X6,X7) <=> (! [X3] : (iof(X3,X6) => iof(X3,X7)) & type_(X6)))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a7_specDef)).
fof(f257,plain,(
  spl12_34),
  inference(avatar_split_clause,[],[f86,f255])).
fof(f86,plain,(
  ( ! [X0,X1] : (~type_(X0) | iof(sK11(X0,X1),X0) | specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f53])).
fof(f252,plain,(
  spl12_33),
  inference(avatar_split_clause,[],[f96,f250])).
fof(f96,plain,(
  ( ! [X2,X0,X1] : (~categorizes(X0,X1) | ~iof(X2,X0) | properSpecializes(X2,X1)) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : ((! [X2] : (properSpecializes(X2,X1) | ~iof(X2,X0)) & type_(X0)) | ~categorizes(X0,X1))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : (categorizes(X0,X1) => (! [X2] : (iof(X2,X0) => properSpecializes(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (categorizes(X0,X1) <=> (! [X2] : (iof(X2,X0) => properSpecializes(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6,X7] : (categorizes(X6,X7) <=> (! [X8] : (iof(X8,X6) => properSpecializes(X8,X7)) & type_(X6)))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a11_categorizationDef)).
fof(f248,plain,(
  spl12_32),
  inference(avatar_split_clause,[],[f94,f246])).
fof(f94,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : ((! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)) | ~isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) => (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) <=> (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : (isPowertypeOf(X6,X7) <=> (! [X8] : (iof(X8,X6) <=> specializes(X8,X7)) & type_(X6)))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a10_isPowertypeOfDef)).
fof(f244,plain,(
  spl12_31),
  inference(avatar_split_clause,[],[f93,f242])).
fof(f93,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | ~specializes(X2,X1) | iof(X2,X0)) )),
  inference(cnf_transformation,[],[f54])).
fof(f240,plain,(
  spl12_30),
  inference(avatar_split_clause,[],[f92,f238])).
fof(f92,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | X0 = X1 | properSpecializes(X0,X1)) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1] : (properSpecializes(X0,X1) <=> (X0 != X1 & specializes(X0,X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X7] : (properSpecializes(X6,X7) <=> (X6 != X7 & specializes(X6,X7)))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a8_properSpecDef)).
fof(f236,plain,(
  spl12_29),
  inference(avatar_split_clause,[],[f88,f234])).
fof(f88,plain,(
  ( ! [X2,X0,X1] : (~iof(X2,X0) | iof(X2,X1) | ~specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f53])).
fof(f182,plain,(
  spl12_18 | ~spl12_2 | ~spl12_13),
  inference(avatar_split_clause,[],[f177,f157,f113,f180])).
fof(f157,plain,(
  spl12_13 <=> ! [X1,X0] : (~categorizes(X0,X1) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_13])])).
fof(f177,plain,(
  type_(sK2) | (~spl12_2 | ~spl12_13)),
  inference(resolution,[],[f158,f114])).
fof(f158,plain,(
  ( ! [X0,X1] : (~categorizes(X0,X1) | type_(X0)) ) | ~spl12_13),
  inference(avatar_component_clause,[],[f157])).
fof(f176,plain,(
  spl12_17 | ~spl12_1 | ~spl12_12),
  inference(avatar_split_clause,[],[f172,f153,f109,f174])).
fof(f153,plain,(
  spl12_12 <=> ! [X1,X0] : (~isPowertypeOf(X0,X1) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_12])])).
fof(f172,plain,(
  type_(sK1) | (~spl12_1 | ~spl12_12)),
  inference(resolution,[],[f154,f110])).
fof(f154,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X0,X1) | type_(X0)) ) | ~spl12_12),
  inference(avatar_component_clause,[],[f153])).
fof(f171,plain,(
  spl12_16),
  inference(avatar_split_clause,[],[f90,f169])).
fof(f90,plain,(
  ( ! [X0,X1] : (specializes(X0,X1) | ~properSpecializes(X0,X1)) )),
  inference(cnf_transformation,[],[f32])).
fof(f167,plain,(
  spl12_15),
  inference(avatar_split_clause,[],[f71,f165])).
fof(f71,plain,(
  ( ! [X0] : (iof(sK7(X0),X0) | ~type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f159,plain,(
  spl12_13),
  inference(avatar_split_clause,[],[f97,f157])).
fof(f97,plain,(
  ( ! [X0,X1] : (~categorizes(X0,X1) | type_(X0)) )),
  inference(cnf_transformation,[],[f55])).
fof(f155,plain,(
  spl12_12),
  inference(avatar_split_clause,[],[f95,f153])).
fof(f95,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X0,X1) | type_(X0)) )),
  inference(cnf_transformation,[],[f54])).
fof(f151,plain,(
  spl12_11),
  inference(avatar_split_clause,[],[f89,f149])).
fof(f89,plain,(
  ( ! [X0,X1] : (type_(X0) | ~specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f53])).
fof(f139,plain,(
  spl12_8),
  inference(avatar_split_clause,[],[f70,f137])).
fof(f123,plain,(
  spl12_4),
  inference(avatar_split_clause,[],[f102,f121])).
fof(f102,plain,(
  ( ! [X1] : (~properSpecializes(X1,X1)) )),
  inference(equality_resolution,[],[f91])).
fof(f91,plain,(
  ( ! [X0,X1] : (X0 != X1 | ~properSpecializes(X0,X1)) )),
  inference(cnf_transformation,[],[f32])).
fof(f119,plain,(
  ~spl12_3),
  inference(avatar_split_clause,[],[f58,f117])).
fof(f58,plain,(
  ~properSpecializes(sK2,sK1)),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ? [X0,X1,X2] : (~properSpecializes(X2,X1) & categorizes(X2,X0) & isPowertypeOf(X1,X0))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ? [X0,X1,X2] : (~properSpecializes(X2,X1) & (categorizes(X2,X0) & isPowertypeOf(X1,X0)))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1,X2] : ((categorizes(X2,X0) & isPowertypeOf(X1,X0)) => properSpecializes(X2,X1))),
  inference(rectify,[],[f22])).
fof(f22,negated_conjecture,(
  ~! [X6,X7,X8] : ((categorizes(X8,X6) & isPowertypeOf(X7,X6)) => properSpecializes(X8,X7))),
  inference(negated_conjecture,[],[f21])).
fof(f21,conjecture,(
  ! [X6,X7,X8] : ((categorizes(X8,X6) & isPowertypeOf(X7,X6)) => properSpecializes(X8,X7))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',t4)).
fof(f115,plain,(
  spl12_2),
  inference(avatar_split_clause,[],[f57,f113])).
fof(f57,plain,(
  categorizes(sK2,sK0)),
  inference(cnf_transformation,[],[f40])).
fof(f111,plain,(
  spl12_1),
  inference(avatar_split_clause,[],[f56,f109])).
fof(f56,plain,(
  isPowertypeOf(sK1,sK0)),
  inference(cnf_transformation,[],[f40])).
% SZS output end Proof for SOT_pdlX
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 9850
% Time elapsed: 0.302 s
% ------------------------------
% ------------------------------
% Success in time 0.355 s
# AutoSched0-Mode selected heuristic G_E___208_C02CMA_F1_SE_CS_SP_PS_S5PRR_RG_S04AN
# and selection function SelectComplexExceptUniqMaxHorn.
#
# Preprocessing time       : 0.028 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4, conjecture, ![X7, X8, X9]:((isPowertypeOf(X8,X7)&categorizes(X9,X7))=>properSpecializes(X9,X8)), file('/tmp/wqEVxBSTg1/SOT_pdlXwd', t4)).
fof(a11_categorizationDef, axiom, ![X7, X8]:(categorizes(X7,X8)<=>(type_(X7)&![X9]:(iof(X9,X7)=>properSpecializes(X9,X8)))), file('/tmp/wqEVxBSTg1/SOT_pdlXwd', a11_categorizationDef)).
fof(a8_properSpecDef, axiom, ![X7, X8]:(properSpecializes(X7,X8)<=>(specializes(X7,X8)&~(X7=X8))), file('/tmp/wqEVxBSTg1/SOT_pdlXwd', a8_properSpecDef)).
fof(a7_specDef, axiom, ![X7, X8]:(specializes(X7,X8)<=>(type_(X7)&![X4]:(iof(X4,X7)=>iof(X4,X8)))), file('/tmp/wqEVxBSTg1/SOT_pdlXwd', a7_specDef)).
fof(a2, axiom, ![X1]:(individual(X1)<=>~(?[X2]:iof(X2,X1))), file('/tmp/wqEVxBSTg1/SOT_pdlXwd', a2)).
fof(a3, axiom, ![X1]:(type_(X1)<=>?[X2]:iof(X2,X1)), file('/tmp/wqEVxBSTg1/SOT_pdlXwd', a3)).
fof(a5_firstOrderTypeDef, axiom, ![X1]:(firstordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>individual(X2)))), file('/tmp/wqEVxBSTg1/SOT_pdlXwd', a5_firstOrderTypeDef)).
fof(a10_isPowertypeOfDef, axiom, ![X7, X8]:(isPowertypeOf(X7,X8)<=>(type_(X7)&![X9]:(iof(X9,X7)<=>specializes(X9,X8)))), file('/tmp/wqEVxBSTg1/SOT_pdlXwd', a10_isPowertypeOfDef)).
fof(c_0_8, negated_conjecture, ~(![X7, X8, X9]:((isPowertypeOf(X8,X7)&categorizes(X9,X7))=>properSpecializes(X9,X8))), inference(assume_negation,[status(cth)],[t4])).
fof(c_0_9, plain, ![X69, X70, X71, X72, X73]:(((type_(X69)|~categorizes(X69,X70))&(~iof(X71,X69)|properSpecializes(X71,X70)|~categorizes(X69,X70)))&((iof(esk14_2(X72,X73),X72)|~type_(X72)|categorizes(X72,X73))&(~properSpecializes(esk14_2(X72,X73),X73)|~type_(X72)|categorizes(X72,X73)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a11_categorizationDef])])])])])])).
fof(c_0_10, negated_conjecture, ((isPowertypeOf(esk23_0,esk22_0)&categorizes(esk24_0,esk22_0))&~properSpecializes(esk24_0,esk23_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
fof(c_0_11, plain, ![X55, X56]:(((specializes(X55,X56)|~properSpecializes(X55,X56))&(X55!=X56|~properSpecializes(X55,X56)))&(~specializes(X55,X56)|X55=X56|properSpecializes(X55,X56))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[a8_properSpecDef])])])])).
cnf(c_0_12, plain, (properSpecializes(X1,X3)|~iof(X1,X2)|~categorizes(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, negated_conjecture, (categorizes(esk24_0,esk22_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, plain, ![X49, X50, X51, X52, X53]:(((type_(X49)|~specializes(X49,X50))&(~iof(X51,X49)|iof(X51,X50)|~specializes(X49,X50)))&((iof(esk11_2(X52,X53),X52)|~type_(X52)|specializes(X52,X53))&(~iof(esk11_2(X52,X53),X53)|~type_(X52)|specializes(X52,X53)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a7_specDef])])])])])])).
cnf(c_0_15, plain, (specializes(X1,X2)|~properSpecializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, negated_conjecture, (properSpecializes(X1,esk22_0)|~iof(X1,esk24_0)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
fof(c_0_17, plain, ![X12, X13, X14]:((~individual(X12)|~iof(X13,X12))&(iof(esk1_1(X14),X14)|individual(X14))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a2])])])])])).
fof(c_0_18, plain, ![X16, X18, X19]:((~type_(X16)|iof(esk2_1(X16),X16))&(~iof(X19,X18)|type_(X18))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a3])])])])])).
cnf(c_0_19, plain, (type_(X1)|~specializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, negated_conjecture, (specializes(X1,esk22_0)|~iof(X1,esk24_0)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_21, plain, (type_(X1)|~categorizes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_22, plain, (~individual(X1)|~iof(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (iof(esk2_1(X1),X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, negated_conjecture, (type_(X1)|~iof(X1,esk24_0)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_25, negated_conjecture, (type_(esk24_0)), inference(spm,[status(thm)],[c_0_21, c_0_13])).
fof(c_0_26, plain, ![X41, X42, X43]:(((type_(X41)|~firstordertype(X41))&(~iof(X42,X41)|individual(X42)|~firstordertype(X41)))&((iof(esk9_1(X43),X43)|~type_(X43)|firstordertype(X43))&(~individual(esk9_1(X43))|~type_(X43)|firstordertype(X43)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a5_firstOrderTypeDef])])])])])])).
fof(c_0_27, plain, ![X62, X63, X64, X65, X66, X67]:(((type_(X62)|~isPowertypeOf(X62,X63))&((~iof(X64,X62)|specializes(X64,X63)|~isPowertypeOf(X62,X63))&(~specializes(X65,X63)|iof(X65,X62)|~isPowertypeOf(X62,X63))))&((~iof(esk13_2(X66,X67),X66)|~specializes(esk13_2(X66,X67),X67)|~type_(X66)|isPowertypeOf(X66,X67))&(iof(esk13_2(X66,X67),X66)|specializes(esk13_2(X66,X67),X67)|~type_(X66)|isPowertypeOf(X66,X67)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a10_isPowertypeOfDef])])])])])])).
cnf(c_0_28, plain, (~type_(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_29, negated_conjecture, (type_(esk2_1(esk24_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_23]), c_0_25])])).
cnf(c_0_30, plain, (individual(X1)|~iof(X1,X2)|~firstordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_31, plain, (type_(X1)|~firstordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, plain, (iof(X1,X3)|~specializes(X1,X2)|~isPowertypeOf(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, negated_conjecture, (isPowertypeOf(esk23_0,esk22_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_34, plain, (iof(esk9_1(X1),X1)|firstordertype(X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35, negated_conjecture, (~individual(esk2_1(esk24_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_36, plain, (individual(esk2_1(X1))|~firstordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_23]), c_0_31])).
cnf(c_0_37, plain, (specializes(X1,X3)|~iof(X1,X2)|~isPowertypeOf(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_38, negated_conjecture, (iof(X1,esk23_0)|~specializes(X1,esk22_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_39, negated_conjecture, (firstordertype(esk24_0)|iof(esk9_1(esk24_0),esk24_0)), inference(spm,[status(thm)],[c_0_34, c_0_25])).
cnf(c_0_40, negated_conjecture, (~firstordertype(esk24_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_41, plain, (X1=X2|properSpecializes(X1,X2)|~specializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_42, negated_conjecture, (specializes(X1,esk22_0)|~iof(X1,esk23_0)), inference(spm,[status(thm)],[c_0_37, c_0_33])).
cnf(c_0_43, negated_conjecture, (iof(X1,esk23_0)|~iof(X1,esk24_0)), inference(spm,[status(thm)],[c_0_38, c_0_20])).
cnf(c_0_44, negated_conjecture, (iof(esk9_1(esk24_0),esk24_0)), inference(sr,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_45, negated_conjecture, (X1=esk22_0|properSpecializes(X1,esk22_0)|~iof(X1,esk23_0)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_46, negated_conjecture, (iof(esk9_1(esk24_0),esk23_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_47, negated_conjecture, (esk9_1(esk24_0)=esk22_0|properSpecializes(esk9_1(esk24_0),esk22_0)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_48, plain, (iof(X1,X3)|~iof(X1,X2)|~specializes(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_49, negated_conjecture, (esk9_1(esk24_0)=esk22_0|specializes(esk9_1(esk24_0),esk22_0)), inference(spm,[status(thm)],[c_0_15, c_0_47])).
cnf(c_0_50, negated_conjecture, (esk9_1(esk24_0)=esk22_0|iof(X1,esk22_0)|~iof(X1,esk9_1(esk24_0))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_51, negated_conjecture, (type_(esk9_1(esk24_0))), inference(spm,[status(thm)],[c_0_24, c_0_44])).
cnf(c_0_52, plain, (X1!=X2|~properSpecializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_53, plain, (type_(X2)|~iof(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_54, plain, (iof(esk1_1(X1),X1)|individual(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_55, negated_conjecture, (esk9_1(esk24_0)=esk22_0|iof(esk2_1(esk9_1(esk24_0)),esk22_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_23]), c_0_51])])).
cnf(c_0_56, plain, (~properSpecializes(X1,X1)), inference(er,[status(thm)],[c_0_52])).
cnf(c_0_57, plain, (iof(esk11_2(X1,X2),X1)|specializes(X1,X2)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_58, plain, (type_(X1)|individual(X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_59, negated_conjecture, (esk9_1(esk24_0)=esk22_0|type_(esk22_0)), inference(spm,[status(thm)],[c_0_53, c_0_55])).
cnf(c_0_60, negated_conjecture, (~iof(esk22_0,esk24_0)), inference(spm,[status(thm)],[c_0_56, c_0_16])).
cnf(c_0_61, plain, (specializes(X1,X2)|iof(esk11_2(X1,X2),X1)|individual(X1)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_62, negated_conjecture, (~individual(esk24_0)), inference(spm,[status(thm)],[c_0_28, c_0_25])).
cnf(c_0_63, negated_conjecture, (type_(esk22_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_59]), c_0_60])).
cnf(c_0_64, plain, (specializes(X1,X2)|~iof(esk11_2(X1,X2),X2)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_65, negated_conjecture, (specializes(esk24_0,X1)|iof(esk11_2(esk24_0,X1),esk23_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_61]), c_0_62])).
cnf(c_0_66, negated_conjecture, (specializes(esk22_0,X1)|iof(esk11_2(esk22_0,X1),esk22_0)), inference(spm,[status(thm)],[c_0_57, c_0_63])).
cnf(c_0_67, negated_conjecture, (specializes(esk24_0,esk23_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_25])])).
cnf(c_0_68, negated_conjecture, (~properSpecializes(esk24_0,esk23_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_69, negated_conjecture, (specializes(esk22_0,esk22_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_66]), c_0_63])])).
cnf(c_0_70, negated_conjecture, (esk24_0=esk23_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_67]), c_0_68])).
cnf(c_0_71, negated_conjecture, (iof(esk22_0,esk23_0)), inference(spm,[status(thm)],[c_0_38, c_0_69])).
cnf(c_0_72, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_70]), c_0_71])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 73
# Proof object clause steps            : 56
# Proof object formula steps           : 17
# Proof object conjectures             : 36
# Proof object clause conjectures      : 33
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 32
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 77
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 72
# Processed clauses                    : 805
# ...of these trivial                  : 1
# ...subsumed                          : 251
# ...remaining for further processing  : 553
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 31
# Backward-rewritten                   : 198
# Generated clauses                    : 1607
# ...of the previous two non-trivial   : 1630
# Contextual simplify-reflections      : 40
# Paramodulations                      : 1604
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 249
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 225
# Current number of unprocessed clauses: 778
# ...number of literals in the above   : 2543
# Current number of archived formulas  : 0
# Current number of archived clauses   : 304
# Clause-clause subsumption calls (NU) : 21319
# Rec. Clause-clause subsumption calls : 17565
# Non-unit clause-clause subsumptions  : 172
# Unit Clause-clause subsumption calls : 697
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 22
# BW rewrite match successes           : 12
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 26574

# -------------------------------------------------
# User time                : 0.067 s
# System time              : 0.009 s
# Total time               : 0.076 s
# Maximum resident set size: 1608 pages
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_pdlX
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_pdlX
% SZS output start Proof for SOT_pdlX
fof(f7063,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f111,f115,f119,f123,f139,f151,f155,f159,f167,f171,f176,f182,f236,f240,f244,f248,f252,f257,f261,f272,f289,f320,f397,f403,f413,f433,f513,f602,f611,f634,f645,f676,f958,f2691,f3707,f3725,f4811,f4826,f4831,f4847,f5004,f5006,f5125,f5827,f6100,f6109,f6627])).
fof(f6627,plain,(
  spl12_3 | ~spl12_30 | spl12_198 | ~spl12_199 | ~spl12_321),
  inference(avatar_contradiction_clause,[],[f6626])).
fof(f6626,plain,(
  $false | (spl12_3 | ~spl12_30 | spl12_198 | ~spl12_199 | ~spl12_321)),
  inference(subsumption_resolution,[],[f6269,f6255])).
fof(f6255,plain,(
  iof(sK8(sK1,sK1),sK1) | (spl12_3 | ~spl12_30 | ~spl12_199 | ~spl12_321)),
  inference(backward_demodulation,[],[f2690,f6123])).
fof(f6123,plain,(
  sK1 = sK2 | (spl12_3 | ~spl12_30 | ~spl12_321)),
  inference(subsumption_resolution,[],[f6120,f118])).
fof(f118,plain,(
  ~properSpecializes(sK2,sK1) | spl12_3),
  inference(avatar_component_clause,[],[f117])).
fof(f117,plain,(
  spl12_3 <=> properSpecializes(sK2,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_3])])).
fof(f6120,plain,(
  sK1 = sK2 | properSpecializes(sK2,sK1) | (~spl12_30 | ~spl12_321)),
  inference(resolution,[],[f6108,f239])).
fof(f239,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | X0 = X1 | properSpecializes(X0,X1)) ) | ~spl12_30),
  inference(avatar_component_clause,[],[f238])).
fof(f238,plain,(
  spl12_30 <=> ! [X1,X0] : (~specializes(X0,X1) | X0 = X1 | properSpecializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_30])])).
fof(f6108,plain,(
  specializes(sK2,sK1) | ~spl12_321),
  inference(avatar_component_clause,[],[f6107])).
fof(f6107,plain,(
  spl12_321 <=> specializes(sK2,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_321])])).
fof(f2690,plain,(
  iof(sK8(sK2,sK1),sK1) | ~spl12_199),
  inference(avatar_component_clause,[],[f2689])).
fof(f2689,plain,(
  spl12_199 <=> iof(sK8(sK2,sK1),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_199])])).
fof(f6269,plain,(
  ~iof(sK8(sK1,sK1),sK1) | (spl12_3 | ~spl12_30 | spl12_198 | ~spl12_321)),
  inference(backward_demodulation,[],[f3706,f6123])).
fof(f3706,plain,(
  ~iof(sK8(sK2,sK1),sK2) | spl12_198),
  inference(avatar_component_clause,[],[f2686])).
fof(f2686,plain,(
  spl12_198 <=> iof(sK8(sK2,sK1),sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_198])])).
fof(f6109,plain,(
  spl12_321 | ~spl12_18 | ~spl12_35 | ~spl12_320),
  inference(avatar_split_clause,[],[f6105,f6098,f259,f180,f6107])).
fof(f180,plain,(
  spl12_18 <=> type_(sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_18])])).
fof(f259,plain,(
  spl12_35 <=> ! [X1,X0] : (~type_(X0) | ~iof(sK11(X0,X1),X1) | specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_35])])).
fof(f6098,plain,(
  spl12_320 <=> ! [X0] : (specializes(sK2,X0) | iof(sK11(sK2,X0),sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_320])])).
fof(f6105,plain,(
  specializes(sK2,sK1) | (~spl12_18 | ~spl12_35 | ~spl12_320)),
  inference(subsumption_resolution,[],[f6104,f181])).
fof(f181,plain,(
  type_(sK2) | ~spl12_18),
  inference(avatar_component_clause,[],[f180])).
fof(f6104,plain,(
  specializes(sK2,sK1) | ~type_(sK2) | (~spl12_35 | ~spl12_320)),
  inference(duplicate_literal_removal,[],[f6101])).
fof(f6101,plain,(
  specializes(sK2,sK1) | ~type_(sK2) | specializes(sK2,sK1) | (~spl12_35 | ~spl12_320)),
  inference(resolution,[],[f6099,f260])).
fof(f260,plain,(
  ( ! [X0,X1] : (~iof(sK11(X0,X1),X1) | ~type_(X0) | specializes(X0,X1)) ) | ~spl12_35),
  inference(avatar_component_clause,[],[f259])).
fof(f6099,plain,(
  ( ! [X0] : (iof(sK11(sK2,X0),sK1) | specializes(sK2,X0)) ) | ~spl12_320),
  inference(avatar_component_clause,[],[f6098])).
fof(f6100,plain,(
  spl12_320 | ~spl12_1 | ~spl12_289),
  inference(avatar_split_clause,[],[f5828,f5825,f109,f6098])).
fof(f109,plain,(
  spl12_1 <=> isPowertypeOf(sK1,sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_1])])).
fof(f5825,plain,(
  spl12_289 <=> ! [X1,X0] : (specializes(sK2,X0) | ~isPowertypeOf(X1,sK0) | iof(sK11(sK2,X0),X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_289])])).
fof(f5828,plain,(
  ( ! [X0] : (specializes(sK2,X0) | iof(sK11(sK2,X0),sK1)) ) | (~spl12_1 | ~spl12_289)),
  inference(resolution,[],[f5826,f110])).
fof(f110,plain,(
  isPowertypeOf(sK1,sK0) | ~spl12_1),
  inference(avatar_component_clause,[],[f109])).
fof(f5826,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X1,sK0) | specializes(sK2,X0) | iof(sK11(sK2,X0),X1)) ) | ~spl12_289),
  inference(avatar_component_clause,[],[f5825])).
fof(f5827,plain,(
  spl12_289 | ~spl12_31 | ~spl12_99),
  inference(avatar_split_clause,[],[f646,f643,f242,f5825])).
fof(f242,plain,(
  spl12_31 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | ~specializes(X2,X1) | iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_31])])).
fof(f643,plain,(
  spl12_99 <=> ! [X0] : (specializes(sK2,X0) | specializes(sK11(sK2,X0),sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_99])])).
fof(f646,plain,(
  ( ! [X0,X1] : (specializes(sK2,X0) | ~isPowertypeOf(X1,sK0) | iof(sK11(sK2,X0),X1)) ) | (~spl12_31 | ~spl12_99)),
  inference(resolution,[],[f644,f243])).
fof(f243,plain,(
  ( ! [X2,X0,X1] : (~specializes(X2,X1) | ~isPowertypeOf(X0,X1) | iof(X2,X0)) ) | ~spl12_31),
  inference(avatar_component_clause,[],[f242])).
fof(f644,plain,(
  ( ! [X0] : (specializes(sK11(sK2,X0),sK0) | specializes(sK2,X0)) ) | ~spl12_99),
  inference(avatar_component_clause,[],[f643])).
fof(f5125,plain,(
  spl12_199 | ~spl12_1 | ~spl12_221),
  inference(avatar_split_clause,[],[f5124,f3723,f109,f2689])).
fof(f3723,plain,(
  spl12_221 <=> ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK8(sK2,sK1),X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_221])])).
fof(f5124,plain,(
  iof(sK8(sK2,sK1),sK1) | (~spl12_1 | ~spl12_221)),
  inference(resolution,[],[f3724,f110])).
fof(f3724,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK8(sK2,sK1),X0)) ) | ~spl12_221),
  inference(avatar_component_clause,[],[f3723])).
fof(f5006,plain,(
  spl12_200 | ~spl12_16 | ~spl12_229),
  inference(avatar_split_clause,[],[f5005,f5002,f169,f3388])).
fof(f3388,plain,(
  spl12_200 <=> specializes(sK8(sK2,sK1),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_200])])).
fof(f169,plain,(
  spl12_16 <=> ! [X1,X0] : (specializes(X0,X1) | ~properSpecializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_16])])).
fof(f5002,plain,(
  spl12_229 <=> properSpecializes(sK8(sK2,sK1),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_229])])).
fof(f5005,plain,(
  specializes(sK8(sK2,sK1),sK0) | (~spl12_16 | ~spl12_229)),
  inference(resolution,[],[f5003,f170])).
fof(f170,plain,(
  ( ! [X0,X1] : (~properSpecializes(X0,X1) | specializes(X0,X1)) ) | ~spl12_16),
  inference(avatar_component_clause,[],[f169])).
fof(f5003,plain,(
  properSpecializes(sK8(sK2,sK1),sK0) | ~spl12_229),
  inference(avatar_component_clause,[],[f5002])).
fof(f5004,plain,(
  spl12_229 | ~spl12_45 | ~spl12_198),
  inference(avatar_split_clause,[],[f4943,f2686,f318,f5002])).
fof(f318,plain,(
  spl12_45 <=> ! [X0] : (~iof(X0,sK2) | properSpecializes(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_45])])).
fof(f4943,plain,(
  properSpecializes(sK8(sK2,sK1),sK0) | (~spl12_45 | ~spl12_198)),
  inference(resolution,[],[f2687,f319])).
fof(f319,plain,(
  ( ! [X0] : (~iof(X0,sK2) | properSpecializes(X0,sK0)) ) | ~spl12_45),
  inference(avatar_component_clause,[],[f318])).
fof(f2687,plain,(
  iof(sK8(sK2,sK1),sK2) | ~spl12_198),
  inference(avatar_component_clause,[],[f2686])).
fof(f4847,plain,(
  spl12_88 | ~spl12_35 | ~spl12_63 | ~spl12_103),
  inference(avatar_split_clause,[],[f4827,f674,f411,f259,f557])).
fof(f557,plain,(
  spl12_88 <=> specializes(sK0,sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_88])])).
fof(f411,plain,(
  spl12_63 <=> type_(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_63])])).
fof(f674,plain,(
  spl12_103 <=> ! [X0] : (iof(sK11(sK0,X0),sK0) | specializes(sK0,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_103])])).
fof(f4827,plain,(
  specializes(sK0,sK0) | (~spl12_35 | ~spl12_63 | ~spl12_103)),
  inference(subsumption_resolution,[],[f679,f412])).
fof(f412,plain,(
  type_(sK0) | ~spl12_63),
  inference(avatar_component_clause,[],[f411])).
fof(f679,plain,(
  specializes(sK0,sK0) | ~type_(sK0) | (~spl12_35 | ~spl12_103)),
  inference(duplicate_literal_removal,[],[f677])).
fof(f677,plain,(
  specializes(sK0,sK0) | ~type_(sK0) | specializes(sK0,sK0) | (~spl12_35 | ~spl12_103)),
  inference(resolution,[],[f675,f260])).
fof(f675,plain,(
  ( ! [X0] : (iof(sK11(sK0,X0),sK0) | specializes(sK0,X0)) ) | ~spl12_103),
  inference(avatar_component_clause,[],[f674])).
fof(f4831,plain,(
  spl12_38 | ~spl12_15 | ~spl12_17 | ~spl12_37),
  inference(avatar_split_clause,[],[f3251,f270,f174,f165,f277])).
fof(f277,plain,(
  spl12_38 <=> specializes(sK7(sK1),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_38])])).
fof(f165,plain,(
  spl12_15 <=> ! [X0] : (iof(sK7(X0),X0) | ~type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_15])])).
fof(f174,plain,(
  spl12_17 <=> type_(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_17])])).
fof(f270,plain,(
  spl12_37 <=> ! [X0] : (specializes(X0,sK0) | ~iof(X0,sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_37])])).
fof(f3251,plain,(
  specializes(sK7(sK1),sK0) | (~spl12_15 | ~spl12_17 | ~spl12_37)),
  inference(subsumption_resolution,[],[f3250,f175])).
fof(f175,plain,(
  type_(sK1) | ~spl12_17),
  inference(avatar_component_clause,[],[f174])).
fof(f3250,plain,(
  specializes(sK7(sK1),sK0) | ~type_(sK1) | (~spl12_15 | ~spl12_37)),
  inference(resolution,[],[f271,f166])).
fof(f166,plain,(
  ( ! [X0] : (iof(sK7(X0),X0) | ~type_(X0)) ) | ~spl12_15),
  inference(avatar_component_clause,[],[f165])).
fof(f271,plain,(
  ( ! [X0] : (~iof(X0,sK1) | specializes(X0,sK0)) ) | ~spl12_37),
  inference(avatar_component_clause,[],[f270])).
fof(f4826,plain,(
  ~spl12_1 | ~spl12_4 | ~spl12_93 | ~spl12_227),
  inference(avatar_contradiction_clause,[],[f4825])).
fof(f4825,plain,(
  $false | (~spl12_1 | ~spl12_4 | ~spl12_93 | ~spl12_227)),
  inference(subsumption_resolution,[],[f612,f4815])).
fof(f4815,plain,(
  ~iof(sK0,sK1) | (~spl12_4 | ~spl12_227)),
  inference(resolution,[],[f4810,f122])).
fof(f122,plain,(
  ( ! [X1] : (~properSpecializes(X1,X1)) ) | ~spl12_4),
  inference(avatar_component_clause,[],[f121])).
fof(f121,plain,(
  spl12_4 <=> ! [X1] : ~properSpecializes(X1,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_4])])).
fof(f4810,plain,(
  ( ! [X0] : (properSpecializes(X0,sK0) | ~iof(X0,sK1)) ) | ~spl12_227),
  inference(avatar_component_clause,[],[f4809])).
fof(f4809,plain,(
  spl12_227 <=> ! [X0] : (~iof(X0,sK1) | properSpecializes(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_227])])).
fof(f612,plain,(
  iof(sK0,sK1) | (~spl12_1 | ~spl12_93)),
  inference(resolution,[],[f601,f110])).
fof(f601,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK0,X0)) ) | ~spl12_93),
  inference(avatar_component_clause,[],[f600])).
fof(f600,plain,(
  spl12_93 <=> ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK0,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_93])])).
fof(f4811,plain,(
  spl12_227 | ~spl12_45 | ~spl12_197),
  inference(avatar_split_clause,[],[f4411,f2683,f318,f4809])).
fof(f2683,plain,(
  spl12_197 <=> sK1 = sK2),
  introduced(avatar_definition,[new_symbols(naming,[spl12_197])])).
fof(f4411,plain,(
  ( ! [X0] : (~iof(X0,sK1) | properSpecializes(X0,sK0)) ) | (~spl12_45 | ~spl12_197)),
  inference(backward_demodulation,[],[f319,f2684])).
fof(f2684,plain,(
  sK1 = sK2 | ~spl12_197),
  inference(avatar_component_clause,[],[f2683])).
fof(f3725,plain,(
  spl12_221 | ~spl12_31 | ~spl12_200),
  inference(avatar_split_clause,[],[f3396,f3388,f242,f3723])).
fof(f3396,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK8(sK2,sK1),X0)) ) | (~spl12_31 | ~spl12_200)),
  inference(resolution,[],[f3389,f243])).
fof(f3389,plain,(
  specializes(sK8(sK2,sK1),sK0) | ~spl12_200),
  inference(avatar_component_clause,[],[f3388])).
fof(f3707,plain,(
  spl12_197 | ~spl12_198 | ~spl12_67 | ~spl12_199),
  inference(avatar_split_clause,[],[f2692,f2689,f431,f2686,f2683])).
fof(f431,plain,(
  spl12_67 <=> ! [X1,X0] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_67])])).
fof(f2692,plain,(
  ~iof(sK8(sK2,sK1),sK2) | sK1 = sK2 | (~spl12_67 | ~spl12_199)),
  inference(resolution,[],[f2690,f432])).
fof(f432,plain,(
  ( ! [X0,X1] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) ) | ~spl12_67),
  inference(avatar_component_clause,[],[f431])).
fof(f2691,plain,(
  spl12_197 | spl12_198 | spl12_199 | ~spl12_18 | ~spl12_138),
  inference(avatar_split_clause,[],[f966,f956,f180,f2689,f2686,f2683])).
fof(f956,plain,(
  spl12_138 <=> ! [X1] : (~type_(X1) | iof(sK8(X1,sK1),sK1) | iof(sK8(X1,sK1),X1) | sK1 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_138])])).
fof(f966,plain,(
  iof(sK8(sK2,sK1),sK1) | iof(sK8(sK2,sK1),sK2) | sK1 = sK2 | (~spl12_18 | ~spl12_138)),
  inference(resolution,[],[f957,f181])).
fof(f957,plain,(
  ( ! [X1] : (~type_(X1) | iof(sK8(X1,sK1),sK1) | iof(sK8(X1,sK1),X1) | sK1 = X1) ) | ~spl12_138),
  inference(avatar_component_clause,[],[f956])).
fof(f958,plain,(
  spl12_138 | ~spl12_17 | ~spl12_94),
  inference(avatar_split_clause,[],[f636,f609,f174,f956])).
fof(f609,plain,(
  spl12_94 <=> ! [X1,X0] : (~type_(X0) | ~type_(X1) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_94])])).
fof(f636,plain,(
  ( ! [X1] : (~type_(X1) | iof(sK8(X1,sK1),sK1) | iof(sK8(X1,sK1),X1) | sK1 = X1) ) | (~spl12_17 | ~spl12_94)),
  inference(resolution,[],[f610,f175])).
fof(f610,plain,(
  ( ! [X0,X1] : (~type_(X1) | ~type_(X0) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1) ) | ~spl12_94),
  inference(avatar_component_clause,[],[f609])).
fof(f676,plain,(
  spl12_103 | ~spl12_34 | ~spl12_63),
  inference(avatar_split_clause,[],[f414,f411,f255,f674])).
fof(f255,plain,(
  spl12_34 <=> ! [X1,X0] : (~type_(X0) | iof(sK11(X0,X1),X0) | specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_34])])).
fof(f414,plain,(
  ( ! [X0] : (iof(sK11(sK0,X0),sK0) | specializes(sK0,X0)) ) | (~spl12_34 | ~spl12_63)),
  inference(resolution,[],[f412,f256])).
fof(f256,plain,(
  ( ! [X0,X1] : (~type_(X0) | iof(sK11(X0,X1),X0) | specializes(X0,X1)) ) | ~spl12_34),
  inference(avatar_component_clause,[],[f255])).
fof(f645,plain,(
  spl12_99 | ~spl12_16 | ~spl12_98),
  inference(avatar_split_clause,[],[f641,f632,f169,f643])).
fof(f632,plain,(
  spl12_98 <=> ! [X0] : (specializes(sK2,X0) | properSpecializes(sK11(sK2,X0),sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_98])])).
fof(f641,plain,(
  ( ! [X0] : (specializes(sK2,X0) | specializes(sK11(sK2,X0),sK0)) ) | (~spl12_16 | ~spl12_98)),
  inference(resolution,[],[f633,f170])).
fof(f633,plain,(
  ( ! [X0] : (properSpecializes(sK11(sK2,X0),sK0) | specializes(sK2,X0)) ) | ~spl12_98),
  inference(avatar_component_clause,[],[f632])).
fof(f634,plain,(
  spl12_98 | ~spl12_45 | ~spl12_82),
  inference(avatar_split_clause,[],[f515,f511,f318,f632])).
fof(f511,plain,(
  spl12_82 <=> ! [X1] : (iof(sK11(sK2,X1),sK2) | specializes(sK2,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_82])])).
fof(f515,plain,(
  ( ! [X0] : (specializes(sK2,X0) | properSpecializes(sK11(sK2,X0),sK0)) ) | (~spl12_45 | ~spl12_82)),
  inference(resolution,[],[f512,f319])).
fof(f512,plain,(
  ( ! [X1] : (iof(sK11(sK2,X1),sK2) | specializes(sK2,X1)) ) | ~spl12_82),
  inference(avatar_component_clause,[],[f511])).
fof(f611,plain,(
  spl12_94),
  inference(avatar_split_clause,[],[f76,f609])).
fof(f76,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~type_(X1) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ((X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))) | ~type_(X1) | ~type_(X0))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : ((X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))) | (~type_(X1) | ~type_(X0)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : ((type_(X1) & type_(X0)) => (X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7] : ((type_(X7) & type_(X6)) => (X6 = X7 <=> ! [X0] : (iof(X0,X6) <=> iof(X0,X7))))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a9_equality)).
fof(f602,plain,(
  spl12_93 | ~spl12_31 | ~spl12_88),
  inference(avatar_split_clause,[],[f560,f557,f242,f600])).
fof(f560,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK0,X0)) ) | (~spl12_31 | ~spl12_88)),
  inference(resolution,[],[f558,f243])).
fof(f558,plain,(
  specializes(sK0,sK0) | ~spl12_88),
  inference(avatar_component_clause,[],[f557])).
fof(f513,plain,(
  spl12_82 | ~spl12_18 | ~spl12_34),
  inference(avatar_split_clause,[],[f285,f255,f180,f511])).
fof(f285,plain,(
  ( ! [X1] : (iof(sK11(sK2,X1),sK2) | specializes(sK2,X1)) ) | (~spl12_18 | ~spl12_34)),
  inference(resolution,[],[f256,f181])).
fof(f433,plain,(
  spl12_67),
  inference(avatar_split_clause,[],[f104,f431])).
fof(f104,plain,(
  ( ! [X0,X1] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(subsumption_resolution,[],[f103,f70])).
fof(f70,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a3)).
fof(f103,plain,(
  ( ! [X0,X1] : (~type_(X1) | ~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(subsumption_resolution,[],[f77,f70])).
fof(f77,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~type_(X1) | ~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f48])).
fof(f413,plain,(
  spl12_63 | ~spl12_8 | ~spl12_61),
  inference(avatar_split_clause,[],[f409,f401,f137,f411])).
fof(f137,plain,(
  spl12_8 <=> ! [X1,X0] : (~iof(X1,X0) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_8])])).
fof(f401,plain,(
  spl12_61 <=> iof(sK7(sK7(sK1)),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_61])])).
fof(f409,plain,(
  type_(sK0) | (~spl12_8 | ~spl12_61)),
  inference(resolution,[],[f402,f138])).
fof(f138,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) ) | ~spl12_8),
  inference(avatar_component_clause,[],[f137])).
fof(f402,plain,(
  iof(sK7(sK7(sK1)),sK0) | ~spl12_61),
  inference(avatar_component_clause,[],[f401])).
fof(f403,plain,(
  spl12_61 | ~spl12_15 | ~spl12_39 | ~spl12_60),
  inference(avatar_split_clause,[],[f399,f395,f287,f165,f401])).
fof(f287,plain,(
  spl12_39 <=> type_(sK7(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_39])])).
fof(f395,plain,(
  spl12_60 <=> ! [X1] : (iof(X1,sK0) | ~iof(X1,sK7(sK1)))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_60])])).
fof(f399,plain,(
  iof(sK7(sK7(sK1)),sK0) | (~spl12_15 | ~spl12_39 | ~spl12_60)),
  inference(subsumption_resolution,[],[f398,f288])).
fof(f288,plain,(
  type_(sK7(sK1)) | ~spl12_39),
  inference(avatar_component_clause,[],[f287])).
fof(f398,plain,(
  iof(sK7(sK7(sK1)),sK0) | ~type_(sK7(sK1)) | (~spl12_15 | ~spl12_60)),
  inference(resolution,[],[f396,f166])).
fof(f396,plain,(
  ( ! [X1] : (~iof(X1,sK7(sK1)) | iof(X1,sK0)) ) | ~spl12_60),
  inference(avatar_component_clause,[],[f395])).
fof(f397,plain,(
  spl12_60 | ~spl12_29 | ~spl12_38),
  inference(avatar_split_clause,[],[f282,f277,f234,f395])).
fof(f234,plain,(
  spl12_29 <=> ! [X1,X0,X2] : (~iof(X2,X0) | iof(X2,X1) | ~specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_29])])).
fof(f282,plain,(
  ( ! [X1] : (iof(X1,sK0) | ~iof(X1,sK7(sK1))) ) | (~spl12_29 | ~spl12_38)),
  inference(resolution,[],[f278,f235])).
fof(f235,plain,(
  ( ! [X2,X0,X1] : (~specializes(X0,X1) | iof(X2,X1) | ~iof(X2,X0)) ) | ~spl12_29),
  inference(avatar_component_clause,[],[f234])).
fof(f278,plain,(
  specializes(sK7(sK1),sK0) | ~spl12_38),
  inference(avatar_component_clause,[],[f277])).
fof(f320,plain,(
  spl12_45 | ~spl12_2 | ~spl12_33),
  inference(avatar_split_clause,[],[f273,f250,f113,f318])).
fof(f113,plain,(
  spl12_2 <=> categorizes(sK2,sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_2])])).
fof(f250,plain,(
  spl12_33 <=> ! [X1,X0,X2] : (~categorizes(X0,X1) | ~iof(X2,X0) | properSpecializes(X2,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_33])])).
fof(f273,plain,(
  ( ! [X0] : (~iof(X0,sK2) | properSpecializes(X0,sK0)) ) | (~spl12_2 | ~spl12_33)),
  inference(resolution,[],[f251,f114])).
fof(f114,plain,(
  categorizes(sK2,sK0) | ~spl12_2),
  inference(avatar_component_clause,[],[f113])).
fof(f251,plain,(
  ( ! [X2,X0,X1] : (~categorizes(X0,X1) | ~iof(X2,X0) | properSpecializes(X2,X1)) ) | ~spl12_33),
  inference(avatar_component_clause,[],[f250])).
fof(f289,plain,(
  spl12_39 | ~spl12_11 | ~spl12_38),
  inference(avatar_split_clause,[],[f283,f277,f149,f287])).
fof(f149,plain,(
  spl12_11 <=> ! [X1,X0] : (type_(X0) | ~specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_11])])).
fof(f283,plain,(
  type_(sK7(sK1)) | (~spl12_11 | ~spl12_38)),
  inference(resolution,[],[f278,f150])).
fof(f150,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) ) | ~spl12_11),
  inference(avatar_component_clause,[],[f149])).
fof(f272,plain,(
  spl12_37 | ~spl12_1 | ~spl12_32),
  inference(avatar_split_clause,[],[f268,f246,f109,f270])).
fof(f246,plain,(
  spl12_32 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_32])])).
fof(f268,plain,(
  ( ! [X0] : (specializes(X0,sK0) | ~iof(X0,sK1)) ) | (~spl12_1 | ~spl12_32)),
  inference(resolution,[],[f247,f110])).
fof(f247,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) ) | ~spl12_32),
  inference(avatar_component_clause,[],[f246])).
fof(f261,plain,(
  spl12_35),
  inference(avatar_split_clause,[],[f87,f259])).
fof(f87,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~iof(sK11(X0,X1),X1) | specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X1) | ~iof(X2,X0)) & type_(X0)))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X0) => iof(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7] : (specializes(X6,X7) <=> (! [X3] : (iof(X3,X6) => iof(X3,X7)) & type_(X6)))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a7_specDef)).
fof(f257,plain,(
  spl12_34),
  inference(avatar_split_clause,[],[f86,f255])).
fof(f86,plain,(
  ( ! [X0,X1] : (~type_(X0) | iof(sK11(X0,X1),X0) | specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f53])).
fof(f252,plain,(
  spl12_33),
  inference(avatar_split_clause,[],[f96,f250])).
fof(f96,plain,(
  ( ! [X2,X0,X1] : (~categorizes(X0,X1) | ~iof(X2,X0) | properSpecializes(X2,X1)) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : ((! [X2] : (properSpecializes(X2,X1) | ~iof(X2,X0)) & type_(X0)) | ~categorizes(X0,X1))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : (categorizes(X0,X1) => (! [X2] : (iof(X2,X0) => properSpecializes(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (categorizes(X0,X1) <=> (! [X2] : (iof(X2,X0) => properSpecializes(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6,X7] : (categorizes(X6,X7) <=> (! [X8] : (iof(X8,X6) => properSpecializes(X8,X7)) & type_(X6)))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a11_categorizationDef)).
fof(f248,plain,(
  spl12_32),
  inference(avatar_split_clause,[],[f94,f246])).
fof(f94,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : ((! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)) | ~isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) => (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) <=> (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : (isPowertypeOf(X6,X7) <=> (! [X8] : (iof(X8,X6) <=> specializes(X8,X7)) & type_(X6)))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a10_isPowertypeOfDef)).
fof(f244,plain,(
  spl12_31),
  inference(avatar_split_clause,[],[f93,f242])).
fof(f93,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | ~specializes(X2,X1) | iof(X2,X0)) )),
  inference(cnf_transformation,[],[f54])).
fof(f240,plain,(
  spl12_30),
  inference(avatar_split_clause,[],[f92,f238])).
fof(f92,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | X0 = X1 | properSpecializes(X0,X1)) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1] : (properSpecializes(X0,X1) <=> (X0 != X1 & specializes(X0,X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X7] : (properSpecializes(X6,X7) <=> (X6 != X7 & specializes(X6,X7)))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a8_properSpecDef)).
fof(f236,plain,(
  spl12_29),
  inference(avatar_split_clause,[],[f88,f234])).
fof(f88,plain,(
  ( ! [X2,X0,X1] : (~iof(X2,X0) | iof(X2,X1) | ~specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f53])).
fof(f182,plain,(
  spl12_18 | ~spl12_2 | ~spl12_13),
  inference(avatar_split_clause,[],[f177,f157,f113,f180])).
fof(f157,plain,(
  spl12_13 <=> ! [X1,X0] : (~categorizes(X0,X1) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_13])])).
fof(f177,plain,(
  type_(sK2) | (~spl12_2 | ~spl12_13)),
  inference(resolution,[],[f158,f114])).
fof(f158,plain,(
  ( ! [X0,X1] : (~categorizes(X0,X1) | type_(X0)) ) | ~spl12_13),
  inference(avatar_component_clause,[],[f157])).
fof(f176,plain,(
  spl12_17 | ~spl12_1 | ~spl12_12),
  inference(avatar_split_clause,[],[f172,f153,f109,f174])).
fof(f153,plain,(
  spl12_12 <=> ! [X1,X0] : (~isPowertypeOf(X0,X1) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_12])])).
fof(f172,plain,(
  type_(sK1) | (~spl12_1 | ~spl12_12)),
  inference(resolution,[],[f154,f110])).
fof(f154,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X0,X1) | type_(X0)) ) | ~spl12_12),
  inference(avatar_component_clause,[],[f153])).
fof(f171,plain,(
  spl12_16),
  inference(avatar_split_clause,[],[f90,f169])).
fof(f90,plain,(
  ( ! [X0,X1] : (specializes(X0,X1) | ~properSpecializes(X0,X1)) )),
  inference(cnf_transformation,[],[f32])).
fof(f167,plain,(
  spl12_15),
  inference(avatar_split_clause,[],[f71,f165])).
fof(f71,plain,(
  ( ! [X0] : (iof(sK7(X0),X0) | ~type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f159,plain,(
  spl12_13),
  inference(avatar_split_clause,[],[f97,f157])).
fof(f97,plain,(
  ( ! [X0,X1] : (~categorizes(X0,X1) | type_(X0)) )),
  inference(cnf_transformation,[],[f55])).
fof(f155,plain,(
  spl12_12),
  inference(avatar_split_clause,[],[f95,f153])).
fof(f95,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X0,X1) | type_(X0)) )),
  inference(cnf_transformation,[],[f54])).
fof(f151,plain,(
  spl12_11),
  inference(avatar_split_clause,[],[f89,f149])).
fof(f89,plain,(
  ( ! [X0,X1] : (type_(X0) | ~specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f53])).
fof(f139,plain,(
  spl12_8),
  inference(avatar_split_clause,[],[f70,f137])).
fof(f123,plain,(
  spl12_4),
  inference(avatar_split_clause,[],[f102,f121])).
fof(f102,plain,(
  ( ! [X1] : (~properSpecializes(X1,X1)) )),
  inference(equality_resolution,[],[f91])).
fof(f91,plain,(
  ( ! [X0,X1] : (X0 != X1 | ~properSpecializes(X0,X1)) )),
  inference(cnf_transformation,[],[f32])).
fof(f119,plain,(
  ~spl12_3),
  inference(avatar_split_clause,[],[f58,f117])).
fof(f58,plain,(
  ~properSpecializes(sK2,sK1)),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ? [X0,X1,X2] : (~properSpecializes(X2,X1) & categorizes(X2,X0) & isPowertypeOf(X1,X0))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ? [X0,X1,X2] : (~properSpecializes(X2,X1) & (categorizes(X2,X0) & isPowertypeOf(X1,X0)))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1,X2] : ((categorizes(X2,X0) & isPowertypeOf(X1,X0)) => properSpecializes(X2,X1))),
  inference(rectify,[],[f22])).
fof(f22,negated_conjecture,(
  ~! [X6,X7,X8] : ((categorizes(X8,X6) & isPowertypeOf(X7,X6)) => properSpecializes(X8,X7))),
  inference(negated_conjecture,[],[f21])).
fof(f21,conjecture,(
  ! [X6,X7,X8] : ((categorizes(X8,X6) & isPowertypeOf(X7,X6)) => properSpecializes(X8,X7))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',t4)).
fof(f115,plain,(
  spl12_2),
  inference(avatar_split_clause,[],[f57,f113])).
fof(f57,plain,(
  categorizes(sK2,sK0)),
  inference(cnf_transformation,[],[f40])).
fof(f111,plain,(
  spl12_1),
  inference(avatar_split_clause,[],[f56,f109])).
fof(f56,plain,(
  isPowertypeOf(sK1,sK0)),
  inference(cnf_transformation,[],[f40])).
% SZS output end Proof for SOT_pdlX
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 9850
% Time elapsed: 0.320 s
% ------------------------------
% ------------------------------
% Success in time 0.369 s
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_pdlX
% SZS output start Proof for SOT_pdlX
fof(f7063,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f111,f115,f119,f123,f139,f151,f155,f159,f167,f171,f176,f182,f236,f240,f244,f248,f252,f257,f261,f272,f289,f320,f397,f403,f413,f433,f513,f602,f611,f634,f645,f676,f958,f2691,f3707,f3725,f4811,f4826,f4831,f4847,f5004,f5006,f5125,f5827,f6100,f6109,f6627])).
fof(f6627,plain,(
  spl12_3 | ~spl12_30 | spl12_198 | ~spl12_199 | ~spl12_321),
  inference(avatar_contradiction_clause,[],[f6626])).
fof(f6626,plain,(
  $false | (spl12_3 | ~spl12_30 | spl12_198 | ~spl12_199 | ~spl12_321)),
  inference(subsumption_resolution,[],[f6269,f6255])).
fof(f6255,plain,(
  iof(sK8(sK1,sK1),sK1) | (spl12_3 | ~spl12_30 | ~spl12_199 | ~spl12_321)),
  inference(backward_demodulation,[],[f2690,f6123])).
fof(f6123,plain,(
  sK1 = sK2 | (spl12_3 | ~spl12_30 | ~spl12_321)),
  inference(subsumption_resolution,[],[f6120,f118])).
fof(f118,plain,(
  ~properSpecializes(sK2,sK1) | spl12_3),
  inference(avatar_component_clause,[],[f117])).
fof(f117,plain,(
  spl12_3 <=> properSpecializes(sK2,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_3])])).
fof(f6120,plain,(
  sK1 = sK2 | properSpecializes(sK2,sK1) | (~spl12_30 | ~spl12_321)),
  inference(resolution,[],[f6108,f239])).
fof(f239,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | X0 = X1 | properSpecializes(X0,X1)) ) | ~spl12_30),
  inference(avatar_component_clause,[],[f238])).
fof(f238,plain,(
  spl12_30 <=> ! [X1,X0] : (~specializes(X0,X1) | X0 = X1 | properSpecializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_30])])).
fof(f6108,plain,(
  specializes(sK2,sK1) | ~spl12_321),
  inference(avatar_component_clause,[],[f6107])).
fof(f6107,plain,(
  spl12_321 <=> specializes(sK2,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_321])])).
fof(f2690,plain,(
  iof(sK8(sK2,sK1),sK1) | ~spl12_199),
  inference(avatar_component_clause,[],[f2689])).
fof(f2689,plain,(
  spl12_199 <=> iof(sK8(sK2,sK1),sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_199])])).
fof(f6269,plain,(
  ~iof(sK8(sK1,sK1),sK1) | (spl12_3 | ~spl12_30 | spl12_198 | ~spl12_321)),
  inference(backward_demodulation,[],[f3706,f6123])).
fof(f3706,plain,(
  ~iof(sK8(sK2,sK1),sK2) | spl12_198),
  inference(avatar_component_clause,[],[f2686])).
fof(f2686,plain,(
  spl12_198 <=> iof(sK8(sK2,sK1),sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_198])])).
fof(f6109,plain,(
  spl12_321 | ~spl12_18 | ~spl12_35 | ~spl12_320),
  inference(avatar_split_clause,[],[f6105,f6098,f259,f180,f6107])).
fof(f180,plain,(
  spl12_18 <=> type_(sK2)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_18])])).
fof(f259,plain,(
  spl12_35 <=> ! [X1,X0] : (~type_(X0) | ~iof(sK11(X0,X1),X1) | specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_35])])).
fof(f6098,plain,(
  spl12_320 <=> ! [X0] : (specializes(sK2,X0) | iof(sK11(sK2,X0),sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_320])])).
fof(f6105,plain,(
  specializes(sK2,sK1) | (~spl12_18 | ~spl12_35 | ~spl12_320)),
  inference(subsumption_resolution,[],[f6104,f181])).
fof(f181,plain,(
  type_(sK2) | ~spl12_18),
  inference(avatar_component_clause,[],[f180])).
fof(f6104,plain,(
  specializes(sK2,sK1) | ~type_(sK2) | (~spl12_35 | ~spl12_320)),
  inference(duplicate_literal_removal,[],[f6101])).
fof(f6101,plain,(
  specializes(sK2,sK1) | ~type_(sK2) | specializes(sK2,sK1) | (~spl12_35 | ~spl12_320)),
  inference(resolution,[],[f6099,f260])).
fof(f260,plain,(
  ( ! [X0,X1] : (~iof(sK11(X0,X1),X1) | ~type_(X0) | specializes(X0,X1)) ) | ~spl12_35),
  inference(avatar_component_clause,[],[f259])).
fof(f6099,plain,(
  ( ! [X0] : (iof(sK11(sK2,X0),sK1) | specializes(sK2,X0)) ) | ~spl12_320),
  inference(avatar_component_clause,[],[f6098])).
fof(f6100,plain,(
  spl12_320 | ~spl12_1 | ~spl12_289),
  inference(avatar_split_clause,[],[f5828,f5825,f109,f6098])).
fof(f109,plain,(
  spl12_1 <=> isPowertypeOf(sK1,sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_1])])).
fof(f5825,plain,(
  spl12_289 <=> ! [X1,X0] : (specializes(sK2,X0) | ~isPowertypeOf(X1,sK0) | iof(sK11(sK2,X0),X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_289])])).
fof(f5828,plain,(
  ( ! [X0] : (specializes(sK2,X0) | iof(sK11(sK2,X0),sK1)) ) | (~spl12_1 | ~spl12_289)),
  inference(resolution,[],[f5826,f110])).
fof(f110,plain,(
  isPowertypeOf(sK1,sK0) | ~spl12_1),
  inference(avatar_component_clause,[],[f109])).
fof(f5826,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X1,sK0) | specializes(sK2,X0) | iof(sK11(sK2,X0),X1)) ) | ~spl12_289),
  inference(avatar_component_clause,[],[f5825])).
fof(f5827,plain,(
  spl12_289 | ~spl12_31 | ~spl12_99),
  inference(avatar_split_clause,[],[f646,f643,f242,f5825])).
fof(f242,plain,(
  spl12_31 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | ~specializes(X2,X1) | iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_31])])).
fof(f643,plain,(
  spl12_99 <=> ! [X0] : (specializes(sK2,X0) | specializes(sK11(sK2,X0),sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_99])])).
fof(f646,plain,(
  ( ! [X0,X1] : (specializes(sK2,X0) | ~isPowertypeOf(X1,sK0) | iof(sK11(sK2,X0),X1)) ) | (~spl12_31 | ~spl12_99)),
  inference(resolution,[],[f644,f243])).
fof(f243,plain,(
  ( ! [X2,X0,X1] : (~specializes(X2,X1) | ~isPowertypeOf(X0,X1) | iof(X2,X0)) ) | ~spl12_31),
  inference(avatar_component_clause,[],[f242])).
fof(f644,plain,(
  ( ! [X0] : (specializes(sK11(sK2,X0),sK0) | specializes(sK2,X0)) ) | ~spl12_99),
  inference(avatar_component_clause,[],[f643])).
fof(f5125,plain,(
  spl12_199 | ~spl12_1 | ~spl12_221),
  inference(avatar_split_clause,[],[f5124,f3723,f109,f2689])).
fof(f3723,plain,(
  spl12_221 <=> ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK8(sK2,sK1),X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_221])])).
fof(f5124,plain,(
  iof(sK8(sK2,sK1),sK1) | (~spl12_1 | ~spl12_221)),
  inference(resolution,[],[f3724,f110])).
fof(f3724,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK8(sK2,sK1),X0)) ) | ~spl12_221),
  inference(avatar_component_clause,[],[f3723])).
fof(f5006,plain,(
  spl12_200 | ~spl12_16 | ~spl12_229),
  inference(avatar_split_clause,[],[f5005,f5002,f169,f3388])).
fof(f3388,plain,(
  spl12_200 <=> specializes(sK8(sK2,sK1),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_200])])).
fof(f169,plain,(
  spl12_16 <=> ! [X1,X0] : (specializes(X0,X1) | ~properSpecializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_16])])).
fof(f5002,plain,(
  spl12_229 <=> properSpecializes(sK8(sK2,sK1),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_229])])).
fof(f5005,plain,(
  specializes(sK8(sK2,sK1),sK0) | (~spl12_16 | ~spl12_229)),
  inference(resolution,[],[f5003,f170])).
fof(f170,plain,(
  ( ! [X0,X1] : (~properSpecializes(X0,X1) | specializes(X0,X1)) ) | ~spl12_16),
  inference(avatar_component_clause,[],[f169])).
fof(f5003,plain,(
  properSpecializes(sK8(sK2,sK1),sK0) | ~spl12_229),
  inference(avatar_component_clause,[],[f5002])).
fof(f5004,plain,(
  spl12_229 | ~spl12_45 | ~spl12_198),
  inference(avatar_split_clause,[],[f4943,f2686,f318,f5002])).
fof(f318,plain,(
  spl12_45 <=> ! [X0] : (~iof(X0,sK2) | properSpecializes(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_45])])).
fof(f4943,plain,(
  properSpecializes(sK8(sK2,sK1),sK0) | (~spl12_45 | ~spl12_198)),
  inference(resolution,[],[f2687,f319])).
fof(f319,plain,(
  ( ! [X0] : (~iof(X0,sK2) | properSpecializes(X0,sK0)) ) | ~spl12_45),
  inference(avatar_component_clause,[],[f318])).
fof(f2687,plain,(
  iof(sK8(sK2,sK1),sK2) | ~spl12_198),
  inference(avatar_component_clause,[],[f2686])).
fof(f4847,plain,(
  spl12_88 | ~spl12_35 | ~spl12_63 | ~spl12_103),
  inference(avatar_split_clause,[],[f4827,f674,f411,f259,f557])).
fof(f557,plain,(
  spl12_88 <=> specializes(sK0,sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_88])])).
fof(f411,plain,(
  spl12_63 <=> type_(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_63])])).
fof(f674,plain,(
  spl12_103 <=> ! [X0] : (iof(sK11(sK0,X0),sK0) | specializes(sK0,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_103])])).
fof(f4827,plain,(
  specializes(sK0,sK0) | (~spl12_35 | ~spl12_63 | ~spl12_103)),
  inference(subsumption_resolution,[],[f679,f412])).
fof(f412,plain,(
  type_(sK0) | ~spl12_63),
  inference(avatar_component_clause,[],[f411])).
fof(f679,plain,(
  specializes(sK0,sK0) | ~type_(sK0) | (~spl12_35 | ~spl12_103)),
  inference(duplicate_literal_removal,[],[f677])).
fof(f677,plain,(
  specializes(sK0,sK0) | ~type_(sK0) | specializes(sK0,sK0) | (~spl12_35 | ~spl12_103)),
  inference(resolution,[],[f675,f260])).
fof(f675,plain,(
  ( ! [X0] : (iof(sK11(sK0,X0),sK0) | specializes(sK0,X0)) ) | ~spl12_103),
  inference(avatar_component_clause,[],[f674])).
fof(f4831,plain,(
  spl12_38 | ~spl12_15 | ~spl12_17 | ~spl12_37),
  inference(avatar_split_clause,[],[f3251,f270,f174,f165,f277])).
fof(f277,plain,(
  spl12_38 <=> specializes(sK7(sK1),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_38])])).
fof(f165,plain,(
  spl12_15 <=> ! [X0] : (iof(sK7(X0),X0) | ~type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_15])])).
fof(f174,plain,(
  spl12_17 <=> type_(sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_17])])).
fof(f270,plain,(
  spl12_37 <=> ! [X0] : (specializes(X0,sK0) | ~iof(X0,sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_37])])).
fof(f3251,plain,(
  specializes(sK7(sK1),sK0) | (~spl12_15 | ~spl12_17 | ~spl12_37)),
  inference(subsumption_resolution,[],[f3250,f175])).
fof(f175,plain,(
  type_(sK1) | ~spl12_17),
  inference(avatar_component_clause,[],[f174])).
fof(f3250,plain,(
  specializes(sK7(sK1),sK0) | ~type_(sK1) | (~spl12_15 | ~spl12_37)),
  inference(resolution,[],[f271,f166])).
fof(f166,plain,(
  ( ! [X0] : (iof(sK7(X0),X0) | ~type_(X0)) ) | ~spl12_15),
  inference(avatar_component_clause,[],[f165])).
fof(f271,plain,(
  ( ! [X0] : (~iof(X0,sK1) | specializes(X0,sK0)) ) | ~spl12_37),
  inference(avatar_component_clause,[],[f270])).
fof(f4826,plain,(
  ~spl12_1 | ~spl12_4 | ~spl12_93 | ~spl12_227),
  inference(avatar_contradiction_clause,[],[f4825])).
fof(f4825,plain,(
  $false | (~spl12_1 | ~spl12_4 | ~spl12_93 | ~spl12_227)),
  inference(subsumption_resolution,[],[f612,f4815])).
fof(f4815,plain,(
  ~iof(sK0,sK1) | (~spl12_4 | ~spl12_227)),
  inference(resolution,[],[f4810,f122])).
fof(f122,plain,(
  ( ! [X1] : (~properSpecializes(X1,X1)) ) | ~spl12_4),
  inference(avatar_component_clause,[],[f121])).
fof(f121,plain,(
  spl12_4 <=> ! [X1] : ~properSpecializes(X1,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_4])])).
fof(f4810,plain,(
  ( ! [X0] : (properSpecializes(X0,sK0) | ~iof(X0,sK1)) ) | ~spl12_227),
  inference(avatar_component_clause,[],[f4809])).
fof(f4809,plain,(
  spl12_227 <=> ! [X0] : (~iof(X0,sK1) | properSpecializes(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_227])])).
fof(f612,plain,(
  iof(sK0,sK1) | (~spl12_1 | ~spl12_93)),
  inference(resolution,[],[f601,f110])).
fof(f601,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK0,X0)) ) | ~spl12_93),
  inference(avatar_component_clause,[],[f600])).
fof(f600,plain,(
  spl12_93 <=> ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK0,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_93])])).
fof(f4811,plain,(
  spl12_227 | ~spl12_45 | ~spl12_197),
  inference(avatar_split_clause,[],[f4411,f2683,f318,f4809])).
fof(f2683,plain,(
  spl12_197 <=> sK1 = sK2),
  introduced(avatar_definition,[new_symbols(naming,[spl12_197])])).
fof(f4411,plain,(
  ( ! [X0] : (~iof(X0,sK1) | properSpecializes(X0,sK0)) ) | (~spl12_45 | ~spl12_197)),
  inference(backward_demodulation,[],[f319,f2684])).
fof(f2684,plain,(
  sK1 = sK2 | ~spl12_197),
  inference(avatar_component_clause,[],[f2683])).
fof(f3725,plain,(
  spl12_221 | ~spl12_31 | ~spl12_200),
  inference(avatar_split_clause,[],[f3396,f3388,f242,f3723])).
fof(f3396,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK8(sK2,sK1),X0)) ) | (~spl12_31 | ~spl12_200)),
  inference(resolution,[],[f3389,f243])).
fof(f3389,plain,(
  specializes(sK8(sK2,sK1),sK0) | ~spl12_200),
  inference(avatar_component_clause,[],[f3388])).
fof(f3707,plain,(
  spl12_197 | ~spl12_198 | ~spl12_67 | ~spl12_199),
  inference(avatar_split_clause,[],[f2692,f2689,f431,f2686,f2683])).
fof(f431,plain,(
  spl12_67 <=> ! [X1,X0] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_67])])).
fof(f2692,plain,(
  ~iof(sK8(sK2,sK1),sK2) | sK1 = sK2 | (~spl12_67 | ~spl12_199)),
  inference(resolution,[],[f2690,f432])).
fof(f432,plain,(
  ( ! [X0,X1] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) ) | ~spl12_67),
  inference(avatar_component_clause,[],[f431])).
fof(f2691,plain,(
  spl12_197 | spl12_198 | spl12_199 | ~spl12_18 | ~spl12_138),
  inference(avatar_split_clause,[],[f966,f956,f180,f2689,f2686,f2683])).
fof(f956,plain,(
  spl12_138 <=> ! [X1] : (~type_(X1) | iof(sK8(X1,sK1),sK1) | iof(sK8(X1,sK1),X1) | sK1 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_138])])).
fof(f966,plain,(
  iof(sK8(sK2,sK1),sK1) | iof(sK8(sK2,sK1),sK2) | sK1 = sK2 | (~spl12_18 | ~spl12_138)),
  inference(resolution,[],[f957,f181])).
fof(f957,plain,(
  ( ! [X1] : (~type_(X1) | iof(sK8(X1,sK1),sK1) | iof(sK8(X1,sK1),X1) | sK1 = X1) ) | ~spl12_138),
  inference(avatar_component_clause,[],[f956])).
fof(f958,plain,(
  spl12_138 | ~spl12_17 | ~spl12_94),
  inference(avatar_split_clause,[],[f636,f609,f174,f956])).
fof(f609,plain,(
  spl12_94 <=> ! [X1,X0] : (~type_(X0) | ~type_(X1) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_94])])).
fof(f636,plain,(
  ( ! [X1] : (~type_(X1) | iof(sK8(X1,sK1),sK1) | iof(sK8(X1,sK1),X1) | sK1 = X1) ) | (~spl12_17 | ~spl12_94)),
  inference(resolution,[],[f610,f175])).
fof(f610,plain,(
  ( ! [X0,X1] : (~type_(X1) | ~type_(X0) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1) ) | ~spl12_94),
  inference(avatar_component_clause,[],[f609])).
fof(f676,plain,(
  spl12_103 | ~spl12_34 | ~spl12_63),
  inference(avatar_split_clause,[],[f414,f411,f255,f674])).
fof(f255,plain,(
  spl12_34 <=> ! [X1,X0] : (~type_(X0) | iof(sK11(X0,X1),X0) | specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_34])])).
fof(f414,plain,(
  ( ! [X0] : (iof(sK11(sK0,X0),sK0) | specializes(sK0,X0)) ) | (~spl12_34 | ~spl12_63)),
  inference(resolution,[],[f412,f256])).
fof(f256,plain,(
  ( ! [X0,X1] : (~type_(X0) | iof(sK11(X0,X1),X0) | specializes(X0,X1)) ) | ~spl12_34),
  inference(avatar_component_clause,[],[f255])).
fof(f645,plain,(
  spl12_99 | ~spl12_16 | ~spl12_98),
  inference(avatar_split_clause,[],[f641,f632,f169,f643])).
fof(f632,plain,(
  spl12_98 <=> ! [X0] : (specializes(sK2,X0) | properSpecializes(sK11(sK2,X0),sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_98])])).
fof(f641,plain,(
  ( ! [X0] : (specializes(sK2,X0) | specializes(sK11(sK2,X0),sK0)) ) | (~spl12_16 | ~spl12_98)),
  inference(resolution,[],[f633,f170])).
fof(f633,plain,(
  ( ! [X0] : (properSpecializes(sK11(sK2,X0),sK0) | specializes(sK2,X0)) ) | ~spl12_98),
  inference(avatar_component_clause,[],[f632])).
fof(f634,plain,(
  spl12_98 | ~spl12_45 | ~spl12_82),
  inference(avatar_split_clause,[],[f515,f511,f318,f632])).
fof(f511,plain,(
  spl12_82 <=> ! [X1] : (iof(sK11(sK2,X1),sK2) | specializes(sK2,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_82])])).
fof(f515,plain,(
  ( ! [X0] : (specializes(sK2,X0) | properSpecializes(sK11(sK2,X0),sK0)) ) | (~spl12_45 | ~spl12_82)),
  inference(resolution,[],[f512,f319])).
fof(f512,plain,(
  ( ! [X1] : (iof(sK11(sK2,X1),sK2) | specializes(sK2,X1)) ) | ~spl12_82),
  inference(avatar_component_clause,[],[f511])).
fof(f611,plain,(
  spl12_94),
  inference(avatar_split_clause,[],[f76,f609])).
fof(f76,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~type_(X1) | iof(sK8(X0,X1),X1) | iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ((X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))) | ~type_(X1) | ~type_(X0))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : ((X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))) | (~type_(X1) | ~type_(X0)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : ((type_(X1) & type_(X0)) => (X0 = X1 <=> ! [X2] : (iof(X2,X0) <=> iof(X2,X1))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7] : ((type_(X7) & type_(X6)) => (X6 = X7 <=> ! [X0] : (iof(X0,X6) <=> iof(X0,X7))))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a9_equality)).
fof(f602,plain,(
  spl12_93 | ~spl12_31 | ~spl12_88),
  inference(avatar_split_clause,[],[f560,f557,f242,f600])).
fof(f560,plain,(
  ( ! [X0] : (~isPowertypeOf(X0,sK0) | iof(sK0,X0)) ) | (~spl12_31 | ~spl12_88)),
  inference(resolution,[],[f558,f243])).
fof(f558,plain,(
  specializes(sK0,sK0) | ~spl12_88),
  inference(avatar_component_clause,[],[f557])).
fof(f513,plain,(
  spl12_82 | ~spl12_18 | ~spl12_34),
  inference(avatar_split_clause,[],[f285,f255,f180,f511])).
fof(f285,plain,(
  ( ! [X1] : (iof(sK11(sK2,X1),sK2) | specializes(sK2,X1)) ) | (~spl12_18 | ~spl12_34)),
  inference(resolution,[],[f256,f181])).
fof(f433,plain,(
  spl12_67),
  inference(avatar_split_clause,[],[f104,f431])).
fof(f104,plain,(
  ( ! [X0,X1] : (~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(subsumption_resolution,[],[f103,f70])).
fof(f70,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a3)).
fof(f103,plain,(
  ( ! [X0,X1] : (~type_(X1) | ~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(subsumption_resolution,[],[f77,f70])).
fof(f77,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~type_(X1) | ~iof(sK8(X0,X1),X1) | ~iof(sK8(X0,X1),X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f48])).
fof(f413,plain,(
  spl12_63 | ~spl12_8 | ~spl12_61),
  inference(avatar_split_clause,[],[f409,f401,f137,f411])).
fof(f137,plain,(
  spl12_8 <=> ! [X1,X0] : (~iof(X1,X0) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_8])])).
fof(f401,plain,(
  spl12_61 <=> iof(sK7(sK7(sK1)),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_61])])).
fof(f409,plain,(
  type_(sK0) | (~spl12_8 | ~spl12_61)),
  inference(resolution,[],[f402,f138])).
fof(f138,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | type_(X0)) ) | ~spl12_8),
  inference(avatar_component_clause,[],[f137])).
fof(f402,plain,(
  iof(sK7(sK7(sK1)),sK0) | ~spl12_61),
  inference(avatar_component_clause,[],[f401])).
fof(f403,plain,(
  spl12_61 | ~spl12_15 | ~spl12_39 | ~spl12_60),
  inference(avatar_split_clause,[],[f399,f395,f287,f165,f401])).
fof(f287,plain,(
  spl12_39 <=> type_(sK7(sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_39])])).
fof(f395,plain,(
  spl12_60 <=> ! [X1] : (iof(X1,sK0) | ~iof(X1,sK7(sK1)))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_60])])).
fof(f399,plain,(
  iof(sK7(sK7(sK1)),sK0) | (~spl12_15 | ~spl12_39 | ~spl12_60)),
  inference(subsumption_resolution,[],[f398,f288])).
fof(f288,plain,(
  type_(sK7(sK1)) | ~spl12_39),
  inference(avatar_component_clause,[],[f287])).
fof(f398,plain,(
  iof(sK7(sK7(sK1)),sK0) | ~type_(sK7(sK1)) | (~spl12_15 | ~spl12_60)),
  inference(resolution,[],[f396,f166])).
fof(f396,plain,(
  ( ! [X1] : (~iof(X1,sK7(sK1)) | iof(X1,sK0)) ) | ~spl12_60),
  inference(avatar_component_clause,[],[f395])).
fof(f397,plain,(
  spl12_60 | ~spl12_29 | ~spl12_38),
  inference(avatar_split_clause,[],[f282,f277,f234,f395])).
fof(f234,plain,(
  spl12_29 <=> ! [X1,X0,X2] : (~iof(X2,X0) | iof(X2,X1) | ~specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_29])])).
fof(f282,plain,(
  ( ! [X1] : (iof(X1,sK0) | ~iof(X1,sK7(sK1))) ) | (~spl12_29 | ~spl12_38)),
  inference(resolution,[],[f278,f235])).
fof(f235,plain,(
  ( ! [X2,X0,X1] : (~specializes(X0,X1) | iof(X2,X1) | ~iof(X2,X0)) ) | ~spl12_29),
  inference(avatar_component_clause,[],[f234])).
fof(f278,plain,(
  specializes(sK7(sK1),sK0) | ~spl12_38),
  inference(avatar_component_clause,[],[f277])).
fof(f320,plain,(
  spl12_45 | ~spl12_2 | ~spl12_33),
  inference(avatar_split_clause,[],[f273,f250,f113,f318])).
fof(f113,plain,(
  spl12_2 <=> categorizes(sK2,sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl12_2])])).
fof(f250,plain,(
  spl12_33 <=> ! [X1,X0,X2] : (~categorizes(X0,X1) | ~iof(X2,X0) | properSpecializes(X2,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_33])])).
fof(f273,plain,(
  ( ! [X0] : (~iof(X0,sK2) | properSpecializes(X0,sK0)) ) | (~spl12_2 | ~spl12_33)),
  inference(resolution,[],[f251,f114])).
fof(f114,plain,(
  categorizes(sK2,sK0) | ~spl12_2),
  inference(avatar_component_clause,[],[f113])).
fof(f251,plain,(
  ( ! [X2,X0,X1] : (~categorizes(X0,X1) | ~iof(X2,X0) | properSpecializes(X2,X1)) ) | ~spl12_33),
  inference(avatar_component_clause,[],[f250])).
fof(f289,plain,(
  spl12_39 | ~spl12_11 | ~spl12_38),
  inference(avatar_split_clause,[],[f283,f277,f149,f287])).
fof(f149,plain,(
  spl12_11 <=> ! [X1,X0] : (type_(X0) | ~specializes(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_11])])).
fof(f283,plain,(
  type_(sK7(sK1)) | (~spl12_11 | ~spl12_38)),
  inference(resolution,[],[f278,f150])).
fof(f150,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) ) | ~spl12_11),
  inference(avatar_component_clause,[],[f149])).
fof(f272,plain,(
  spl12_37 | ~spl12_1 | ~spl12_32),
  inference(avatar_split_clause,[],[f268,f246,f109,f270])).
fof(f246,plain,(
  spl12_32 <=> ! [X1,X0,X2] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_32])])).
fof(f268,plain,(
  ( ! [X0] : (specializes(X0,sK0) | ~iof(X0,sK1)) ) | (~spl12_1 | ~spl12_32)),
  inference(resolution,[],[f247,f110])).
fof(f247,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) ) | ~spl12_32),
  inference(avatar_component_clause,[],[f246])).
fof(f261,plain,(
  spl12_35),
  inference(avatar_split_clause,[],[f87,f259])).
fof(f87,plain,(
  ( ! [X0,X1] : (~type_(X0) | ~iof(sK11(X0,X1),X1) | specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X1) | ~iof(X2,X0)) & type_(X0)))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X0) => iof(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7] : (specializes(X6,X7) <=> (! [X3] : (iof(X3,X6) => iof(X3,X7)) & type_(X6)))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a7_specDef)).
fof(f257,plain,(
  spl12_34),
  inference(avatar_split_clause,[],[f86,f255])).
fof(f86,plain,(
  ( ! [X0,X1] : (~type_(X0) | iof(sK11(X0,X1),X0) | specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f53])).
fof(f252,plain,(
  spl12_33),
  inference(avatar_split_clause,[],[f96,f250])).
fof(f96,plain,(
  ( ! [X2,X0,X1] : (~categorizes(X0,X1) | ~iof(X2,X0) | properSpecializes(X2,X1)) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : ((! [X2] : (properSpecializes(X2,X1) | ~iof(X2,X0)) & type_(X0)) | ~categorizes(X0,X1))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : (categorizes(X0,X1) => (! [X2] : (iof(X2,X0) => properSpecializes(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (categorizes(X0,X1) <=> (! [X2] : (iof(X2,X0) => properSpecializes(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6,X7] : (categorizes(X6,X7) <=> (! [X8] : (iof(X8,X6) => properSpecializes(X8,X7)) & type_(X6)))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a11_categorizationDef)).
fof(f248,plain,(
  spl12_32),
  inference(avatar_split_clause,[],[f94,f246])).
fof(f94,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | specializes(X2,X1) | ~iof(X2,X0)) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : ((! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)) | ~isPowertypeOf(X0,X1))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) => (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : (isPowertypeOf(X0,X1) <=> (! [X2] : (iof(X2,X0) <=> specializes(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : (isPowertypeOf(X6,X7) <=> (! [X8] : (iof(X8,X6) <=> specializes(X8,X7)) & type_(X6)))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a10_isPowertypeOfDef)).
fof(f244,plain,(
  spl12_31),
  inference(avatar_split_clause,[],[f93,f242])).
fof(f93,plain,(
  ( ! [X2,X0,X1] : (~isPowertypeOf(X0,X1) | ~specializes(X2,X1) | iof(X2,X0)) )),
  inference(cnf_transformation,[],[f54])).
fof(f240,plain,(
  spl12_30),
  inference(avatar_split_clause,[],[f92,f238])).
fof(f92,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | X0 = X1 | properSpecializes(X0,X1)) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1] : (properSpecializes(X0,X1) <=> (X0 != X1 & specializes(X0,X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X7] : (properSpecializes(X6,X7) <=> (X6 != X7 & specializes(X6,X7)))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',a8_properSpecDef)).
fof(f236,plain,(
  spl12_29),
  inference(avatar_split_clause,[],[f88,f234])).
fof(f88,plain,(
  ( ! [X2,X0,X1] : (~iof(X2,X0) | iof(X2,X1) | ~specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f53])).
fof(f182,plain,(
  spl12_18 | ~spl12_2 | ~spl12_13),
  inference(avatar_split_clause,[],[f177,f157,f113,f180])).
fof(f157,plain,(
  spl12_13 <=> ! [X1,X0] : (~categorizes(X0,X1) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_13])])).
fof(f177,plain,(
  type_(sK2) | (~spl12_2 | ~spl12_13)),
  inference(resolution,[],[f158,f114])).
fof(f158,plain,(
  ( ! [X0,X1] : (~categorizes(X0,X1) | type_(X0)) ) | ~spl12_13),
  inference(avatar_component_clause,[],[f157])).
fof(f176,plain,(
  spl12_17 | ~spl12_1 | ~spl12_12),
  inference(avatar_split_clause,[],[f172,f153,f109,f174])).
fof(f153,plain,(
  spl12_12 <=> ! [X1,X0] : (~isPowertypeOf(X0,X1) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_12])])).
fof(f172,plain,(
  type_(sK1) | (~spl12_1 | ~spl12_12)),
  inference(resolution,[],[f154,f110])).
fof(f154,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X0,X1) | type_(X0)) ) | ~spl12_12),
  inference(avatar_component_clause,[],[f153])).
fof(f171,plain,(
  spl12_16),
  inference(avatar_split_clause,[],[f90,f169])).
fof(f90,plain,(
  ( ! [X0,X1] : (specializes(X0,X1) | ~properSpecializes(X0,X1)) )),
  inference(cnf_transformation,[],[f32])).
fof(f167,plain,(
  spl12_15),
  inference(avatar_split_clause,[],[f71,f165])).
fof(f71,plain,(
  ( ! [X0] : (iof(sK7(X0),X0) | ~type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f159,plain,(
  spl12_13),
  inference(avatar_split_clause,[],[f97,f157])).
fof(f97,plain,(
  ( ! [X0,X1] : (~categorizes(X0,X1) | type_(X0)) )),
  inference(cnf_transformation,[],[f55])).
fof(f155,plain,(
  spl12_12),
  inference(avatar_split_clause,[],[f95,f153])).
fof(f95,plain,(
  ( ! [X0,X1] : (~isPowertypeOf(X0,X1) | type_(X0)) )),
  inference(cnf_transformation,[],[f54])).
fof(f151,plain,(
  spl12_11),
  inference(avatar_split_clause,[],[f89,f149])).
fof(f89,plain,(
  ( ! [X0,X1] : (type_(X0) | ~specializes(X0,X1)) )),
  inference(cnf_transformation,[],[f53])).
fof(f139,plain,(
  spl12_8),
  inference(avatar_split_clause,[],[f70,f137])).
fof(f123,plain,(
  spl12_4),
  inference(avatar_split_clause,[],[f102,f121])).
fof(f102,plain,(
  ( ! [X1] : (~properSpecializes(X1,X1)) )),
  inference(equality_resolution,[],[f91])).
fof(f91,plain,(
  ( ! [X0,X1] : (X0 != X1 | ~properSpecializes(X0,X1)) )),
  inference(cnf_transformation,[],[f32])).
fof(f119,plain,(
  ~spl12_3),
  inference(avatar_split_clause,[],[f58,f117])).
fof(f58,plain,(
  ~properSpecializes(sK2,sK1)),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ? [X0,X1,X2] : (~properSpecializes(X2,X1) & categorizes(X2,X0) & isPowertypeOf(X1,X0))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ? [X0,X1,X2] : (~properSpecializes(X2,X1) & (categorizes(X2,X0) & isPowertypeOf(X1,X0)))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1,X2] : ((categorizes(X2,X0) & isPowertypeOf(X1,X0)) => properSpecializes(X2,X1))),
  inference(rectify,[],[f22])).
fof(f22,negated_conjecture,(
  ~! [X6,X7,X8] : ((categorizes(X8,X6) & isPowertypeOf(X7,X6)) => properSpecializes(X8,X7))),
  inference(negated_conjecture,[],[f21])).
fof(f21,conjecture,(
  ! [X6,X7,X8] : ((categorizes(X8,X6) & isPowertypeOf(X7,X6)) => properSpecializes(X8,X7))),
  file('/tmp/wqEVxBSTg1/SOT_pdlXwd',t4)).
fof(f115,plain,(
  spl12_2),
  inference(avatar_split_clause,[],[f57,f113])).
fof(f57,plain,(
  categorizes(sK2,sK0)),
  inference(cnf_transformation,[],[f40])).
fof(f111,plain,(
  spl12_1),
  inference(avatar_split_clause,[],[f56,f109])).
fof(f56,plain,(
  isPowertypeOf(sK1,sK0)),
  inference(cnf_transformation,[],[f40])).
% SZS output end Proof for SOT_pdlX
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 9850
% Time elapsed: 0.299 s
% ------------------------------
% ------------------------------
% Success in time 0.362 s

% END OF SYSTEM OUTPUT
% RESULT: SOT_pdlXwd - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.27 
% OUTPUT: SOT_pdlXwd - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.27 

% END OF SYSTEM OUTPUT
% RESULT: SOT_pdlXwd - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.36 
% OUTPUT: SOT_pdlXwd - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.37 

% END OF SYSTEM OUTPUT
% RESULT: SOT_pdlXwd - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.39 
% OUTPUT: SOT_pdlXwd - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.39 

% END OF SYSTEM OUTPUT

% END OF SYSTEM OUTPUT

% END OF SYSTEM OUTPUT
% RESULT: SOT_pdlXwd - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.38 
% OUTPUT: SOT_pdlXwd - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.39 
% RESULT: SOT_pdlXwd - E---FNT-2.5 says Theorem - CPU = 0.00 WC = 0.07 
% OUTPUT: SOT_pdlXwd - E---FNT-2.5 says CNFRefutation - CPU = 0.00 WC = 0.07 
% RESULT: SOT_pdlXwd - E---FNT-2.5 says Theorem - CPU = 0.00 WC = 0.08 
% OUTPUT: SOT_pdlXwd - E---FNT-2.5 says CNFRefutation - CPU = 0.00 WC = 0.08 
Vampire---SAT-4.5 has ended (PID 97259)
Vampire---SAT-4.5 has ended (PID 97263)
Vampire---SAT-4.5 has ended (PID 97266)
Vampire---SAT-4.5 has ended (PID 97268)
E---FNT-2.5 has ended (PID 97270)
E---FNT-2.5 has ended (PID 97273)
RESULT: SOT_pdlXwd - SSCPA---0.0 says Theorem - CPU = 5.0 WC = 6.7  SolvingSystem = Vampire---SAT-4.5
</PRE>

</body>
</html>