<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>System Results</title>
<meta name="ROBOTS" content="INDEX,NOFOLLOW" />
<script type="text/javascript">//<![CDATA[

  var tstp_dump;
  function openSoTSTP (dump) {   
    var tstp_url = 'http://www.tptp.org/cgi-bin/SystemOnTSTP';
    var tstp_browser = window.open(tstp_url, '_blank');
    tstp_dump = dump;
  }
  function getTSTPDump () {
    return tstp_dump;
  }

//]]></script>
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>

% SZS start RequiredInformation<BR>
% Congratulations - you have become a registered power user of SystemOnTPTP, at IP address 191.6.38.118.<BR>
% Please consider donating to the TPTP project - see www.tptp.org for details.<BR>
% When you donate this message will disappear.<BR>
% If you do not donate a random delay might be added to your processing time.<BR>
% SZS end RequiredInformation<BR>
<PRE>
Getting system recommendations for generic problem /tmp/0CirYUgLs0/SOT_QlsZwL
This is a FOF_THM_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.95
Vampire---TF1-4.5    recommended  0.95
Vampire---4.5        recommended  0.95
Vampire---SAT-4.5    recommended  0.95
ET---2.0             recommended  0.80
E---2.5              recommended  0.78
E---FNT-2.5          recommended  0.78
CSE_E---1.0          recommended  0.76
E---FNT-2.5          recommended  0.74
Isabelle---2TH0      recommended  0.69
Isabelle---2TF0      recommended  0.69
Isabelle---2FOF      recommended  0.69
Isabelle---2016      recommended  0.69
Isabelle-HOT---2016  recommended  0.68
CVC4---1.7           recommended  0.68
CVC4---SAT-1.7       recommended  0.68
Leo-III---1.4        recommended  0.61
Vampire---THF-4.5    recommended  0.59
Vampire---TF1-4.5    recommended  0.59
Vampire---SAT-4.5    recommended  0.59
iProver---3.0        recommended  0.58
iProver---SAT-3.0    recommended  0.58
Fampire---1.3        recommended  0.57
SInE---0.4           recommended  0.53
iProver---SAT-3.0    recommended  0.52
ePrincess---1.0      recommended  0.50
Princess---170717    recommended  0.48
SPASS---3.9          recommended  0.48
Z3---4.4.1           recommended  0.46
Zipperpin---1.5      recommended  0.40
CVC4---SAT-1.7       recommended  0.40
Prover9---1109a      recommended  0.39
Equinox---6.0.1a     recommended  0.37
LEO-II---1.7.0       recommended  0.35
nanoCoP---1.1        recommended  0.34
Muscadet---4.5       recommended  0.25
Twee---2.2           recommended  0.23
Metis---2.4          recommended  0.23
Geo-III---2018C      recommended  0.18
iProverMo---2.5-0.1  recommended  0.14
Beagle---0.9.47      subsumed     0.41
leanCoP---2.2        subsumed     0.36
E-Darwin---1.5       subsumed     0.31
CSE---1.0            subsumed     0.28
Darwin---1.4.5       subsumed     0.28
Bliksem---1.12       subsumed     0.21
Zenon---0.7.1        subsumed     0.16
ZenonModulo---0.4.2  subsumed     0.15
SOS---2.0            subsumed     0.11
Otter---3.3          subsumed     0.04
This is a FOF_CSA_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.98
Vampire---TF1-4.5    recommended  0.98
Vampire---SAT-4.5    recommended  0.98
Paradox---4.0        recommended  0.60
Nitpick---2016       subsumed     0.55
DarwinFM---1.4.5     subsumed     0.43
CVC4---SAT-1.7       subsumed     0.40
Vampire---THF-4.5    subsumed     0.34
Vampire---TF1-4.5    subsumed     0.34
Vampire---4.5        subsumed     0.34
Vampire---SAT-4.5    subsumed     0.34
E---2.5              subsumed     0.34
E---FNT-2.5          subsumed     0.34
E---FNT-2.5          subsumed     0.34
iProver---SAT-3.0    subsumed     0.32
Metis---2.4          subsumed     0.32
SPASS---3.9          subsumed     0.28
Beagle---0.9.47      subsumed     0.26
Geo-III---2018C      subsumed     0.25
Z3---4.4.1           subsumed     0.25
Princess---170717    subsumed     0.25
Equinox---6.0.1a     subsumed     0.19
iProver---3.0        subsumed     0.11
iProver---SAT-3.0    subsumed     0.11
Darwin---1.4.5       subsumed     0.06
E-Darwin---1.5       subsumed     0.04
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use ET,2.0
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use CSE_E,1.0
CSE_E---1.0 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use Isabelle,2TH0
Isabelle---2016 can be used
Check if we can use Isabelle,2TF0
Isabelle---2016 can be used
Check if we can use Isabelle,2FOF
Isabelle---2016 can be used
Check if we can use Isabelle,2016
Isabelle---2016 can be used
Check if we can use Isabelle-HOT,2016
Isabelle-HOT---2016 can be used
Check if we can use CVC4,1.7
CVC4---SAT-1.7 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Leo-III,1.4
Leo-III---1.4 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Fampire,1.3
Check if we can use SInE,0.4
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use ePrincess,1.0
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Zipperpin,1.5
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Prover9,1109a
Prover9---1109a can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use LEO-II,1.7.0
LEO-II---1.7.0 can be used
Check if we can use nanoCoP,1.1
nanoCoP---1.1 can be used
Check if we can use Muscadet,4.5
Muscadet---4.5 can be used
Check if we can use Twee,2.2
Twee---2.2 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use iProverMo,2.5-0.1
iProverMo---2.5-0.1 can be used
Check if we can use Beagle,0.9.47
Check if we can use leanCoP,2.2
leanCoP---2.2 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Check if we can use CSE,1.0
CSE---1.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use Bliksem,1.12
Bliksem---1.12 can be used
Check if we can use Zenon,0.7.1
Zenon---0.7.1 can be used
Check if we can use ZenonModulo,0.4.2
Check if we can use SOS,2.0
SOS---2.0 can be used
Check if we can use Otter,3.3
Otter---3.3 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use Paradox,4.0
Paradox---4.0 can be used
Check if we can use Nitpick,2016
Nitpick---2016 can be used
Check if we can use DarwinFM,1.4.5
DarwinFM---1.4.5 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Beagle,0.9.47
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
E---FNT-2.5 will be used
E---FNT-2.5 will be used
Vampire---SAT-4.5 is given 499 seconds in -cE mode
Vampire---SAT-4.5 is given 250 seconds in -cE mode
Vampire---SAT-4.5 is given 125 seconds in -cE mode
Vampire---SAT-4.5 is given 62 seconds in -cE mode
E---FNT-2.5 is given 31 seconds in -cE mode
E---FNT-2.5 is given 32 seconds in -cE mode
Vampire---SAT-4.5 has started (PID 97257)
Vampire---SAT-4.5 has started (PID 97260)
Vampire---SAT-4.5 has started (PID 97264)
Vampire---SAT-4.5 has started (PID 97267)
E---FNT-2.5 has started (PID 97269)
E---FNT-2.5 has started (PID 97272)
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  499  /tmp/0CirYUgLs0/SOT_QlsZwL
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  250  /tmp/0CirYUgLs0/SOT_QlsZwL
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  125  /tmp/0CirYUgLs0/SOT_QlsZwL
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  62  /tmp/0CirYUgLs0/SOT_QlsZwL
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  31  /tmp/0CirYUgLs0/SOT_QlsZwL
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  32  /tmp/0CirYUgLs0/SOT_QlsZwL
% START OF SYSTEM OUTPUT
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_QlsZ
% START OF SYSTEM OUTPUT
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_QlsZ
% START OF SYSTEM OUTPUT
% START OF SYSTEM OUTPUT
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_QlsZ
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_QlsZ
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 15 seconds
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 15 seconds
% Time limit reached!
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 49252
% Time elapsed: 4.200 s
% ------------------------------
% ------------------------------
% fmb+10_1_av=off:fmbsr=1.6:lma=on:nm=64:nwc=3:sp=reverse_arity:urr=on_258 on SOT_QlsZ
TRYING [1]
TRYING [2]
TRYING [3]
TRYING [4]
TRYING [5]
TRYING [6]
% Time limit reached!
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 47589
% Time elapsed: 4.200 s
% ------------------------------
% ------------------------------
TRYING [7]
TRYING [8]
% fmb+10_1_av=off:fmbsr=1.6:lma=on:nm=64:nwc=3:sp=reverse_arity:urr=on_258 on SOT_QlsZ
TRYING [1]
TRYING [2]
TRYING [3]
TRYING [4]
TRYING [5]
TRYING [6]
TRYING [7]
TRYING [9]
TRYING [8]
TRYING [9]
TRYING [10]
% Time limit reached!
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 54753
% Time elapsed: 4.200 s
% ------------------------------
% ------------------------------
% Time limit reached!
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 49636
% Time elapsed: 4.200 s
% ------------------------------
% ------------------------------
TRYING [10]
% fmb+10_1_av=off:fmbsr=1.6:lma=on:nm=64:nwc=3:sp=reverse_arity:urr=on_258 on SOT_QlsZ
TRYING [1]
TRYING [2]
TRYING [3]
TRYING [4]
% fmb+10_1_av=off:fmbsr=1.6:lma=on:nm=64:nwc=3:sp=reverse_arity:urr=on_258 on SOT_QlsZ
TRYING [5]
TRYING [1]
TRYING [2]
TRYING [3]
TRYING [4]
TRYING [6]
TRYING [5]
TRYING [6]
TRYING [7]
TRYING [7]
TRYING [8]
TRYING [8]
TRYING [9]
TRYING [9]
TRYING [11]
TRYING [10]
TRYING [10]
TRYING [11]
TRYING [11]
TRYING [11]
TRYING [12]
TRYING [12]
TRYING [12]
TRYING [12]
TRYING [13]
TRYING [13]
% 10 wall clock seconds gone, waiting for system output
TRYING [13]
TRYING [13]
# AutoSched0-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S5PRR_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.017 s
# No success with AutoSched0
# Trying AutoSched1 for 7 seconds
# AutoSched0-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S5PRR_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.029 s
# No success with AutoSched0
# Trying AutoSched1 for 7 seconds
% 20 wall clock seconds gone, waiting for system output
# AutoSched1-Mode selected heuristic G_E___301_C18_F1_URBAN_S5PRR_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.014 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t5, conjecture, ![X7, X8, X9]:((partitions(X7,X9)&partitions(X8,X9))=>~(properSpecializes(X7,X8))), file('/tmp/0CirYUgLs0/SOT_QlsZwL', t5)).
fof(partitioningDefinition, axiom, ![X7, X8]:(partitions(X7,X8)<=>(completelyCategorizes(X7,X8)&disjointlyCategorizes(X7,X8))), file('/tmp/0CirYUgLs0/SOT_QlsZwL', partitioningDefinition)).
fof(a8_properSpecDef, axiom, ![X7, X8]:(properSpecializes(X7,X8)<=>(specializes(X7,X8)&~(X7=X8))), file('/tmp/0CirYUgLs0/SOT_QlsZwL', a8_properSpecDef)).
fof(disjointCategorizationDefinition, axiom, ![X7, X8]:(disjointlyCategorizes(X7,X8)<=>(categorizes(X7,X8)&![X4, X9, X10]:((((iof(X9,X7)&iof(X10,X7))&iof(X4,X9))&iof(X4,X10))=>X9=X10))), file('/tmp/0CirYUgLs0/SOT_QlsZwL', disjointCategorizationDefinition)).
fof(a7_specDef, axiom, ![X7, X8]:(specializes(X7,X8)<=>(type_(X7)&![X4]:(iof(X4,X7)=>iof(X4,X8)))), file('/tmp/0CirYUgLs0/SOT_QlsZwL', a7_specDef)).
fof(completeCategorizationDefinition, axiom, ![X7, X8]:(completelyCategorizes(X7,X8)<=>(categorizes(X7,X8)&![X4]:(iof(X4,X8)=>?[X9]:(iof(X4,X9)&iof(X9,X7))))), file('/tmp/0CirYUgLs0/SOT_QlsZwL', completeCategorizationDefinition)).
fof(a11_categorizationDef, axiom, ![X7, X8]:(categorizes(X7,X8)<=>(type_(X7)&![X9]:(iof(X9,X7)=>properSpecializes(X9,X8)))), file('/tmp/0CirYUgLs0/SOT_QlsZwL', a11_categorizationDef)).
fof(a9_equality, axiom, ![X7, X8]:((type_(X7)&type_(X8))=>(X7=X8<=>![X1]:(iof(X1,X7)<=>iof(X1,X8)))), file('/tmp/0CirYUgLs0/SOT_QlsZwL', a9_equality)).
fof(a3, axiom, ![X1]:(type_(X1)<=>?[X2]:iof(X2,X1)), file('/tmp/0CirYUgLs0/SOT_QlsZwL', a3)).
fof(c_0_9, negated_conjecture, ~(![X7, X8, X9]:((partitions(X7,X9)&partitions(X8,X9))=>~(properSpecializes(X7,X8)))), inference(assume_negation,[status(cth)],[t5])).
fof(c_0_10, plain, ![X101, X102]:(((completelyCategorizes(X101,X102)|~partitions(X101,X102))&(disjointlyCategorizes(X101,X102)|~partitions(X101,X102)))&(~completelyCategorizes(X101,X102)|~disjointlyCategorizes(X101,X102)|partitions(X101,X102))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[partitioningDefinition])])])).
fof(c_0_11, negated_conjecture, ((partitions(esk22_0,esk24_0)&partitions(esk23_0,esk24_0))&properSpecializes(esk22_0,esk23_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])).
fof(c_0_12, plain, ![X55, X56]:(((specializes(X55,X56)|~properSpecializes(X55,X56))&(X55!=X56|~properSpecializes(X55,X56)))&(~specializes(X55,X56)|X55=X56|properSpecializes(X55,X56))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[a8_properSpecDef])])])])).
fof(c_0_13, plain, ![X91, X92, X93, X94, X95, X96, X97]:(((categorizes(X91,X92)|~disjointlyCategorizes(X91,X92))&(~iof(X94,X91)|~iof(X95,X91)|~iof(X93,X94)|~iof(X93,X95)|X94=X95|~disjointlyCategorizes(X91,X92)))&(((((iof(esk20_2(X96,X97),X96)|~categorizes(X96,X97)|disjointlyCategorizes(X96,X97))&(iof(esk21_2(X96,X97),X96)|~categorizes(X96,X97)|disjointlyCategorizes(X96,X97)))&(iof(esk19_2(X96,X97),esk20_2(X96,X97))|~categorizes(X96,X97)|disjointlyCategorizes(X96,X97)))&(iof(esk19_2(X96,X97),esk21_2(X96,X97))|~categorizes(X96,X97)|disjointlyCategorizes(X96,X97)))&(esk20_2(X96,X97)!=esk21_2(X96,X97)|~categorizes(X96,X97)|disjointlyCategorizes(X96,X97)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[disjointCategorizationDefinition])])])])])])).
cnf(c_0_14, plain, (disjointlyCategorizes(X1,X2)|~partitions(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (partitions(esk23_0,esk24_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_16, plain, ![X49, X50, X51, X52, X53]:(((type_(X49)|~specializes(X49,X50))&(~iof(X51,X49)|iof(X51,X50)|~specializes(X49,X50)))&((iof(esk11_2(X52,X53),X52)|~type_(X52)|specializes(X52,X53))&(~iof(esk11_2(X52,X53),X53)|~type_(X52)|specializes(X52,X53)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a7_specDef])])])])])])).
cnf(c_0_17, plain, (specializes(X1,X2)|~properSpecializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, negated_conjecture, (properSpecializes(esk22_0,esk23_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, plain, (X1=X3|~iof(X1,X2)|~iof(X3,X2)|~iof(X4,X1)|~iof(X4,X3)|~disjointlyCategorizes(X2,X5)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, negated_conjecture, (disjointlyCategorizes(esk23_0,esk24_0)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_21, plain, (iof(X1,X3)|~iof(X1,X2)|~specializes(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, negated_conjecture, (specializes(esk22_0,esk23_0)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
fof(c_0_23, plain, ![X83, X84, X85, X87, X88, X90]:(((categorizes(X83,X84)|~completelyCategorizes(X83,X84))&((iof(X85,esk17_3(X83,X84,X85))|~iof(X85,X84)|~completelyCategorizes(X83,X84))&(iof(esk17_3(X83,X84,X85),X83)|~iof(X85,X84)|~completelyCategorizes(X83,X84))))&((iof(esk18_2(X87,X88),X88)|~categorizes(X87,X88)|completelyCategorizes(X87,X88))&(~iof(esk18_2(X87,X88),X90)|~iof(X90,X87)|~categorizes(X87,X88)|completelyCategorizes(X87,X88)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[completeCategorizationDefinition])])])])])])).
cnf(c_0_24, plain, (completelyCategorizes(X1,X2)|~partitions(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_25, negated_conjecture, (X1=X2|~iof(X2,esk23_0)|~iof(X1,esk23_0)|~iof(X3,X2)|~iof(X3,X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26, negated_conjecture, (iof(X1,esk23_0)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
fof(c_0_27, plain, ![X69, X70, X71, X72, X73]:(((type_(X69)|~categorizes(X69,X70))&(~iof(X71,X69)|properSpecializes(X71,X70)|~categorizes(X69,X70)))&((iof(esk14_2(X72,X73),X72)|~type_(X72)|categorizes(X72,X73))&(~properSpecializes(esk14_2(X72,X73),X73)|~type_(X72)|categorizes(X72,X73)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a11_categorizationDef])])])])])])).
cnf(c_0_28, plain, (categorizes(X1,X2)|~completelyCategorizes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, negated_conjecture, (completelyCategorizes(esk23_0,esk24_0)), inference(spm,[status(thm)],[c_0_24, c_0_15])).
cnf(c_0_30, negated_conjecture, (X1=X2|~iof(X1,esk23_0)|~iof(X2,esk22_0)|~iof(X3,X2)|~iof(X3,X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_31, plain, (iof(esk17_3(X1,X2,X3),X1)|~iof(X3,X2)|~completelyCategorizes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, plain, (properSpecializes(X1,X3)|~iof(X1,X2)|~categorizes(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, negated_conjecture, (categorizes(esk23_0,esk24_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, negated_conjecture, (esk17_3(esk23_0,X1,X2)=X3|~completelyCategorizes(esk23_0,X1)|~iof(X4,esk17_3(esk23_0,X1,X2))|~iof(X3,esk22_0)|~iof(X4,X3)|~iof(X2,X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_35, plain, (iof(X1,esk17_3(X2,X3,X1))|~iof(X1,X3)|~completelyCategorizes(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_36, negated_conjecture, (properSpecializes(X1,esk24_0)|~iof(X1,esk23_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_37, negated_conjecture, (esk17_3(esk23_0,X1,X2)=X3|~completelyCategorizes(esk23_0,X1)|~iof(X3,esk22_0)|~iof(X2,X3)|~iof(X2,X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_38, negated_conjecture, (specializes(X1,esk24_0)|~iof(X1,esk23_0)), inference(spm,[status(thm)],[c_0_17, c_0_36])).
fof(c_0_39, plain, ![X57, X58, X59, X60]:(((~iof(X59,X57)|iof(X59,X58)|X57!=X58|(~type_(X57)|~type_(X58)))&(~iof(X60,X58)|iof(X60,X57)|X57!=X58|(~type_(X57)|~type_(X58))))&((~iof(esk12_2(X57,X58),X57)|~iof(esk12_2(X57,X58),X58)|X57=X58|(~type_(X57)|~type_(X58)))&(iof(esk12_2(X57,X58),X57)|iof(esk12_2(X57,X58),X58)|X57=X58|(~type_(X57)|~type_(X58))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a9_equality])])])])])])).
cnf(c_0_40, plain, (type_(X1)|~categorizes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_41, negated_conjecture, (esk17_3(esk23_0,X1,X2)=esk17_3(esk22_0,X3,X4)|~completelyCategorizes(esk23_0,X1)|~completelyCategorizes(esk22_0,X3)|~iof(X2,esk17_3(esk22_0,X3,X4))|~iof(X2,X1)|~iof(X4,X3)), inference(spm,[status(thm)],[c_0_37, c_0_31])).
cnf(c_0_42, negated_conjecture, (iof(X1,esk24_0)|~iof(X2,esk23_0)|~iof(X1,X2)), inference(spm,[status(thm)],[c_0_21, c_0_38])).
cnf(c_0_43, plain, (iof(esk12_2(X1,X2),X1)|iof(esk12_2(X1,X2),X2)|X1=X2|~type_(X1)|~type_(X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_44, negated_conjecture, (type_(esk23_0)), inference(spm,[status(thm)],[c_0_40, c_0_33])).
cnf(c_0_45, plain, (type_(X1)|~specializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_46, plain, ![X16, X18, X19]:((~type_(X16)|iof(esk2_1(X16),X16))&(~iof(X19,X18)|type_(X18))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a3])])])])])).
cnf(c_0_47, negated_conjecture, (esk17_3(esk23_0,X1,X2)=esk17_3(esk22_0,X3,X2)|~completelyCategorizes(esk23_0,X1)|~completelyCategorizes(esk22_0,X3)|~iof(X2,X1)|~iof(X2,X3)), inference(spm,[status(thm)],[c_0_41, c_0_35])).
cnf(c_0_48, negated_conjecture, (partitions(esk22_0,esk24_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_49, negated_conjecture, (iof(X1,esk24_0)|~completelyCategorizes(esk23_0,X2)|~iof(X1,esk17_3(esk23_0,X2,X3))|~iof(X3,X2)), inference(spm,[status(thm)],[c_0_42, c_0_31])).
cnf(c_0_50, negated_conjecture, (X1=esk17_3(esk23_0,X2,X3)|~completelyCategorizes(esk23_0,X2)|~iof(X4,esk17_3(esk23_0,X2,X3))|~iof(X1,esk23_0)|~iof(X4,X1)|~iof(X3,X2)), inference(spm,[status(thm)],[c_0_25, c_0_31])).
cnf(c_0_51, negated_conjecture, (X1=esk23_0|iof(esk12_2(X1,esk23_0),esk23_0)|iof(esk12_2(X1,esk23_0),X1)|~type_(X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_52, negated_conjecture, (type_(esk22_0)), inference(spm,[status(thm)],[c_0_45, c_0_22])).
cnf(c_0_53, plain, (X1=X2|~iof(esk12_2(X1,X2),X1)|~iof(esk12_2(X1,X2),X2)|~type_(X1)|~type_(X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_54, plain, (type_(X2)|~iof(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_55, negated_conjecture, (iof(esk17_3(esk23_0,X1,X2),esk22_0)|~completelyCategorizes(esk22_0,X3)|~completelyCategorizes(esk23_0,X1)|~iof(X2,X3)|~iof(X2,X1)), inference(spm,[status(thm)],[c_0_31, c_0_47])).
cnf(c_0_56, negated_conjecture, (completelyCategorizes(esk22_0,esk24_0)), inference(spm,[status(thm)],[c_0_24, c_0_48])).
cnf(c_0_57, negated_conjecture, (iof(X1,esk24_0)|~completelyCategorizes(esk23_0,X2)|~iof(X1,X2)), inference(spm,[status(thm)],[c_0_49, c_0_35])).
cnf(c_0_58, negated_conjecture, (X1=esk17_3(esk23_0,X2,X3)|~completelyCategorizes(esk23_0,X2)|~iof(X1,esk23_0)|~iof(X3,X1)|~iof(X3,X2)), inference(spm,[status(thm)],[c_0_50, c_0_35])).
cnf(c_0_59, negated_conjecture, (esk23_0=esk22_0|iof(esk12_2(esk22_0,esk23_0),esk23_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_26])).
cnf(c_0_60, plain, (X1=X2|~iof(esk12_2(X1,X2),X2)|~iof(esk12_2(X1,X2),X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_53, c_0_54]), c_0_54])).
cnf(c_0_61, negated_conjecture, (iof(esk17_3(esk23_0,X1,X2),esk22_0)|~completelyCategorizes(esk23_0,X1)|~iof(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57])).
cnf(c_0_62, negated_conjecture, (esk17_3(esk23_0,X1,X2)=esk12_2(esk22_0,esk23_0)|esk23_0=esk22_0|~completelyCategorizes(esk23_0,X1)|~iof(X2,esk12_2(esk22_0,esk23_0))|~iof(X2,X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_63, negated_conjecture, (esk23_0=esk22_0|~iof(esk12_2(esk22_0,esk23_0),esk22_0)), inference(spm,[status(thm)],[c_0_60, c_0_59])).
cnf(c_0_64, negated_conjecture, (esk23_0=esk22_0|~completelyCategorizes(esk23_0,X1)|~iof(X2,esk12_2(esk22_0,esk23_0))|~iof(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])).
cnf(c_0_65, plain, (iof(esk2_1(X1),X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_66, negated_conjecture, (esk23_0=esk22_0|iof(X1,esk24_0)|~iof(X1,esk12_2(esk22_0,esk23_0))), inference(spm,[status(thm)],[c_0_42, c_0_59])).
cnf(c_0_67, negated_conjecture, (esk23_0=esk22_0|~completelyCategorizes(esk23_0,X1)|~type_(esk12_2(esk22_0,esk23_0))|~iof(esk2_1(esk12_2(esk22_0,esk23_0)),X1)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_68, negated_conjecture, (esk23_0=esk22_0|iof(esk2_1(esk12_2(esk22_0,esk23_0)),esk24_0)|~type_(esk12_2(esk22_0,esk23_0))), inference(spm,[status(thm)],[c_0_66, c_0_65])).
cnf(c_0_69, negated_conjecture, (esk23_0=esk22_0|~type_(esk12_2(esk22_0,esk23_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_29])])).
cnf(c_0_70, negated_conjecture, (type_(X1)|~iof(X1,esk23_0)), inference(spm,[status(thm)],[c_0_45, c_0_38])).
cnf(c_0_71, plain, (X1!=X2|~properSpecializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_72, negated_conjecture, (esk23_0=esk22_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_59])).
cnf(c_0_73, plain, (~properSpecializes(X1,X1)), inference(er,[status(thm)],[c_0_71])).
cnf(c_0_74, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_72]), c_0_73]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 75
# Proof object clause steps            : 56
# Proof object formula steps           : 19
# Proof object conjectures             : 41
# Proof object clause conjectures      : 38
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 34
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 77
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 72
# Processed clauses                    : 9487
# ...of these trivial                  : 2
# ...subsumed                          : 4016
# ...remaining for further processing  : 5469
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 968
# Backward-rewritten                   : 2771
# Generated clauses                    : 48278
# ...of the previous two non-trivial   : 48942
# Contextual simplify-reflections      : 597
# Paramodulations                      : 48270
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 1
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.051
#    Propositional solver time         : 0.065
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1722
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 1708
# Current number of unprocessed clauses: 36389
# ...number of literals in the above   : 241379
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3747
# Clause-clause subsumption calls (NU) : 3228256
# Rec. Clause-clause subsumption calls : 287753
# Non-unit clause-clause subsumptions  : 3720
# Unit Clause-clause subsumption calls : 6079
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1612990

# -------------------------------------------------
# User time                : 16.982 s
# System time              : 0.462 s
# Total time               : 17.444 s
# Maximum resident set size: 1608 pages

% END OF SYSTEM OUTPUT
% RESULT: SOT_QlsZwL - E---FNT-2.5 says Theorem - CPU = 17.39 WC = 17.52 
% OUTPUT: SOT_QlsZwL - E---FNT-2.5 says CNFRefutation - CPU = 17.39 WC = 17.52 

% END OF SYSTEM OUTPUT
% RESULT: SOT_QlsZwL - E---FNT-2.5 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_QlsZwL - E---FNT-2.5 says None - CPU = 0 WC = 0 
% 30 wall clock seconds gone, waiting for system output
% 40 wall clock seconds gone, waiting for system output
% 50 wall clock seconds gone, waiting for system output
% 60 wall clock seconds gone, waiting for system output

% END OF SYSTEM OUTPUT
% RESULT: SOT_QlsZwL - Vampire---SAT-4.5 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_QlsZwL - Vampire---SAT-4.5 says None - CPU = 0 WC = 0 
% 70 wall clock seconds gone, waiting for system output
% 80 wall clock seconds gone, waiting for system output
% 90 wall clock seconds gone, waiting for system output
% 100 wall clock seconds gone, waiting for system output
% 110 wall clock seconds gone, waiting for system output
% 120 wall clock seconds gone, waiting for system output

% END OF SYSTEM OUTPUT
% RESULT: SOT_QlsZwL - Vampire---SAT-4.5 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_QlsZwL - Vampire---SAT-4.5 says None - CPU = 0 WC = 0 
% 130 wall clock seconds gone, waiting for system output
% 140 wall clock seconds gone, waiting for system output
% 150 wall clock seconds gone, waiting for system output
% 160 wall clock seconds gone, waiting for system output
% 170 wall clock seconds gone, waiting for system output
% 180 wall clock seconds gone, waiting for system output
% 190 wall clock seconds gone, waiting for system output
% 200 wall clock seconds gone, waiting for system output
% 210 wall clock seconds gone, waiting for system output
% 220 wall clock seconds gone, waiting for system output
% 230 wall clock seconds gone, waiting for system output
% 240 wall clock seconds gone, waiting for system output
% 250 wall clock seconds gone, waiting for system output

% END OF SYSTEM OUTPUT
% RESULT: SOT_QlsZwL - Vampire---SAT-4.5 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_QlsZwL - Vampire---SAT-4.5 says None - CPU = 0 WC = 0 
% 260 wall clock seconds gone, waiting for system output
% 270 wall clock seconds gone, waiting for system output
% 280 wall clock seconds gone, waiting for system output
% 290 wall clock seconds gone, waiting for system output
% 300 wall clock seconds gone, waiting for system output
% 310 wall clock seconds gone, waiting for system output
% 320 wall clock seconds gone, waiting for system output
% 330 wall clock seconds gone, waiting for system output
% 340 wall clock seconds gone, waiting for system output
% 350 wall clock seconds gone, waiting for system output
% 360 wall clock seconds gone, waiting for system output
% 370 wall clock seconds gone, waiting for system output
% 380 wall clock seconds gone, waiting for system output
% 390 wall clock seconds gone, waiting for system output
% 400 wall clock seconds gone, waiting for system output
% 410 wall clock seconds gone, waiting for system output
% 420 wall clock seconds gone, waiting for system output
% 430 wall clock seconds gone, waiting for system output
% 440 wall clock seconds gone, waiting for system output
% 450 wall clock seconds gone, waiting for system output
% 460 wall clock seconds gone, waiting for system output
% 470 wall clock seconds gone, waiting for system output
% 480 wall clock seconds gone, waiting for system output
% 490 wall clock seconds gone, waiting for system output
% 500 wall clock seconds gone, waiting for system output

% END OF SYSTEM OUTPUT
% RESULT: SOT_QlsZwL - Vampire---SAT-4.5 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_QlsZwL - Vampire---SAT-4.5 says None - CPU = 0 WC = 0 
Vampire---SAT-4.5 has ended (PID 97257)
Vampire---SAT-4.5 has ended (PID 97260)
Vampire---SAT-4.5 has ended (PID 97264)
Vampire---SAT-4.5 has ended (PID 97267)
E---FNT-2.5 has ended (PID 97269)
E---FNT-2.5 has ended (PID 97272)
RESULT: SOT_QlsZwL - SSCPA---0.0 says Theorem - CPU = 970.5 WC = 501.6  SolvingSystem = E---FNT-2.5
</PRE>

</body>
</html>