<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>System Results</title>
<meta name="ROBOTS" content="INDEX,NOFOLLOW" />
<script type="text/javascript">//<![CDATA[

  var tstp_dump;
  function openSoTSTP (dump) {   
    var tstp_url = 'http://www.tptp.org/cgi-bin/SystemOnTSTP';
    var tstp_browser = window.open(tstp_url, '_blank');
    tstp_dump = dump;
  }
  function getTSTPDump () {
    return tstp_dump;
  }

//]]></script>
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>

% SZS start RequiredInformation<BR>
% Congratulations - you have become a registered power user of SystemOnTPTP, at IP address 191.6.38.118.<BR>
% Please consider donating to the TPTP project - see www.tptp.org for details.<BR>
% When you donate this message will disappear.<BR>
% If you do not donate a random delay might be added to your processing time.<BR>
% SZS end RequiredInformation<BR>
<PRE>
Getting system recommendations for generic problem /tmp/2fvw3M_yeU/SOT_hbsLD0
This is a FOF_THM_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.95
Vampire---TF1-4.5    recommended  0.95
Vampire---4.5        recommended  0.95
Vampire---SAT-4.5    recommended  0.95
ET---2.0             recommended  0.80
E---2.5              recommended  0.78
E---FNT-2.5          recommended  0.78
CSE_E---1.0          recommended  0.76
E---FNT-2.5          recommended  0.74
Isabelle---2TH0      recommended  0.69
Isabelle---2TF0      recommended  0.69
Isabelle---2FOF      recommended  0.69
Isabelle---2016      recommended  0.69
Isabelle-HOT---2016  recommended  0.68
CVC4---1.7           recommended  0.68
CVC4---SAT-1.7       recommended  0.68
Leo-III---1.4        recommended  0.61
Vampire---THF-4.5    recommended  0.59
Vampire---TF1-4.5    recommended  0.59
Vampire---SAT-4.5    recommended  0.59
iProver---3.0        recommended  0.58
iProver---SAT-3.0    recommended  0.58
Fampire---1.3        recommended  0.57
SInE---0.4           recommended  0.53
iProver---SAT-3.0    recommended  0.52
ePrincess---1.0      recommended  0.50
Princess---170717    recommended  0.48
SPASS---3.9          recommended  0.48
Z3---4.4.1           recommended  0.46
Zipperpin---1.5      recommended  0.40
CVC4---SAT-1.7       recommended  0.40
Prover9---1109a      recommended  0.39
Equinox---6.0.1a     recommended  0.37
LEO-II---1.7.0       recommended  0.35
nanoCoP---1.1        recommended  0.34
Muscadet---4.5       recommended  0.25
Twee---2.2           recommended  0.23
Metis---2.4          recommended  0.23
Geo-III---2018C      recommended  0.18
iProverMo---2.5-0.1  recommended  0.14
Beagle---0.9.47      subsumed     0.41
leanCoP---2.2        subsumed     0.36
E-Darwin---1.5       subsumed     0.31
CSE---1.0            subsumed     0.28
Darwin---1.4.5       subsumed     0.28
Bliksem---1.12       subsumed     0.21
Zenon---0.7.1        subsumed     0.16
ZenonModulo---0.4.2  subsumed     0.15
SOS---2.0            subsumed     0.11
Otter---3.3          subsumed     0.04
This is a FOF_CSA_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.98
Vampire---TF1-4.5    recommended  0.98
Vampire---SAT-4.5    recommended  0.98
Paradox---4.0        recommended  0.60
Nitpick---2016       subsumed     0.55
DarwinFM---1.4.5     subsumed     0.43
CVC4---SAT-1.7       subsumed     0.40
Vampire---THF-4.5    subsumed     0.34
Vampire---TF1-4.5    subsumed     0.34
Vampire---4.5        subsumed     0.34
Vampire---SAT-4.5    subsumed     0.34
E---2.5              subsumed     0.34
E---FNT-2.5          subsumed     0.34
E---FNT-2.5          subsumed     0.34
iProver---SAT-3.0    subsumed     0.32
Metis---2.4          subsumed     0.32
SPASS---3.9          subsumed     0.28
Beagle---0.9.47      subsumed     0.26
Geo-III---2018C      subsumed     0.25
Z3---4.4.1           subsumed     0.25
Princess---170717    subsumed     0.25
Equinox---6.0.1a     subsumed     0.19
iProver---3.0        subsumed     0.11
iProver---SAT-3.0    subsumed     0.11
Darwin---1.4.5       subsumed     0.06
E-Darwin---1.5       subsumed     0.04
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use ET,2.0
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use CSE_E,1.0
CSE_E---1.0 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use Isabelle,2TH0
Isabelle---2016 can be used
Check if we can use Isabelle,2TF0
Isabelle---2016 can be used
Check if we can use Isabelle,2FOF
Isabelle---2016 can be used
Check if we can use Isabelle,2016
Isabelle---2016 can be used
Check if we can use Isabelle-HOT,2016
Isabelle-HOT---2016 can be used
Check if we can use CVC4,1.7
CVC4---SAT-1.7 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Leo-III,1.4
Leo-III---1.4 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Fampire,1.3
Check if we can use SInE,0.4
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use ePrincess,1.0
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Zipperpin,1.5
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Prover9,1109a
Prover9---1109a can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use LEO-II,1.7.0
LEO-II---1.7.0 can be used
Check if we can use nanoCoP,1.1
nanoCoP---1.1 can be used
Check if we can use Muscadet,4.5
Muscadet---4.5 can be used
Check if we can use Twee,2.2
Twee---2.2 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use iProverMo,2.5-0.1
iProverMo---2.5-0.1 can be used
Check if we can use Beagle,0.9.47
Check if we can use leanCoP,2.2
leanCoP---2.2 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Check if we can use CSE,1.0
CSE---1.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use Bliksem,1.12
Bliksem---1.12 can be used
Check if we can use Zenon,0.7.1
Zenon---0.7.1 can be used
Check if we can use ZenonModulo,0.4.2
Check if we can use SOS,2.0
SOS---2.0 can be used
Check if we can use Otter,3.3
Otter---3.3 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use Paradox,4.0
Paradox---4.0 can be used
Check if we can use Nitpick,2016
Nitpick---2016 can be used
Check if we can use DarwinFM,1.4.5
DarwinFM---1.4.5 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Beagle,0.9.47
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
E---FNT-2.5 will be used
E---FNT-2.5 will be used
CSE_E---1.0 will be used
E---FNT-2.5 will be used
Isabelle---2016 will be used
Isabelle---2016 will be used
Vampire---SAT-4.5 is given 30 seconds in -cS mode
Vampire---SAT-4.5 is given 30 seconds in -cS mode
Vampire---SAT-4.5 is given 30 seconds in -cS mode
Vampire---SAT-4.5 is given 30 seconds in -cS mode
E---FNT-2.5 is given 30 seconds in -cS mode
E---FNT-2.5 is given 30 seconds in -cS mode
CSE_E---1.0 is given 30 seconds in -cS mode
E---FNT-2.5 is given 30 seconds in -cS mode
Isabelle---2016 is given 30 seconds in -cS mode
Isabelle---2016 is given 30 seconds in -cS mode
Vampire---SAT-4.5 has started (PID 102627)
Vampire---SAT-4.5 has started (PID 102628)
Vampire---SAT-4.5 has started (PID 102629)
Vampire---SAT-4.5 has started (PID 102630)
E---FNT-2.5 has started (PID 102631)
E---FNT-2.5 has started (PID 102632)
CSE_E---1.0 has started (PID 102633)
E---FNT-2.5 has started (PID 102634)
Isabelle---2016 has started (PID 102635)
Isabelle---2016 has started (PID 102636)
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  30  /tmp/2fvw3M_yeU/SOT_hbsLD0
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  30  /tmp/2fvw3M_yeU/SOT_hbsLD0
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  30  /tmp/2fvw3M_yeU/SOT_hbsLD0
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  30  /tmp/2fvw3M_yeU/SOT_hbsLD0
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  30  /tmp/2fvw3M_yeU/SOT_hbsLD0
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  30  /tmp/2fvw3M_yeU/SOT_hbsLD0
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 CSE_E---1.0  30  /tmp/2fvw3M_yeU/SOT_hbsLD0
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  30  /tmp/2fvw3M_yeU/SOT_hbsLD0
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Isabelle---2016  30  /tmp/2fvw3M_yeU/SOT_hbsLD0
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Isabelle---2016  30  /tmp/2fvw3M_yeU/SOT_hbsLD0
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 14 seconds
% START OF SYSTEM OUTPUT
% START OF SYSTEM OUTPUT
% START OF SYSTEM OUTPUT
# AutoSched0-Mode selected heuristic G_E___208_C02CMA_F1_SE_CS_SP_PS_S5PRR_RG_S04AN
# and selection function SelectComplexExceptUniqMaxHorn.
#
# Preprocessing time       : 0.016 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t9_ap3, conjecture, ~(?[X11]:(iof(X11,c_sot)&specializes(X11,c_sot))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', t9_ap3)).
fof(a2, axiom, ![X1]:(individual(X1)<=>~(?[X2]:iof(X2,X1))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a2)).
fof(a3, axiom, ![X1]:(type_(X1)<=>?[X2]:iof(X2,X1)), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a3)).
fof(a7_specDef, axiom, ![X7, X8]:(specializes(X7,X8)<=>(type_(X7)&![X4]:(iof(X4,X7)=>iof(X4,X8)))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a7_specDef)).
fof(a15_sotConstantDef, axiom, ![X11]:(X11=c_sot<=>![X1]:(iof(X1,X11)<=>secondordertype(X1))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a15_sotConstantDef)).
fof(a6_secondOrderTypeDef, axiom, ![X1]:(secondordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>firstordertype(X2)))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a6_secondOrderTypeDef)).
fof(a5_firstOrderTypeDef, axiom, ![X1]:(firstordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>individual(X2)))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a5_firstOrderTypeDef)).
fof(c_0_7, negated_conjecture, ~(~(?[X11]:(iof(X11,c_sot)&specializes(X11,c_sot)))), inference(assume_negation,[status(cth)],[t9_ap3])).
fof(c_0_8, plain, ![X13, X14, X15]:((~individual(X13)|~iof(X14,X13))&(iof(esk1_1(X15),X15)|individual(X15))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a2])])])])])).
fof(c_0_9, plain, ![X17, X19, X20]:((~type_(X17)|iof(esk2_1(X17),X17))&(~iof(X20,X19)|type_(X19))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a3])])])])])).
fof(c_0_10, plain, ![X50, X51, X52, X53, X54]:(((type_(X50)|~specializes(X50,X51))&(~iof(X52,X50)|iof(X52,X51)|~specializes(X50,X51)))&((iof(esk11_2(X53,X54),X53)|~type_(X53)|specializes(X53,X54))&(~iof(esk11_2(X53,X54),X54)|~type_(X53)|specializes(X53,X54)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a7_specDef])])])])])])).
fof(c_0_11, negated_conjecture, (iof(esk25_0,c_sot)&specializes(esk25_0,c_sot)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
cnf(c_0_12, plain, (~individual(X1)|~iof(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (iof(esk2_1(X1),X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (type_(X1)|~specializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (specializes(esk25_0,c_sot)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_16, plain, ![X114, X115, X116, X117]:(((~iof(X115,X114)|secondordertype(X115)|X114!=c_sot)&(~secondordertype(X116)|iof(X116,X114)|X114!=c_sot))&((~iof(esk24_1(X117),X117)|~secondordertype(esk24_1(X117))|X117=c_sot)&(iof(esk24_1(X117),X117)|secondordertype(esk24_1(X117))|X117=c_sot))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a15_sotConstantDef])])])])])])).
cnf(c_0_17, plain, (iof(X1,X3)|~iof(X1,X2)|~specializes(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, plain, (~type_(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_19, negated_conjecture, (type_(esk25_0)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
fof(c_0_20, plain, ![X46, X47, X48]:(((type_(X46)|~secondordertype(X46))&(~iof(X47,X46)|firstordertype(X47)|~secondordertype(X46)))&((iof(esk10_1(X48),X48)|~type_(X48)|secondordertype(X48))&(~firstordertype(esk10_1(X48))|~type_(X48)|secondordertype(X48)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a6_secondOrderTypeDef])])])])])])).
fof(c_0_21, plain, ![X42, X43, X44]:(((type_(X42)|~firstordertype(X42))&(~iof(X43,X42)|individual(X43)|~firstordertype(X42)))&((iof(esk9_1(X44),X44)|~type_(X44)|firstordertype(X44))&(~individual(esk9_1(X44))|~type_(X44)|firstordertype(X44)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a5_firstOrderTypeDef])])])])])])).
cnf(c_0_22, plain, (secondordertype(X1)|~iof(X1,X2)|X2!=c_sot), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, negated_conjecture, (iof(X1,c_sot)|~iof(X1,esk25_0)), inference(spm,[status(thm)],[c_0_17, c_0_15])).
cnf(c_0_24, plain, (iof(esk1_1(X1),X1)|individual(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_25, negated_conjecture, (~individual(esk25_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_26, plain, (type_(X1)|~secondordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (iof(esk9_1(X1),X1)|firstordertype(X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (secondordertype(X1)|~iof(X1,c_sot)), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_29, negated_conjecture, (iof(esk1_1(esk25_0),c_sot)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
cnf(c_0_30, plain, (type_(X1)|~firstordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (firstordertype(X1)|~iof(X1,X2)|~secondordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_32, plain, (~secondordertype(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_18, c_0_26])).
cnf(c_0_33, negated_conjecture, (firstordertype(esk25_0)|iof(esk9_1(esk25_0),esk25_0)), inference(spm,[status(thm)],[c_0_27, c_0_19])).
cnf(c_0_34, negated_conjecture, (secondordertype(esk1_1(esk25_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, plain, (individual(X1)|~iof(X1,X2)|~firstordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_36, plain, (~firstordertype(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_18, c_0_30])).
cnf(c_0_37, negated_conjecture, (iof(esk25_0,c_sot)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_38, plain, (firstordertype(esk1_1(X1))|~secondordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_24]), c_0_32])).
cnf(c_0_39, negated_conjecture, (firstordertype(esk25_0)|iof(esk9_1(esk25_0),c_sot)), inference(spm,[status(thm)],[c_0_23, c_0_33])).
cnf(c_0_40, negated_conjecture, (~individual(esk1_1(esk25_0))), inference(spm,[status(thm)],[c_0_32, c_0_34])).
cnf(c_0_41, plain, (individual(esk1_1(X1))|~firstordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_24]), c_0_36])).
cnf(c_0_42, negated_conjecture, (secondordertype(esk25_0)), inference(spm,[status(thm)],[c_0_28, c_0_37])).
cnf(c_0_43, plain, (~secondordertype(X1)|~individual(esk1_1(X1))), inference(spm,[status(thm)],[c_0_36, c_0_38])).
cnf(c_0_44, negated_conjecture, (secondordertype(esk9_1(esk25_0))|firstordertype(esk25_0)), inference(spm,[status(thm)],[c_0_28, c_0_39])).
cnf(c_0_45, negated_conjecture, (~firstordertype(esk25_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_46, negated_conjecture, (firstordertype(esk9_1(esk25_0))|firstordertype(esk25_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_33]), c_0_42])])).
cnf(c_0_47, plain, (~secondordertype(X1)|~firstordertype(X1)), inference(spm,[status(thm)],[c_0_43, c_0_41])).
cnf(c_0_48, negated_conjecture, (secondordertype(esk9_1(esk25_0))), inference(sr,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_49, negated_conjecture, (firstordertype(esk9_1(esk25_0))), inference(sr,[status(thm)],[c_0_46, c_0_45])).
cnf(c_0_50, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 51
# Proof object clause steps            : 36
# Proof object formula steps           : 15
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 10
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 24
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 88
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 83
# Processed clauses                    : 561
# ...of these trivial                  : 2
# ...subsumed                          : 146
# ...remaining for further processing  : 413
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 14
# Backward-rewritten                   : 8
# Generated clauses                    : 854
# ...of the previous two non-trivial   : 770
# Contextual simplify-reflections      : 27
# Paramodulations                      : 836
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 290
#    Positive orientable unit clauses  : 29
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 18
#    Non-unit-clauses                  : 243
# Current number of unprocessed clauses: 344
# ...number of literals in the above   : 1081
# Current number of archived formulas  : 0
# Current number of archived clauses   : 117
# Clause-clause subsumption calls (NU) : 13501
# Rec. Clause-clause subsumption calls : 10424
# Non-unit clause-clause subsumptions  : 68
# Unit Clause-clause subsumption calls : 1095
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 15308

# -------------------------------------------------
# User time                : 0.038 s
# System time              : 0.006 s
# Total time               : 0.045 s
# Maximum resident set size: 1612 pages
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 14 seconds
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_hbsL
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_hbsL
# AutoSched0-Mode selected heuristic G_E___208_C02CMA_F1_SE_CS_SP_PS_S5PRR_RG_S04AN
# and selection function SelectComplexExceptUniqMaxHorn.
#
# Preprocessing time       : 0.016 s
# Presaturation interreduction done

# Proof found!
% START OF SYSTEM OUTPUT
# SZS status Theorem
# SZS output start CNFRefutation
fof(t9_ap3, conjecture, ~(?[X11]:(iof(X11,c_sot)&specializes(X11,c_sot))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', t9_ap3)).
fof(a2, axiom, ![X1]:(individual(X1)<=>~(?[X2]:iof(X2,X1))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a2)).
fof(a3, axiom, ![X1]:(type_(X1)<=>?[X2]:iof(X2,X1)), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a3)).
% START OF SYSTEM OUTPUT
fof(a7_specDef, axiom, ![X7, X8]:(specializes(X7,X8)<=>(type_(X7)&![X4]:(iof(X4,X7)=>iof(X4,X8)))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a7_specDef)).
fof(a15_sotConstantDef, axiom, ![X11]:(X11=c_sot<=>![X1]:(iof(X1,X11)<=>secondordertype(X1))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a15_sotConstantDef)).
fof(a6_secondOrderTypeDef, axiom, ![X1]:(secondordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>firstordertype(X2)))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a6_secondOrderTypeDef)).
fof(a5_firstOrderTypeDef, axiom, ![X1]:(firstordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>individual(X2)))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a5_firstOrderTypeDef)).
fof(c_0_7, negated_conjecture, ~(~(?[X11]:(iof(X11,c_sot)&specializes(X11,c_sot)))), inference(assume_negation,[status(cth)],[t9_ap3])).
fof(c_0_8, plain, ![X13, X14, X15]:((~individual(X13)|~iof(X14,X13))&(iof(esk1_1(X15),X15)|individual(X15))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a2])])])])])).
fof(c_0_9, plain, ![X17, X19, X20]:((~type_(X17)|iof(esk2_1(X17),X17))&(~iof(X20,X19)|type_(X19))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a3])])])])])).
fof(c_0_10, plain, ![X50, X51, X52, X53, X54]:(((type_(X50)|~specializes(X50,X51))&(~iof(X52,X50)|iof(X52,X51)|~specializes(X50,X51)))&((iof(esk11_2(X53,X54),X53)|~type_(X53)|specializes(X53,X54))&(~iof(esk11_2(X53,X54),X54)|~type_(X53)|specializes(X53,X54)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a7_specDef])])])])])])).
fof(c_0_11, negated_conjecture, (iof(esk25_0,c_sot)&specializes(esk25_0,c_sot)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
cnf(c_0_12, plain, (~individual(X1)|~iof(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (iof(esk2_1(X1),X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (type_(X1)|~specializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (specializes(esk25_0,c_sot)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_16, plain, ![X114, X115, X116, X117]:(((~iof(X115,X114)|secondordertype(X115)|X114!=c_sot)&(~secondordertype(X116)|iof(X116,X114)|X114!=c_sot))&((~iof(esk24_1(X117),X117)|~secondordertype(esk24_1(X117))|X117=c_sot)&(iof(esk24_1(X117),X117)|secondordertype(esk24_1(X117))|X117=c_sot))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a15_sotConstantDef])])])])])])).
cnf(c_0_17, plain, (iof(X1,X3)|~iof(X1,X2)|~specializes(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, plain, (~type_(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_19, negated_conjecture, (type_(esk25_0)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
fof(c_0_20, plain, ![X46, X47, X48]:(((type_(X46)|~secondordertype(X46))&(~iof(X47,X46)|firstordertype(X47)|~secondordertype(X46)))&((iof(esk10_1(X48),X48)|~type_(X48)|secondordertype(X48))&(~firstordertype(esk10_1(X48))|~type_(X48)|secondordertype(X48)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a6_secondOrderTypeDef])])])])])])).
fof(c_0_21, plain, ![X42, X43, X44]:(((type_(X42)|~firstordertype(X42))&(~iof(X43,X42)|individual(X43)|~firstordertype(X42)))&((iof(esk9_1(X44),X44)|~type_(X44)|firstordertype(X44))&(~individual(esk9_1(X44))|~type_(X44)|firstordertype(X44)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a5_firstOrderTypeDef])])])])])])).
cnf(c_0_22, plain, (secondordertype(X1)|~iof(X1,X2)|X2!=c_sot), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, negated_conjecture, (iof(X1,c_sot)|~iof(X1,esk25_0)), inference(spm,[status(thm)],[c_0_17, c_0_15])).
cnf(c_0_24, plain, (iof(esk1_1(X1),X1)|individual(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_25, negated_conjecture, (~individual(esk25_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_26, plain, (type_(X1)|~secondordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (iof(esk9_1(X1),X1)|firstordertype(X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (secondordertype(X1)|~iof(X1,c_sot)), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_29, negated_conjecture, (iof(esk1_1(esk25_0),c_sot)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
cnf(c_0_30, plain, (type_(X1)|~firstordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (firstordertype(X1)|~iof(X1,X2)|~secondordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_32, plain, (~secondordertype(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_18, c_0_26])).
cnf(c_0_33, negated_conjecture, (firstordertype(esk25_0)|iof(esk9_1(esk25_0),esk25_0)), inference(spm,[status(thm)],[c_0_27, c_0_19])).
cnf(c_0_34, negated_conjecture, (secondordertype(esk1_1(esk25_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, plain, (individual(X1)|~iof(X1,X2)|~firstordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_36, plain, (~firstordertype(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_18, c_0_30])).
cnf(c_0_37, negated_conjecture, (iof(esk25_0,c_sot)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_38, plain, (firstordertype(esk1_1(X1))|~secondordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_24]), c_0_32])).
cnf(c_0_39, negated_conjecture, (firstordertype(esk25_0)|iof(esk9_1(esk25_0),c_sot)), inference(spm,[status(thm)],[c_0_23, c_0_33])).
cnf(c_0_40, negated_conjecture, (~individual(esk1_1(esk25_0))), inference(spm,[status(thm)],[c_0_32, c_0_34])).
cnf(c_0_41, plain, (individual(esk1_1(X1))|~firstordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_24]), c_0_36])).
cnf(c_0_42, negated_conjecture, (secondordertype(esk25_0)), inference(spm,[status(thm)],[c_0_28, c_0_37])).
cnf(c_0_43, plain, (~secondordertype(X1)|~individual(esk1_1(X1))), inference(spm,[status(thm)],[c_0_36, c_0_38])).
cnf(c_0_44, negated_conjecture, (secondordertype(esk9_1(esk25_0))|firstordertype(esk25_0)), inference(spm,[status(thm)],[c_0_28, c_0_39])).
cnf(c_0_45, negated_conjecture, (~firstordertype(esk25_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_46, negated_conjecture, (firstordertype(esk9_1(esk25_0))|firstordertype(esk25_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_33]), c_0_42])])).
cnf(c_0_47, plain, (~secondordertype(X1)|~firstordertype(X1)), inference(spm,[status(thm)],[c_0_43, c_0_41])).
cnf(c_0_48, negated_conjecture, (secondordertype(esk9_1(esk25_0))), inference(sr,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_49, negated_conjecture, (firstordertype(esk9_1(esk25_0))), inference(sr,[status(thm)],[c_0_46, c_0_45])).
cnf(c_0_50, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 51
# Proof object clause steps            : 36
# Proof object formula steps           : 15
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 10
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 24
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 88
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 83
# Processed clauses                    : 561
# ...of these trivial                  : 2
# ...subsumed                          : 146
# ...remaining for further processing  : 413
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 14
# Backward-rewritten                   : 8
# Generated clauses                    : 854
# ...of the previous two non-trivial   : 770
# Contextual simplify-reflections      : 27
# Paramodulations                      : 836
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 290
#    Positive orientable unit clauses  : 29
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 18
#    Non-unit-clauses                  : 243
# Current number of unprocessed clauses: 344
# ...number of literals in the above   : 1081
# Current number of archived formulas  : 0
# Current number of archived clauses   : 117
# Clause-clause subsumption calls (NU) : 13501
# Rec. Clause-clause subsumption calls : 10424
# Non-unit clause-clause subsumptions  : 68
# Unit Clause-clause subsumption calls : 1095
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 15308

# -------------------------------------------------
# User time                : 0.035 s
# System time              : 0.007 s
# Total time               : 0.042 s
# Maximum resident set size: 1612 pages
% START OF SYSTEM OUTPUT
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_hbsL
% START OF SYSTEM OUTPUT
Moving to /home/tptp/Systems/CSE_E---1.0
java -jar /home/tptp/Systems/CSE_E---1.0/mcs_epr.jar 30 /tmp/2fvw3M_yeU/SOT_hbsLD0
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_hbsL
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 14 seconds
# AutoSched0-Mode selected heuristic G_E___208_C02CMA_F1_SE_CS_SP_PS_S5PRR_RG_S04AN
# and selection function SelectComplexExceptUniqMaxHorn.
#
# Preprocessing time       : 0.016 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t9_ap3, conjecture, ~(?[X11]:(iof(X11,c_sot)&specializes(X11,c_sot))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', t9_ap3)).
fof(a2, axiom, ![X1]:(individual(X1)<=>~(?[X2]:iof(X2,X1))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a2)).
fof(a3, axiom, ![X1]:(type_(X1)<=>?[X2]:iof(X2,X1)), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a3)).
fof(a7_specDef, axiom, ![X7, X8]:(specializes(X7,X8)<=>(type_(X7)&![X4]:(iof(X4,X7)=>iof(X4,X8)))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a7_specDef)).
fof(a15_sotConstantDef, axiom, ![X11]:(X11=c_sot<=>![X1]:(iof(X1,X11)<=>secondordertype(X1))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a15_sotConstantDef)).
fof(a6_secondOrderTypeDef, axiom, ![X1]:(secondordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>firstordertype(X2)))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a6_secondOrderTypeDef)).
fof(a5_firstOrderTypeDef, axiom, ![X1]:(firstordertype(X1)<=>(type_(X1)&![X2]:(iof(X2,X1)=>individual(X2)))), file('/tmp/2fvw3M_yeU/SOT_hbsLD0', a5_firstOrderTypeDef)).
fof(c_0_7, negated_conjecture, ~(~(?[X11]:(iof(X11,c_sot)&specializes(X11,c_sot)))), inference(assume_negation,[status(cth)],[t9_ap3])).
fof(c_0_8, plain, ![X13, X14, X15]:((~individual(X13)|~iof(X14,X13))&(iof(esk1_1(X15),X15)|individual(X15))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a2])])])])])).
fof(c_0_9, plain, ![X17, X19, X20]:((~type_(X17)|iof(esk2_1(X17),X17))&(~iof(X20,X19)|type_(X19))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a3])])])])])).
fof(c_0_10, plain, ![X50, X51, X52, X53, X54]:(((type_(X50)|~specializes(X50,X51))&(~iof(X52,X50)|iof(X52,X51)|~specializes(X50,X51)))&((iof(esk11_2(X53,X54),X53)|~type_(X53)|specializes(X53,X54))&(~iof(esk11_2(X53,X54),X54)|~type_(X53)|specializes(X53,X54)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a7_specDef])])])])])])).
fof(c_0_11, negated_conjecture, (iof(esk25_0,c_sot)&specializes(esk25_0,c_sot)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
cnf(c_0_12, plain, (~individual(X1)|~iof(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (iof(esk2_1(X1),X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (type_(X1)|~specializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (specializes(esk25_0,c_sot)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_16, plain, ![X114, X115, X116, X117]:(((~iof(X115,X114)|secondordertype(X115)|X114!=c_sot)&(~secondordertype(X116)|iof(X116,X114)|X114!=c_sot))&((~iof(esk24_1(X117),X117)|~secondordertype(esk24_1(X117))|X117=c_sot)&(iof(esk24_1(X117),X117)|secondordertype(esk24_1(X117))|X117=c_sot))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a15_sotConstantDef])])])])])])).
cnf(c_0_17, plain, (iof(X1,X3)|~iof(X1,X2)|~specializes(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, plain, (~type_(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_19, negated_conjecture, (type_(esk25_0)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
fof(c_0_20, plain, ![X46, X47, X48]:(((type_(X46)|~secondordertype(X46))&(~iof(X47,X46)|firstordertype(X47)|~secondordertype(X46)))&((iof(esk10_1(X48),X48)|~type_(X48)|secondordertype(X48))&(~firstordertype(esk10_1(X48))|~type_(X48)|secondordertype(X48)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a6_secondOrderTypeDef])])])])])])).
fof(c_0_21, plain, ![X42, X43, X44]:(((type_(X42)|~firstordertype(X42))&(~iof(X43,X42)|individual(X43)|~firstordertype(X42)))&((iof(esk9_1(X44),X44)|~type_(X44)|firstordertype(X44))&(~individual(esk9_1(X44))|~type_(X44)|firstordertype(X44)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a5_firstOrderTypeDef])])])])])])).
cnf(c_0_22, plain, (secondordertype(X1)|~iof(X1,X2)|X2!=c_sot), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, negated_conjecture, (iof(X1,c_sot)|~iof(X1,esk25_0)), inference(spm,[status(thm)],[c_0_17, c_0_15])).
cnf(c_0_24, plain, (iof(esk1_1(X1),X1)|individual(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_25, negated_conjecture, (~individual(esk25_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_26, plain, (type_(X1)|~secondordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (iof(esk9_1(X1),X1)|firstordertype(X1)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (secondordertype(X1)|~iof(X1,c_sot)), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_29, negated_conjecture, (iof(esk1_1(esk25_0),c_sot)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
cnf(c_0_30, plain, (type_(X1)|~firstordertype(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (firstordertype(X1)|~iof(X1,X2)|~secondordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_32, plain, (~secondordertype(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_18, c_0_26])).
cnf(c_0_33, negated_conjecture, (firstordertype(esk25_0)|iof(esk9_1(esk25_0),esk25_0)), inference(spm,[status(thm)],[c_0_27, c_0_19])).
cnf(c_0_34, negated_conjecture, (secondordertype(esk1_1(esk25_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, plain, (individual(X1)|~iof(X1,X2)|~firstordertype(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_36, plain, (~firstordertype(X1)|~individual(X1)), inference(spm,[status(thm)],[c_0_18, c_0_30])).
cnf(c_0_37, negated_conjecture, (iof(esk25_0,c_sot)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_38, plain, (firstordertype(esk1_1(X1))|~secondordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_24]), c_0_32])).
cnf(c_0_39, negated_conjecture, (firstordertype(esk25_0)|iof(esk9_1(esk25_0),c_sot)), inference(spm,[status(thm)],[c_0_23, c_0_33])).
cnf(c_0_40, negated_conjecture, (~individual(esk1_1(esk25_0))), inference(spm,[status(thm)],[c_0_32, c_0_34])).
cnf(c_0_41, plain, (individual(esk1_1(X1))|~firstordertype(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_24]), c_0_36])).
cnf(c_0_42, negated_conjecture, (secondordertype(esk25_0)), inference(spm,[status(thm)],[c_0_28, c_0_37])).
cnf(c_0_43, plain, (~secondordertype(X1)|~individual(esk1_1(X1))), inference(spm,[status(thm)],[c_0_36, c_0_38])).
cnf(c_0_44, negated_conjecture, (secondordertype(esk9_1(esk25_0))|firstordertype(esk25_0)), inference(spm,[status(thm)],[c_0_28, c_0_39])).
cnf(c_0_45, negated_conjecture, (~firstordertype(esk25_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_46, negated_conjecture, (firstordertype(esk9_1(esk25_0))|firstordertype(esk25_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_33]), c_0_42])])).
cnf(c_0_47, plain, (~secondordertype(X1)|~firstordertype(X1)), inference(spm,[status(thm)],[c_0_43, c_0_41])).
cnf(c_0_48, negated_conjecture, (secondordertype(esk9_1(esk25_0))), inference(sr,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_49, negated_conjecture, (firstordertype(esk9_1(esk25_0))), inference(sr,[status(thm)],[c_0_46, c_0_45])).
cnf(c_0_50, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 51
# Proof object clause steps            : 36
# Proof object formula steps           : 15
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 10
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 24
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 88
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 83
# Processed clauses                    : 561
# ...of these trivial                  : 2
# ...subsumed                          : 146
# ...remaining for further processing  : 413
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 14
# Backward-rewritten                   : 8
# Generated clauses                    : 854
# ...of the previous two non-trivial   : 770
# Contextual simplify-reflections      : 27
# Paramodulations                      : 836
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 290
#    Positive orientable unit clauses  : 29
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 18
#    Non-unit-clauses                  : 243
# Current number of unprocessed clauses: 344
# ...number of literals in the above   : 1081
# Current number of archived formulas  : 0
# Current number of archived clauses   : 117
# Clause-clause subsumption calls (NU) : 13501
# Rec. Clause-clause subsumption calls : 10424
# Non-unit clause-clause subsumptions  : 68
# Unit Clause-clause subsumption calls : 1095
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 15308

# -------------------------------------------------
# User time                : 0.039 s
# System time              : 0.006 s
# Total time               : 0.045 s
# Maximum resident set size: 1612 pages
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_hbsL
% SZS output start Proof for SOT_hbsL
fof(f9378,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f129,f133,f145,f153,f169,f189,f201,f205,f226,f244,f336,f348,f396,f626,f728,f795,f802,f818,f917,f1239,f1249,f4014,f4093,f4537,f9353,f9360,f9368,f9377])).
fof(f9377,plain,(
  ~spl15_7 | ~spl15_20 | ~spl15_368 | ~spl15_369),
  inference(avatar_contradiction_clause,[],[f9376])).
fof(f9376,plain,(
  $false | (~spl15_7 | ~spl15_20 | ~spl15_368 | ~spl15_369)),
  inference(subsumption_resolution,[],[f9372,f9364])).
fof(f9364,plain,(
  ( ! [X0] : (~iof(X0,sK8(sK7(sK0)))) ) | (~spl15_7 | ~spl15_368)),
  inference(resolution,[],[f9359,f152])).
fof(f152,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) ) | ~spl15_7),
  inference(avatar_component_clause,[],[f151])).
fof(f151,plain,(
  spl15_7 <=> ! [X1,X0] : (~iof(X1,X0) | ~individual(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_7])])).
fof(f9359,plain,(
  individual(sK8(sK7(sK0))) | ~spl15_368),
  inference(avatar_component_clause,[],[f9358])).
fof(f9358,plain,(
  spl15_368 <=> individual(sK8(sK7(sK0)))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_368])])).
fof(f9372,plain,(
  iof(sK8(sK8(sK7(sK0))),sK8(sK7(sK0))) | (~spl15_20 | ~spl15_369)),
  inference(resolution,[],[f9367,f204])).
fof(f204,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK8(X0),X0)) ) | ~spl15_20),
  inference(avatar_component_clause,[],[f203])).
fof(f203,plain,(
  spl15_20 <=> ! [X0] : (iof(sK8(X0),X0) | ~type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_20])])).
fof(f9367,plain,(
  type_(sK8(sK7(sK0))) | ~spl15_369),
  inference(avatar_component_clause,[],[f9366])).
fof(f9366,plain,(
  spl15_369 <=> type_(sK8(sK7(sK0)))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_369])])).
fof(f9368,plain,(
  spl15_369 | ~spl15_5 | ~spl15_367),
  inference(avatar_split_clause,[],[f9356,f9351,f143,f9366])).
fof(f143,plain,(
  spl15_5 <=> ! [X0] : (type_(X0) | ~firstordertype(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_5])])).
fof(f9351,plain,(
  spl15_367 <=> firstordertype(sK8(sK7(sK0)))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_367])])).
fof(f9356,plain,(
  type_(sK8(sK7(sK0))) | (~spl15_5 | ~spl15_367)),
  inference(resolution,[],[f9352,f144])).
% Refutation found. Thanks to Tanya!
fof(f144,plain,(
  ( ! [X0] : (~firstordertype(X0) | type_(X0)) ) | ~spl15_5),
  inference(avatar_component_clause,[],[f143])).
fof(f9352,plain,(
  firstordertype(sK8(sK7(sK0))) | ~spl15_367),
% SZS status Theorem for SOT_hbsL
  inference(avatar_component_clause,[],[f9351])).
fof(f9360,plain,(
  spl15_368 | ~spl15_298 | ~spl15_344),
% SZS output start Proof for SOT_hbsL
  inference(avatar_split_clause,[],[f9345,f4535,f4012,f9358])).
fof(f4012,plain,(
  spl15_298 <=> ! [X0] : (individual(X0) | ~iof(X0,sK7(sK0)))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_298])])).
fof(f4535,plain,(
  spl15_344 <=> iof(sK8(sK7(sK0)),sK7(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_344])])).
fof(f9345,plain,(
  individual(sK8(sK7(sK0))) | (~spl15_298 | ~spl15_344)),
  inference(resolution,[],[f4536,f4013])).
fof(f4013,plain,(
  ( ! [X0] : (~iof(X0,sK7(sK0)) | individual(X0)) ) | ~spl15_298),
  inference(avatar_component_clause,[],[f4012])).
fof(f4536,plain,(
  iof(sK8(sK7(sK0)),sK7(sK0)) | ~spl15_344),
  inference(avatar_component_clause,[],[f4535])).
fof(f9353,plain,(
  spl15_367 | ~spl15_303 | ~spl15_344),
  inference(avatar_split_clause,[],[f9344,f4535,f4091,f9351])).
fof(f9378,plain,(
fof(f4091,plain,(
  spl15_303 <=> ! [X0] : (firstordertype(X0) | ~iof(X0,sK7(sK0)))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_303])])).
  $false),
fof(f9344,plain,(
  firstordertype(sK8(sK7(sK0))) | (~spl15_303 | ~spl15_344)),
  inference(resolution,[],[f4536,f4092])).
fof(f4092,plain,(
  inference(avatar_sat_refutation,[],[f129,f133,f145,f153,f169,f189,f201,f205,f226,f244,f336,f348,f396,f626,f728,f795,f802,f818,f917,f1239,f1249,f4014,f4093,f4537,f9353,f9360,f9368,f9377])).
  ( ! [X0] : (~iof(X0,sK7(sK0)) | firstordertype(X0)) ) | ~spl15_303),
  inference(avatar_component_clause,[],[f4091])).
fof(f9377,plain,(
fof(f4537,plain,(
  spl15_344 | ~spl15_20 | ~spl15_90),
  ~spl15_7 | ~spl15_20 | ~spl15_368 | ~spl15_369),
  inference(avatar_split_clause,[],[f809,f800,f203,f4535])).
fof(f800,plain,(
  spl15_90 <=> type_(sK7(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_90])])).
% Refutation found. Thanks to Tanya!
  inference(avatar_contradiction_clause,[],[f9376])).
fof(f809,plain,(
  iof(sK8(sK7(sK0)),sK7(sK0)) | (~spl15_20 | ~spl15_90)),
fof(f9376,plain,(
  inference(resolution,[],[f801,f204])).
% SZS status Theorem for SOT_hbsL
  $false | (~spl15_7 | ~spl15_20 | ~spl15_368 | ~spl15_369)),
fof(f801,plain,(
  type_(sK7(sK0)) | ~spl15_90),
  inference(subsumption_resolution,[],[f9372,f9364])).
% SZS output start Proof for SOT_hbsL
  inference(avatar_component_clause,[],[f800])).
fof(f4093,plain,(
fof(f9364,plain,(
  spl15_303 | ~spl15_35 | ~spl15_149),
  ( ! [X0] : (~iof(X0,sK8(sK7(sK0)))) ) | (~spl15_7 | ~spl15_368)),
  inference(avatar_split_clause,[],[f1250,f1247,f334,f4091])).
fof(f334,plain,(
  spl15_35 <=> ! [X1,X0] : (~iof(X1,X0) | firstordertype(X1) | ~secondordertype(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_35])])).
fof(f1247,plain,(
  inference(resolution,[],[f9359,f152])).
  spl15_149 <=> secondordertype(sK7(sK0))),
fof(f152,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_149])])).
fof(f1250,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) ) | ~spl15_7),
  ( ! [X0] : (firstordertype(X0) | ~iof(X0,sK7(sK0))) ) | (~spl15_35 | ~spl15_149)),
  inference(resolution,[],[f1248,f335])).
  inference(avatar_component_clause,[],[f151])).
fof(f335,plain,(
  ( ! [X0,X1] : (~secondordertype(X0) | firstordertype(X1) | ~iof(X1,X0)) ) | ~spl15_35),
fof(f151,plain,(
  inference(avatar_component_clause,[],[f334])).
fof(f1248,plain,(
  secondordertype(sK7(sK0)) | ~spl15_149),
  inference(avatar_component_clause,[],[f1247])).
  spl15_7 <=> ! [X1,X0] : (~iof(X1,X0) | ~individual(X0))),
fof(f4014,plain,(
  spl15_298 | ~spl15_38 | ~spl15_89),
  introduced(avatar_definition,[new_symbols(naming,[spl15_7])])).
  inference(avatar_split_clause,[],[f823,f793,f346,f4012])).
fof(f346,plain,(
fof(f9359,plain,(
  spl15_38 <=> ! [X1,X0] : (~iof(X1,X0) | individual(X1) | ~firstordertype(X0))),
fof(f9378,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_38])])).
  individual(sK8(sK7(sK0))) | ~spl15_368),
  $false),
fof(f793,plain,(
  spl15_89 <=> firstordertype(sK7(sK0))),
  inference(avatar_sat_refutation,[],[f129,f133,f145,f153,f169,f189,f201,f205,f226,f244,f336,f348,f396,f626,f728,f795,f802,f818,f917,f1239,f1249,f4014,f4093,f4537,f9353,f9360,f9368,f9377])).
  inference(avatar_component_clause,[],[f9358])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_89])])).
fof(f9377,plain,(
fof(f823,plain,(
  ~spl15_7 | ~spl15_20 | ~spl15_368 | ~spl15_369),
  ( ! [X0] : (individual(X0) | ~iof(X0,sK7(sK0))) ) | (~spl15_38 | ~spl15_89)),
  inference(avatar_contradiction_clause,[],[f9376])).
  inference(resolution,[],[f794,f347])).
fof(f9376,plain,(
fof(f347,plain,(
  $false | (~spl15_7 | ~spl15_20 | ~spl15_368 | ~spl15_369)),
fof(f9358,plain,(
  ( ! [X0,X1] : (~firstordertype(X0) | individual(X1) | ~iof(X1,X0)) ) | ~spl15_38),
  inference(subsumption_resolution,[],[f9372,f9364])).
  inference(avatar_component_clause,[],[f346])).
fof(f9364,plain,(
  spl15_368 <=> individual(sK8(sK7(sK0)))),
fof(f794,plain,(
  ( ! [X0] : (~iof(X0,sK8(sK7(sK0)))) ) | (~spl15_7 | ~spl15_368)),
  firstordertype(sK7(sK0)) | ~spl15_89),
  inference(resolution,[],[f9359,f152])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_368])])).
  inference(avatar_component_clause,[],[f793])).
fof(f152,plain,(
fof(f1249,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) ) | ~spl15_7),
fof(f9372,plain,(
  spl15_149 | ~spl15_16 | ~spl15_147),
  inference(avatar_component_clause,[],[f151])).
  inference(avatar_split_clause,[],[f1240,f1237,f187,f1247])).
fof(f151,plain,(
  iof(sK8(sK8(sK7(sK0))),sK8(sK7(sK0))) | (~spl15_20 | ~spl15_369)),
  spl15_7 <=> ! [X1,X0] : (~iof(X1,X0) | ~individual(X0))),
fof(f187,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_7])])).
  spl15_16 <=> ! [X1] : (secondordertype(X1) | ~iof(X1,c_sot))),
  inference(resolution,[],[f9367,f204])).
fof(f9359,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_16])])).
  individual(sK8(sK7(sK0))) | ~spl15_368),
fof(f1237,plain,(
fof(f204,plain,(
  inference(avatar_component_clause,[],[f9358])).
  spl15_147 <=> iof(sK7(sK0),c_sot)),
fof(f9358,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_147])])).
  ( ! [X0] : (~type_(X0) | iof(sK8(X0),X0)) ) | ~spl15_20),
  spl15_368 <=> individual(sK8(sK7(sK0)))),
fof(f1240,plain,(
  inference(avatar_component_clause,[],[f203])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_368])])).
  secondordertype(sK7(sK0)) | (~spl15_16 | ~spl15_147)),
fof(f9372,plain,(
  inference(resolution,[],[f1238,f188])).
fof(f203,plain,(
  iof(sK8(sK8(sK7(sK0))),sK8(sK7(sK0))) | (~spl15_20 | ~spl15_369)),
fof(f188,plain,(
  ( ! [X1] : (~iof(X1,c_sot) | secondordertype(X1)) ) | ~spl15_16),
  spl15_20 <=> ! [X0] : (iof(sK8(X0),X0) | ~type_(X0))),
  inference(avatar_component_clause,[],[f187])).
  inference(resolution,[],[f9367,f204])).
fof(f1238,plain,(
fof(f204,plain,(
  iof(sK7(sK0),c_sot) | ~spl15_147),
  ( ! [X0] : (~type_(X0) | iof(sK8(X0),X0)) ) | ~spl15_20),
  inference(avatar_component_clause,[],[f1237])).
  inference(avatar_component_clause,[],[f203])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_20])])).
fof(f1239,plain,(
fof(f203,plain,(
  spl15_88 | spl15_147 | ~spl15_19 | ~spl15_107),
  spl15_20 <=> ! [X0] : (iof(sK8(X0),X0) | ~type_(X0))),
fof(f9367,plain,(
  inference(avatar_split_clause,[],[f918,f915,f199,f1237,f790])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_20])])).
fof(f790,plain,(
  type_(sK8(sK7(sK0))) | ~spl15_369),
fof(f9367,plain,(
  spl15_88 <=> individual(sK0)),
  type_(sK8(sK7(sK0))) | ~spl15_369),
  introduced(avatar_definition,[new_symbols(naming,[spl15_88])])).
  inference(avatar_component_clause,[],[f9366])).
  inference(avatar_component_clause,[],[f9366])).
fof(f199,plain,(
  spl15_19 <=> ! [X0] : (iof(sK7(X0),X0) | individual(X0))),
fof(f9366,plain,(
fof(f9366,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_19])])).
  spl15_369 <=> type_(sK8(sK7(sK0)))),
fof(f915,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_369])])).
  spl15_107 <=> ! [X0] : (~iof(X0,sK0) | iof(X0,c_sot))),
fof(f9368,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_107])])).
  spl15_369 | ~spl15_5 | ~spl15_367),
  spl15_369 <=> type_(sK8(sK7(sK0)))),
fof(f918,plain,(
  inference(avatar_split_clause,[],[f9356,f9351,f143,f9366])).
  iof(sK7(sK0),c_sot) | individual(sK0) | (~spl15_19 | ~spl15_107)),
fof(f143,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_369])])).
  inference(resolution,[],[f916,f200])).
  spl15_5 <=> ! [X0] : (type_(X0) | ~firstordertype(X0))),
fof(f200,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_5])])).
fof(f9368,plain,(
  ( ! [X0] : (iof(sK7(X0),X0) | individual(X0)) ) | ~spl15_19),
fof(f9351,plain,(
  spl15_369 | ~spl15_5 | ~spl15_367),
  inference(avatar_component_clause,[],[f199])).
  spl15_367 <=> firstordertype(sK8(sK7(sK0)))),
fof(f916,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_367])])).
  inference(avatar_split_clause,[],[f9356,f9351,f143,f9366])).
  ( ! [X0] : (~iof(X0,sK0) | iof(X0,c_sot)) ) | ~spl15_107),
fof(f9356,plain,(
fof(f143,plain,(
  type_(sK8(sK7(sK0))) | (~spl15_5 | ~spl15_367)),
  inference(avatar_component_clause,[],[f915])).
fof(f917,plain,(
  inference(resolution,[],[f9352,f144])).
  spl15_107 | ~spl15_2 | ~spl15_68),
fof(f144,plain,(
  inference(avatar_split_clause,[],[f861,f624,f131,f915])).
  ( ! [X0] : (~firstordertype(X0) | type_(X0)) ) | ~spl15_5),
fof(f131,plain,(
  inference(avatar_component_clause,[],[f143])).
  spl15_5 <=> ! [X0] : (type_(X0) | ~firstordertype(X0))),
  spl15_2 <=> specializes(sK0,c_sot)),
fof(f9352,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_2])])).
  firstordertype(sK8(sK7(sK0))) | ~spl15_367),
  introduced(avatar_definition,[new_symbols(naming,[spl15_5])])).
fof(f624,plain,(
  inference(avatar_component_clause,[],[f9351])).
fof(f9351,plain,(
fof(f9360,plain,(
  spl15_68 <=> ! [X1,X0,X2] : (~specializes(X0,X1) | ~iof(X2,X0) | iof(X2,X1))),
  spl15_368 | ~spl15_298 | ~spl15_344),
  introduced(avatar_definition,[new_symbols(naming,[spl15_68])])).
  spl15_367 <=> firstordertype(sK8(sK7(sK0)))),
  inference(avatar_split_clause,[],[f9345,f4535,f4012,f9358])).
fof(f861,plain,(
fof(f4012,plain,(
  ( ! [X0] : (~iof(X0,sK0) | iof(X0,c_sot)) ) | (~spl15_2 | ~spl15_68)),
  spl15_298 <=> ! [X0] : (individual(X0) | ~iof(X0,sK7(sK0)))),
  inference(resolution,[],[f625,f132])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_298])])).
fof(f132,plain,(
fof(f4535,plain,(
  specializes(sK0,c_sot) | ~spl15_2),
  introduced(avatar_definition,[new_symbols(naming,[spl15_367])])).
  spl15_344 <=> iof(sK8(sK7(sK0)),sK7(sK0))),
  inference(avatar_component_clause,[],[f131])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_344])])).
fof(f9356,plain,(
fof(f625,plain,(
fof(f9345,plain,(
  ( ! [X2,X0,X1] : (~specializes(X0,X1) | ~iof(X2,X0) | iof(X2,X1)) ) | ~spl15_68),
  individual(sK8(sK7(sK0))) | (~spl15_298 | ~spl15_344)),
  type_(sK8(sK7(sK0))) | (~spl15_5 | ~spl15_367)),
  inference(avatar_component_clause,[],[f624])).
  inference(resolution,[],[f4536,f4013])).
fof(f818,plain,(
  inference(resolution,[],[f9352,f144])).
fof(f4013,plain,(
  ~spl15_7 | ~spl15_45 | ~spl15_88),
  ( ! [X0] : (~iof(X0,sK7(sK0)) | individual(X0)) ) | ~spl15_298),
  inference(avatar_contradiction_clause,[],[f817])).
fof(f144,plain,(
  inference(avatar_component_clause,[],[f4012])).
fof(f817,plain,(
fof(f4536,plain,(
  $false | (~spl15_7 | ~spl15_45 | ~spl15_88)),
  ( ! [X0] : (~firstordertype(X0) | type_(X0)) ) | ~spl15_5),
  iof(sK8(sK7(sK0)),sK7(sK0)) | ~spl15_344),
  inference(subsumption_resolution,[],[f395,f815])).
  inference(avatar_component_clause,[],[f4535])).
  inference(avatar_component_clause,[],[f143])).
fof(f815,plain,(
fof(f9353,plain,(
  ( ! [X0] : (~iof(X0,sK0)) ) | (~spl15_7 | ~spl15_88)),
  spl15_367 | ~spl15_303 | ~spl15_344),
fof(f9352,plain,(
  inference(resolution,[],[f791,f152])).
  inference(avatar_split_clause,[],[f9344,f4535,f4091,f9351])).
fof(f791,plain,(
fof(f4091,plain,(
  firstordertype(sK8(sK7(sK0))) | ~spl15_367),
  individual(sK0) | ~spl15_88),
  spl15_303 <=> ! [X0] : (firstordertype(X0) | ~iof(X0,sK7(sK0)))),
  inference(avatar_component_clause,[],[f790])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_303])])).
fof(f395,plain,(
fof(f9344,plain,(
  iof(sK8(sK0),sK0) | ~spl15_45),
  firstordertype(sK8(sK7(sK0))) | (~spl15_303 | ~spl15_344)),
  inference(avatar_component_clause,[],[f9351])).
  inference(avatar_component_clause,[],[f394])).
  inference(resolution,[],[f4536,f4092])).
fof(f394,plain,(
fof(f9360,plain,(
fof(f4092,plain,(
  spl15_45 <=> iof(sK8(sK0),sK0)),
  ( ! [X0] : (~iof(X0,sK7(sK0)) | firstordertype(X0)) ) | ~spl15_303),
  spl15_368 | ~spl15_298 | ~spl15_344),
  introduced(avatar_definition,[new_symbols(naming,[spl15_45])])).
  inference(avatar_component_clause,[],[f4091])).
fof(f802,plain,(
fof(f4537,plain,(
  inference(avatar_split_clause,[],[f9345,f4535,f4012,f9358])).
  spl15_90 | ~spl15_5 | ~spl15_89),
  spl15_344 | ~spl15_20 | ~spl15_90),
  inference(avatar_split_clause,[],[f798,f793,f143,f800])).
fof(f4012,plain,(
fof(f798,plain,(
  inference(avatar_split_clause,[],[f809,f800,f203,f4535])).
  type_(sK7(sK0)) | (~spl15_5 | ~spl15_89)),
  spl15_298 <=> ! [X0] : (individual(X0) | ~iof(X0,sK7(sK0)))),
fof(f800,plain,(
  inference(resolution,[],[f794,f144])).
  spl15_90 <=> type_(sK7(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_298])])).
fof(f795,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_90])])).
  spl15_88 | spl15_89 | ~spl15_19 | ~spl15_78),
fof(f809,plain,(
  inference(avatar_split_clause,[],[f729,f726,f199,f793,f790])).
  iof(sK8(sK7(sK0)),sK7(sK0)) | (~spl15_20 | ~spl15_90)),
fof(f726,plain,(
  inference(resolution,[],[f801,f204])).
fof(f4535,plain,(
  spl15_78 <=> ! [X0] : (firstordertype(X0) | ~iof(X0,sK0))),
fof(f801,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_78])])).
  type_(sK7(sK0)) | ~spl15_90),
  spl15_344 <=> iof(sK8(sK7(sK0)),sK7(sK0))),
fof(f729,plain,(
  inference(avatar_component_clause,[],[f800])).
  firstordertype(sK7(sK0)) | individual(sK0) | (~spl15_19 | ~spl15_78)),
fof(f4093,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_344])])).
  inference(resolution,[],[f727,f200])).
  spl15_303 | ~spl15_35 | ~spl15_149),
fof(f727,plain,(
  inference(avatar_split_clause,[],[f1250,f1247,f334,f4091])).
fof(f9345,plain,(
  ( ! [X0] : (~iof(X0,sK0) | firstordertype(X0)) ) | ~spl15_78),
fof(f334,plain,(
  inference(avatar_component_clause,[],[f726])).
  individual(sK8(sK7(sK0))) | (~spl15_298 | ~spl15_344)),
  spl15_35 <=> ! [X1,X0] : (~iof(X1,X0) | firstordertype(X1) | ~secondordertype(X0))),
fof(f728,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_35])])).
  spl15_78 | ~spl15_27 | ~spl15_35),
fof(f1247,plain,(
  inference(avatar_split_clause,[],[f417,f334,f242,f726])).
  spl15_149 <=> secondordertype(sK7(sK0))),
fof(f242,plain,(
  inference(resolution,[],[f4536,f4013])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_149])])).
  spl15_27 <=> secondordertype(sK0)),
fof(f1250,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_27])])).
fof(f4013,plain,(
  ( ! [X0] : (firstordertype(X0) | ~iof(X0,sK7(sK0))) ) | (~spl15_35 | ~spl15_149)),
fof(f417,plain,(
  inference(resolution,[],[f1248,f335])).
  ( ! [X0] : (firstordertype(X0) | ~iof(X0,sK0)) ) | (~spl15_27 | ~spl15_35)),
  ( ! [X0] : (~iof(X0,sK7(sK0)) | individual(X0)) ) | ~spl15_298),
fof(f335,plain,(
  inference(resolution,[],[f335,f243])).
  ( ! [X0,X1] : (~secondordertype(X0) | firstordertype(X1) | ~iof(X1,X0)) ) | ~spl15_35),
  inference(avatar_component_clause,[],[f4012])).
fof(f243,plain,(
  inference(avatar_component_clause,[],[f334])).
  secondordertype(sK0) | ~spl15_27),
fof(f1248,plain,(
fof(f4536,plain,(
  inference(avatar_component_clause,[],[f242])).
  secondordertype(sK7(sK0)) | ~spl15_149),
fof(f626,plain,(
  inference(avatar_component_clause,[],[f1247])).
  spl15_68),
fof(f4014,plain,(
  inference(avatar_split_clause,[],[f109,f624])).
  spl15_298 | ~spl15_38 | ~spl15_89),
  iof(sK8(sK7(sK0)),sK7(sK0)) | ~spl15_344),
fof(f109,plain,(
  inference(avatar_split_clause,[],[f823,f793,f346,f4012])).
  ( ! [X2,X0,X1] : (~specializes(X0,X1) | ~iof(X2,X0) | iof(X2,X1)) )),
fof(f346,plain,(
  inference(avatar_component_clause,[],[f4535])).
  inference(cnf_transformation,[],[f58])).
  spl15_38 <=> ! [X1,X0] : (~iof(X1,X0) | individual(X1) | ~firstordertype(X0))),
fof(f58,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_38])])).
fof(f9353,plain,(
  ! [X0,X1] : ((! [X2] : (iof(X2,X1) | ~iof(X2,X0)) & type_(X0)) | ~specializes(X0,X1))),
fof(f793,plain,(
  spl15_367 | ~spl15_303 | ~spl15_344),
  inference(ennf_transformation,[],[f42])).
  spl15_89 <=> firstordertype(sK7(sK0))),
fof(f42,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_89])])).
  inference(avatar_split_clause,[],[f9344,f4535,f4091,f9351])).
  ! [X0,X1] : (specializes(X0,X1) => (! [X2] : (iof(X2,X0) => iof(X2,X1)) & type_(X0)))),
fof(f823,plain,(
  inference(unused_predicate_definition_removal,[],[f41])).
  ( ! [X0] : (individual(X0) | ~iof(X0,sK7(sK0))) ) | (~spl15_38 | ~spl15_89)),
fof(f4091,plain,(
fof(f41,plain,(
  inference(resolution,[],[f794,f347])).
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X0) => iof(X2,X1)) & type_(X0)))),
fof(f347,plain,(
  spl15_303 <=> ! [X0] : (firstordertype(X0) | ~iof(X0,sK7(sK0)))),
  inference(rectify,[],[f12])).
  ( ! [X0,X1] : (~firstordertype(X0) | individual(X1) | ~iof(X1,X0)) ) | ~spl15_38),
  introduced(avatar_definition,[new_symbols(naming,[spl15_303])])).
fof(f12,axiom,(
  inference(avatar_component_clause,[],[f346])).
fof(f794,plain,(
  ! [X6,X7] : (specializes(X6,X7) <=> (! [X3] : (iof(X3,X6) => iof(X3,X7)) & type_(X6)))),
  firstordertype(sK7(sK0)) | ~spl15_89),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a7_specDef)).
  inference(avatar_component_clause,[],[f793])).
fof(f396,plain,(
fof(f9344,plain,(
fof(f1249,plain,(
  spl15_45 | ~spl15_20 | ~spl15_24),
  spl15_149 | ~spl15_16 | ~spl15_147),
  inference(avatar_split_clause,[],[f359,f224,f203,f394])).
  firstordertype(sK8(sK7(sK0))) | (~spl15_303 | ~spl15_344)),
  inference(avatar_split_clause,[],[f1240,f1237,f187,f1247])).
fof(f224,plain,(
fof(f187,plain,(
  spl15_24 <=> type_(sK0)),
  inference(resolution,[],[f4536,f4092])).
  spl15_16 <=> ! [X1] : (secondordertype(X1) | ~iof(X1,c_sot))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_24])])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_16])])).
fof(f4092,plain,(
fof(f359,plain,(
fof(f1237,plain,(
  iof(sK8(sK0),sK0) | (~spl15_20 | ~spl15_24)),
  spl15_147 <=> iof(sK7(sK0),c_sot)),
  ( ! [X0] : (~iof(X0,sK7(sK0)) | firstordertype(X0)) ) | ~spl15_303),
  inference(resolution,[],[f204,f225])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_147])])).
fof(f225,plain,(
  inference(avatar_component_clause,[],[f4091])).
fof(f1240,plain,(
  type_(sK0) | ~spl15_24),
  secondordertype(sK7(sK0)) | (~spl15_16 | ~spl15_147)),
  inference(avatar_component_clause,[],[f224])).
  inference(resolution,[],[f1238,f188])).
fof(f348,plain,(
fof(f188,plain,(
  spl15_38),
fof(f4537,plain,(
  ( ! [X1] : (~iof(X1,c_sot) | secondordertype(X1)) ) | ~spl15_16),
  inference(avatar_split_clause,[],[f77,f346])).
  inference(avatar_component_clause,[],[f187])).
fof(f77,plain,(
  spl15_344 | ~spl15_20 | ~spl15_90),
fof(f1238,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | individual(X1) | ~firstordertype(X0)) )),
  iof(sK7(sK0),c_sot) | ~spl15_147),
  inference(cnf_transformation,[],[f49])).
  inference(avatar_split_clause,[],[f809,f800,f203,f4535])).
  inference(avatar_component_clause,[],[f1237])).
fof(f49,plain,(
fof(f1239,plain,(
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (individual(X1) | ~iof(X1,X0)) & type_(X0)))),
fof(f800,plain,(
  spl15_88 | spl15_147 | ~spl15_19 | ~spl15_107),
  inference(ennf_transformation,[],[f10])).
  inference(avatar_split_clause,[],[f918,f915,f199,f1237,f790])).
  spl15_90 <=> type_(sK7(sK0))),
fof(f10,axiom,(
fof(f790,plain,(
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (iof(X1,X0) => individual(X1)) & type_(X0)))),
  spl15_88 <=> individual(sK0)),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a5_firstOrderTypeDef)).
  introduced(avatar_definition,[new_symbols(naming,[spl15_88])])).
fof(f336,plain,(
fof(f199,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_90])])).
  spl15_35),
  spl15_19 <=> ! [X0] : (iof(sK7(X0),X0) | individual(X0))),
  inference(avatar_split_clause,[],[f73,f334])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_19])])).
fof(f809,plain,(
fof(f73,plain,(
fof(f915,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | firstordertype(X1) | ~secondordertype(X0)) )),
  spl15_107 <=> ! [X0] : (~iof(X0,sK0) | iof(X0,c_sot))),
  iof(sK8(sK7(sK0)),sK7(sK0)) | (~spl15_20 | ~spl15_90)),
  inference(cnf_transformation,[],[f48])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_107])])).
fof(f48,plain,(
fof(f918,plain,(
  inference(resolution,[],[f801,f204])).
  iof(sK7(sK0),c_sot) | individual(sK0) | (~spl15_19 | ~spl15_107)),
  ! [X0] : (secondordertype(X0) <=> (! [X1] : (firstordertype(X1) | ~iof(X1,X0)) & type_(X0)))),
  inference(resolution,[],[f916,f200])).
  inference(ennf_transformation,[],[f11])).
fof(f801,plain,(
fof(f200,plain,(
fof(f11,axiom,(
  ( ! [X0] : (iof(sK7(X0),X0) | individual(X0)) ) | ~spl15_19),
  ! [X0] : (secondordertype(X0) <=> (! [X1] : (iof(X1,X0) => firstordertype(X1)) & type_(X0)))),
  inference(avatar_component_clause,[],[f199])).
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a6_secondOrderTypeDef)).
fof(f916,plain,(
fof(f244,plain,(
  type_(sK7(sK0)) | ~spl15_90),
  ( ! [X0] : (~iof(X0,sK0) | iof(X0,c_sot)) ) | ~spl15_107),
  spl15_27 | ~spl15_1 | ~spl15_16),
  inference(avatar_component_clause,[],[f915])).
  inference(avatar_split_clause,[],[f240,f187,f127,f242])).
  inference(avatar_component_clause,[],[f800])).
fof(f127,plain,(
fof(f917,plain,(
  spl15_1 <=> iof(sK0,c_sot)),
fof(f4093,plain,(
  spl15_107 | ~spl15_2 | ~spl15_68),
  introduced(avatar_definition,[new_symbols(naming,[spl15_1])])).
  inference(avatar_split_clause,[],[f861,f624,f131,f915])).
  spl15_303 | ~spl15_35 | ~spl15_149),
fof(f240,plain,(
fof(f131,plain,(
  secondordertype(sK0) | (~spl15_1 | ~spl15_16)),
  spl15_2 <=> specializes(sK0,c_sot)),
  inference(avatar_split_clause,[],[f1250,f1247,f334,f4091])).
  inference(resolution,[],[f188,f128])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_2])])).
fof(f128,plain,(
fof(f624,plain,(
fof(f334,plain,(
  iof(sK0,c_sot) | ~spl15_1),
  spl15_68 <=> ! [X1,X0,X2] : (~specializes(X0,X1) | ~iof(X2,X0) | iof(X2,X1))),
  inference(avatar_component_clause,[],[f127])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_68])])).
  spl15_35 <=> ! [X1,X0] : (~iof(X1,X0) | firstordertype(X1) | ~secondordertype(X0))),
fof(f226,plain,(
fof(f861,plain,(
  spl15_24 | ~spl15_2 | ~spl15_11),
  ( ! [X0] : (~iof(X0,sK0) | iof(X0,c_sot)) ) | (~spl15_2 | ~spl15_68)),
  introduced(avatar_definition,[new_symbols(naming,[spl15_35])])).
  inference(avatar_split_clause,[],[f222,f167,f131,f224])).
  inference(resolution,[],[f625,f132])).
fof(f167,plain,(
fof(f132,plain,(
  specializes(sK0,c_sot) | ~spl15_2),
  spl15_11 <=> ! [X1,X0] : (~specializes(X0,X1) | type_(X0))),
  inference(avatar_component_clause,[],[f131])).
fof(f1247,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_11])])).
fof(f625,plain,(
fof(f222,plain,(
  ( ! [X2,X0,X1] : (~specializes(X0,X1) | ~iof(X2,X0) | iof(X2,X1)) ) | ~spl15_68),
  spl15_149 <=> secondordertype(sK7(sK0))),
  type_(sK0) | (~spl15_2 | ~spl15_11)),
  inference(avatar_component_clause,[],[f624])).
  inference(resolution,[],[f168,f132])).
fof(f818,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_149])])).
fof(f168,plain,(
  ~spl15_7 | ~spl15_45 | ~spl15_88),
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) ) | ~spl15_11),
fof(f1250,plain,(
  inference(avatar_contradiction_clause,[],[f817])).
  inference(avatar_component_clause,[],[f167])).
fof(f817,plain,(
fof(f205,plain,(
  ( ! [X0] : (firstordertype(X0) | ~iof(X0,sK7(sK0))) ) | (~spl15_35 | ~spl15_149)),
  $false | (~spl15_7 | ~spl15_45 | ~spl15_88)),
  spl15_20),
  inference(subsumption_resolution,[],[f395,f815])).
  inference(resolution,[],[f1248,f335])).
  inference(avatar_split_clause,[],[f82,f203])).
fof(f815,plain,(
fof(f82,plain,(
  ( ! [X0] : (~iof(X0,sK0)) ) | (~spl15_7 | ~spl15_88)),
  ( ! [X0] : (iof(sK8(X0),X0) | ~type_(X0)) )),
  inference(resolution,[],[f791,f152])).
  inference(cnf_transformation,[],[f3])).
fof(f791,plain,(
fof(f335,plain,(
fof(f3,axiom,(
  individual(sK0) | ~spl15_88),
  inference(avatar_component_clause,[],[f790])).
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  ( ! [X0,X1] : (~secondordertype(X0) | firstordertype(X1) | ~iof(X1,X0)) ) | ~spl15_35),
fof(f395,plain,(
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a3)).
  iof(sK8(sK0),sK0) | ~spl15_45),
fof(f201,plain,(
  inference(avatar_component_clause,[],[f334])).
  inference(avatar_component_clause,[],[f394])).
  spl15_19),
fof(f394,plain,(
fof(f1248,plain,(
  inference(avatar_split_clause,[],[f79,f199])).
  spl15_45 <=> iof(sK8(sK0),sK0)),
fof(f79,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_45])])).
  secondordertype(sK7(sK0)) | ~spl15_149),
  ( ! [X0] : (iof(sK7(X0),X0) | individual(X0)) )),
fof(f802,plain,(
  inference(cnf_transformation,[],[f50])).
  spl15_90 | ~spl15_5 | ~spl15_89),
fof(f50,plain,(
  inference(avatar_split_clause,[],[f798,f793,f143,f800])).
  ! [X0] : (individual(X0) <=> ! [X1] : ~iof(X1,X0))),
fof(f798,plain,(
  inference(avatar_component_clause,[],[f1247])).
  inference(ennf_transformation,[],[f2])).
  type_(sK7(sK0)) | (~spl15_5 | ~spl15_89)),
  inference(resolution,[],[f794,f144])).
fof(f2,axiom,(
fof(f4014,plain,(
fof(f795,plain,(
  ! [X0] : (individual(X0) <=> ~? [X1] : iof(X1,X0))),
  spl15_88 | spl15_89 | ~spl15_19 | ~spl15_78),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a2)).
  spl15_298 | ~spl15_38 | ~spl15_89),
  inference(avatar_split_clause,[],[f729,f726,f199,f793,f790])).
fof(f189,plain,(
  inference(avatar_split_clause,[],[f823,f793,f346,f4012])).
fof(f726,plain,(
  spl15_16),
  spl15_78 <=> ! [X0] : (firstordertype(X0) | ~iof(X0,sK0))),
  inference(avatar_split_clause,[],[f115,f187])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_78])])).
fof(f115,plain,(
fof(f729,plain,(
  ( ! [X1] : (secondordertype(X1) | ~iof(X1,c_sot)) )),
  firstordertype(sK7(sK0)) | individual(sK0) | (~spl15_19 | ~spl15_78)),
fof(f346,plain,(
  inference(equality_resolution,[],[f92])).
  inference(resolution,[],[f727,f200])).
fof(f92,plain,(
fof(f727,plain,(
  spl15_38 <=> ! [X1,X0] : (~iof(X1,X0) | individual(X1) | ~firstordertype(X0))),
  ( ! [X0,X1] : (secondordertype(X1) | ~iof(X1,X0) | c_sot != X0) )),
  ( ! [X0] : (~iof(X0,sK0) | firstordertype(X0)) ) | ~spl15_78),
  inference(cnf_transformation,[],[f31])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_38])])).
  inference(avatar_component_clause,[],[f726])).
fof(f31,plain,(
fof(f728,plain,(
  ! [X0] : (c_sot = X0 <=> ! [X1] : (iof(X1,X0) <=> secondordertype(X1)))),
fof(f793,plain,(
  spl15_78 | ~spl15_27 | ~spl15_35),
  inference(rectify,[],[f23])).
  inference(avatar_split_clause,[],[f417,f334,f242,f726])).
fof(f23,axiom,(
  spl15_89 <=> firstordertype(sK7(sK0))),
fof(f242,plain,(
  ! [X10] : (c_sot = X10 <=> ! [X0] : (iof(X0,X10) <=> secondordertype(X0)))),
  spl15_27 <=> secondordertype(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl15_89])])).
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a15_sotConstantDef)).
  introduced(avatar_definition,[new_symbols(naming,[spl15_27])])).
fof(f169,plain,(
fof(f823,plain,(
fof(f417,plain,(
  spl15_11),
  ( ! [X0] : (firstordertype(X0) | ~iof(X0,sK0)) ) | (~spl15_27 | ~spl15_35)),
  inference(avatar_split_clause,[],[f110,f167])).
  ( ! [X0] : (individual(X0) | ~iof(X0,sK7(sK0))) ) | (~spl15_38 | ~spl15_89)),
  inference(resolution,[],[f335,f243])).
fof(f110,plain,(
fof(f243,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) )),
  inference(resolution,[],[f794,f347])).
  secondordertype(sK0) | ~spl15_27),
  inference(cnf_transformation,[],[f58])).
  inference(avatar_component_clause,[],[f242])).
fof(f153,plain,(
fof(f626,plain,(
  spl15_7),
  spl15_68),
  inference(avatar_split_clause,[],[f80,f151])).
fof(f347,plain,(
  inference(avatar_split_clause,[],[f109,f624])).
fof(f80,plain,(
fof(f109,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | ~individual(X0)) )),
  ( ! [X0,X1] : (~firstordertype(X0) | individual(X1) | ~iof(X1,X0)) ) | ~spl15_38),
  ( ! [X2,X0,X1] : (~specializes(X0,X1) | ~iof(X2,X0) | iof(X2,X1)) )),
  inference(cnf_transformation,[],[f50])).
  inference(cnf_transformation,[],[f58])).
  inference(avatar_component_clause,[],[f346])).
fof(f145,plain,(
fof(f58,plain,(
  spl15_5),
  ! [X0,X1] : ((! [X2] : (iof(X2,X1) | ~iof(X2,X0)) & type_(X0)) | ~specializes(X0,X1))),
fof(f794,plain,(
  inference(avatar_split_clause,[],[f78,f143])).
  inference(ennf_transformation,[],[f42])).
fof(f78,plain,(
  firstordertype(sK7(sK0)) | ~spl15_89),
fof(f42,plain,(
  ( ! [X0] : (type_(X0) | ~firstordertype(X0)) )),
  ! [X0,X1] : (specializes(X0,X1) => (! [X2] : (iof(X2,X0) => iof(X2,X1)) & type_(X0)))),
  inference(cnf_transformation,[],[f49])).
  inference(unused_predicate_definition_removal,[],[f41])).
fof(f133,plain,(
fof(f41,plain,(
  spl15_2),
  inference(avatar_component_clause,[],[f793])).
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X0) => iof(X2,X1)) & type_(X0)))),
  inference(avatar_split_clause,[],[f60,f131])).
  inference(rectify,[],[f12])).
fof(f60,plain,(
fof(f1249,plain,(
fof(f12,axiom,(
  specializes(sK0,c_sot)),
  ! [X6,X7] : (specializes(X6,X7) <=> (! [X3] : (iof(X3,X6) => iof(X3,X7)) & type_(X6)))),
  inference(cnf_transformation,[],[f27])).
  spl15_149 | ~spl15_16 | ~spl15_147),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a7_specDef)).
fof(f27,plain,(
fof(f396,plain,(
  inference(avatar_split_clause,[],[f1240,f1237,f187,f1247])).
  ? [X0] : (specializes(X0,c_sot) & iof(X0,c_sot))),
  spl15_45 | ~spl15_20 | ~spl15_24),
  inference(flattening,[],[f26])).
  inference(avatar_split_clause,[],[f359,f224,f203,f394])).
fof(f187,plain,(
fof(f26,plain,(
fof(f224,plain,(
  ~~? [X0] : (specializes(X0,c_sot) & iof(X0,c_sot))),
  spl15_24 <=> type_(sK0)),
  inference(rectify,[],[f25])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_24])])).
fof(f25,negated_conjecture,(
fof(f359,plain,(
  spl15_16 <=> ! [X1] : (secondordertype(X1) | ~iof(X1,c_sot))),
  ~~? [X10] : (specializes(X10,c_sot) & iof(X10,c_sot))),
  iof(sK8(sK0),sK0) | (~spl15_20 | ~spl15_24)),
  inference(negated_conjecture,[],[f24])).
  inference(resolution,[],[f204,f225])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_16])])).
fof(f24,conjecture,(
fof(f225,plain,(
fof(f1237,plain,(
  ~? [X10] : (specializes(X10,c_sot) & iof(X10,c_sot))),
  type_(sK0) | ~spl15_24),
  inference(avatar_component_clause,[],[f224])).
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',t9_ap3)).
  spl15_147 <=> iof(sK7(sK0),c_sot)),
fof(f348,plain,(
fof(f129,plain,(
  spl15_38),
  spl15_1),
  introduced(avatar_definition,[new_symbols(naming,[spl15_147])])).
  inference(avatar_split_clause,[],[f77,f346])).
  inference(avatar_split_clause,[],[f59,f127])).
fof(f77,plain,(
fof(f59,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | individual(X1) | ~firstordertype(X0)) )),
  iof(sK0,c_sot)),
fof(f1240,plain,(
  inference(cnf_transformation,[],[f49])).
  inference(cnf_transformation,[],[f27])).
fof(f49,plain,(
% SZS output end Proof for SOT_hbsL
  secondordertype(sK7(sK0)) | (~spl15_16 | ~spl15_147)),
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (individual(X1) | ~iof(X1,X0)) & type_(X0)))),
% ------------------------------
  inference(ennf_transformation,[],[f10])).
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
  inference(resolution,[],[f1238,f188])).
fof(f10,axiom,(
% Termination reason: Refutation
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (iof(X1,X0) => individual(X1)) & type_(X0)))),
fof(f188,plain,(

  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a5_firstOrderTypeDef)).
% Memory used [KB]: 10106
  ( ! [X1] : (~iof(X1,c_sot) | secondordertype(X1)) ) | ~spl15_16),
fof(f336,plain,(
% Time elapsed: 0.210 s
  spl15_35),
% ------------------------------
  inference(avatar_component_clause,[],[f187])).
  inference(avatar_split_clause,[],[f73,f334])).
% ------------------------------
fof(f73,plain,(
% Success in time 0.255 s
fof(f1238,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | firstordertype(X1) | ~secondordertype(X0)) )),
  inference(cnf_transformation,[],[f48])).
  iof(sK7(sK0),c_sot) | ~spl15_147),
fof(f48,plain,(
  ! [X0] : (secondordertype(X0) <=> (! [X1] : (firstordertype(X1) | ~iof(X1,X0)) & type_(X0)))),
  inference(ennf_transformation,[],[f11])).
fof(f11,axiom,(
  inference(avatar_component_clause,[],[f1237])).
  ! [X0] : (secondordertype(X0) <=> (! [X1] : (iof(X1,X0) => firstordertype(X1)) & type_(X0)))),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a6_secondOrderTypeDef)).
fof(f1239,plain,(
fof(f244,plain,(
  spl15_27 | ~spl15_1 | ~spl15_16),
  spl15_88 | spl15_147 | ~spl15_19 | ~spl15_107),
  inference(avatar_split_clause,[],[f240,f187,f127,f242])).
  inference(avatar_split_clause,[],[f918,f915,f199,f1237,f790])).
fof(f127,plain,(
  spl15_1 <=> iof(sK0,c_sot)),
fof(f790,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_1])])).
fof(f240,plain,(
  spl15_88 <=> individual(sK0)),
  secondordertype(sK0) | (~spl15_1 | ~spl15_16)),
  inference(resolution,[],[f188,f128])).
fof(f128,plain,(
  iof(sK0,c_sot) | ~spl15_1),
  introduced(avatar_definition,[new_symbols(naming,[spl15_88])])).
  inference(avatar_component_clause,[],[f127])).
fof(f226,plain,(
fof(f199,plain,(
  spl15_24 | ~spl15_2 | ~spl15_11),
  spl15_19 <=> ! [X0] : (iof(sK7(X0),X0) | individual(X0))),
  inference(avatar_split_clause,[],[f222,f167,f131,f224])).
fof(f167,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_19])])).
  spl15_11 <=> ! [X1,X0] : (~specializes(X0,X1) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_11])])).
fof(f915,plain,(
fof(f222,plain,(
  type_(sK0) | (~spl15_2 | ~spl15_11)),
  inference(resolution,[],[f168,f132])).
fof(f168,plain,(
  spl15_107 <=> ! [X0] : (~iof(X0,sK0) | iof(X0,c_sot))),
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) ) | ~spl15_11),
  inference(avatar_component_clause,[],[f167])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_107])])).
fof(f205,plain,(
  spl15_20),
fof(f918,plain,(
  inference(avatar_split_clause,[],[f82,f203])).
  iof(sK7(sK0),c_sot) | individual(sK0) | (~spl15_19 | ~spl15_107)),
fof(f82,plain,(
  ( ! [X0] : (iof(sK8(X0),X0) | ~type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  inference(resolution,[],[f916,f200])).
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a3)).
fof(f201,plain,(
fof(f200,plain,(
  spl15_19),
  ( ! [X0] : (iof(sK7(X0),X0) | individual(X0)) ) | ~spl15_19),
  inference(avatar_split_clause,[],[f79,f199])).
fof(f79,plain,(
  inference(avatar_component_clause,[],[f199])).
  ( ! [X0] : (iof(sK7(X0),X0) | individual(X0)) )),
  inference(cnf_transformation,[],[f50])).
fof(f916,plain,(
fof(f50,plain,(
  ( ! [X0] : (~iof(X0,sK0) | iof(X0,c_sot)) ) | ~spl15_107),
  ! [X0] : (individual(X0) <=> ! [X1] : ~iof(X1,X0))),
  inference(ennf_transformation,[],[f2])).
  inference(avatar_component_clause,[],[f915])).
fof(f2,axiom,(
  ! [X0] : (individual(X0) <=> ~? [X1] : iof(X1,X0))),
fof(f917,plain,(
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a2)).
  spl15_107 | ~spl15_2 | ~spl15_68),
fof(f189,plain,(
  spl15_16),
  inference(avatar_split_clause,[],[f115,f187])).
fof(f115,plain,(
  inference(avatar_split_clause,[],[f861,f624,f131,f915])).
  ( ! [X1] : (secondordertype(X1) | ~iof(X1,c_sot)) )),
  inference(equality_resolution,[],[f92])).
fof(f131,plain,(
fof(f92,plain,(
  ( ! [X0,X1] : (secondordertype(X1) | ~iof(X1,X0) | c_sot != X0) )),
  spl15_2 <=> specializes(sK0,c_sot)),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  introduced(avatar_definition,[new_symbols(naming,[spl15_2])])).
  ! [X0] : (c_sot = X0 <=> ! [X1] : (iof(X1,X0) <=> secondordertype(X1)))),
  inference(rectify,[],[f23])).
fof(f624,plain,(
fof(f23,axiom,(
  ! [X10] : (c_sot = X10 <=> ! [X0] : (iof(X0,X10) <=> secondordertype(X0)))),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a15_sotConstantDef)).
fof(f169,plain,(
  spl15_68 <=> ! [X1,X0,X2] : (~specializes(X0,X1) | ~iof(X2,X0) | iof(X2,X1))),
  spl15_11),
  inference(avatar_split_clause,[],[f110,f167])).
  introduced(avatar_definition,[new_symbols(naming,[spl15_68])])).
fof(f110,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) )),
fof(f861,plain,(
  inference(cnf_transformation,[],[f58])).
  ( ! [X0] : (~iof(X0,sK0) | iof(X0,c_sot)) ) | (~spl15_2 | ~spl15_68)),
fof(f153,plain,(
  spl15_7),
  inference(resolution,[],[f625,f132])).
  inference(avatar_split_clause,[],[f80,f151])).
fof(f80,plain,(
fof(f132,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | ~individual(X0)) )),
  inference(cnf_transformation,[],[f50])).
fof(f145,plain,(
  spl15_5),
  specializes(sK0,c_sot) | ~spl15_2),
  inference(avatar_split_clause,[],[f78,f143])).
fof(f78,plain,(
  inference(avatar_component_clause,[],[f131])).
  ( ! [X0] : (type_(X0) | ~firstordertype(X0)) )),
fof(f625,plain,(
  inference(cnf_transformation,[],[f49])).
fof(f133,plain,(
  ( ! [X2,X0,X1] : (~specializes(X0,X1) | ~iof(X2,X0) | iof(X2,X1)) ) | ~spl15_68),
  spl15_2),
  inference(avatar_split_clause,[],[f60,f131])).
fof(f60,plain,(
  specializes(sK0,c_sot)),
  inference(avatar_component_clause,[],[f624])).
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
fof(f818,plain,(
  ? [X0] : (specializes(X0,c_sot) & iof(X0,c_sot))),
  inference(flattening,[],[f26])).
  ~spl15_7 | ~spl15_45 | ~spl15_88),
fof(f26,plain,(
  ~~? [X0] : (specializes(X0,c_sot) & iof(X0,c_sot))),
  inference(avatar_contradiction_clause,[],[f817])).
  inference(rectify,[],[f25])).
fof(f817,plain,(
fof(f25,negated_conjecture,(
  ~~? [X10] : (specializes(X10,c_sot) & iof(X10,c_sot))),
  $false | (~spl15_7 | ~spl15_45 | ~spl15_88)),
  inference(negated_conjecture,[],[f24])).
fof(f24,conjecture,(
  inference(subsumption_resolution,[],[f395,f815])).
  ~? [X10] : (specializes(X10,c_sot) & iof(X10,c_sot))),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',t9_ap3)).
fof(f815,plain,(
fof(f129,plain,(
  spl15_1),
  inference(avatar_split_clause,[],[f59,f127])).
fof(f59,plain,(
  ( ! [X0] : (~iof(X0,sK0)) ) | (~spl15_7 | ~spl15_88)),
  iof(sK0,c_sot)),
  inference(cnf_transformation,[],[f27])).
  inference(resolution,[],[f791,f152])).
% SZS output end Proof for SOT_hbsL
% ------------------------------
fof(f791,plain,(
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
  individual(sK0) | ~spl15_88),
% Termination reason: Refutation

  inference(avatar_component_clause,[],[f790])).
% Memory used [KB]: 10106
% Time elapsed: 0.218 s
fof(f395,plain,(
% ------------------------------
% ------------------------------
% Success in time 0.258 s
  iof(sK8(sK0),sK0) | ~spl15_45),
  inference(avatar_component_clause,[],[f394])).
fof(f394,plain,(
  spl15_45 <=> iof(sK8(sK0),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl15_45])])).
fof(f802,plain,(
  spl15_90 | ~spl15_5 | ~spl15_89),
  inference(avatar_split_clause,[],[f798,f793,f143,f800])).
fof(f798,plain,(
  type_(sK7(sK0)) | (~spl15_5 | ~spl15_89)),
  inference(resolution,[],[f794,f144])).
fof(f795,plain,(
  spl15_88 | spl15_89 | ~spl15_19 | ~spl15_78),
  inference(avatar_split_clause,[],[f729,f726,f199,f793,f790])).
fof(f726,plain,(
  spl15_78 <=> ! [X0] : (firstordertype(X0) | ~iof(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_78])])).
fof(f729,plain,(
  firstordertype(sK7(sK0)) | individual(sK0) | (~spl15_19 | ~spl15_78)),
  inference(resolution,[],[f727,f200])).
fof(f727,plain,(
  ( ! [X0] : (~iof(X0,sK0) | firstordertype(X0)) ) | ~spl15_78),
  inference(avatar_component_clause,[],[f726])).
fof(f728,plain,(
  spl15_78 | ~spl15_27 | ~spl15_35),
  inference(avatar_split_clause,[],[f417,f334,f242,f726])).
fof(f242,plain,(
  spl15_27 <=> secondordertype(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl15_27])])).
fof(f417,plain,(
  ( ! [X0] : (firstordertype(X0) | ~iof(X0,sK0)) ) | (~spl15_27 | ~spl15_35)),
  inference(resolution,[],[f335,f243])).
fof(f243,plain,(
  secondordertype(sK0) | ~spl15_27),
  inference(avatar_component_clause,[],[f242])).
fof(f626,plain,(
  spl15_68),
  inference(avatar_split_clause,[],[f109,f624])).
fof(f109,plain,(
  ( ! [X2,X0,X1] : (~specializes(X0,X1) | ~iof(X2,X0) | iof(X2,X1)) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : ((! [X2] : (iof(X2,X1) | ~iof(X2,X0)) & type_(X0)) | ~specializes(X0,X1))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : (specializes(X0,X1) => (! [X2] : (iof(X2,X0) => iof(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X0) => iof(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7] : (specializes(X6,X7) <=> (! [X3] : (iof(X3,X6) => iof(X3,X7)) & type_(X6)))),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a7_specDef)).
fof(f396,plain,(
  spl15_45 | ~spl15_20 | ~spl15_24),
  inference(avatar_split_clause,[],[f359,f224,f203,f394])).
fof(f224,plain,(
  spl15_24 <=> type_(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl15_24])])).
fof(f359,plain,(
  iof(sK8(sK0),sK0) | (~spl15_20 | ~spl15_24)),
  inference(resolution,[],[f204,f225])).
fof(f225,plain,(
  type_(sK0) | ~spl15_24),
  inference(avatar_component_clause,[],[f224])).
fof(f348,plain,(
  spl15_38),
  inference(avatar_split_clause,[],[f77,f346])).
fof(f77,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | individual(X1) | ~firstordertype(X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (individual(X1) | ~iof(X1,X0)) & type_(X0)))),
  inference(ennf_transformation,[],[f10])).
fof(f10,axiom,(
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (iof(X1,X0) => individual(X1)) & type_(X0)))),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a5_firstOrderTypeDef)).
fof(f336,plain,(
  spl15_35),
  inference(avatar_split_clause,[],[f73,f334])).
fof(f73,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | firstordertype(X1) | ~secondordertype(X0)) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : (secondordertype(X0) <=> (! [X1] : (firstordertype(X1) | ~iof(X1,X0)) & type_(X0)))),
  inference(ennf_transformation,[],[f11])).
fof(f11,axiom,(
  ! [X0] : (secondordertype(X0) <=> (! [X1] : (iof(X1,X0) => firstordertype(X1)) & type_(X0)))),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a6_secondOrderTypeDef)).
fof(f244,plain,(
  spl15_27 | ~spl15_1 | ~spl15_16),
  inference(avatar_split_clause,[],[f240,f187,f127,f242])).
fof(f127,plain,(
  spl15_1 <=> iof(sK0,c_sot)),
  introduced(avatar_definition,[new_symbols(naming,[spl15_1])])).
fof(f240,plain,(
  secondordertype(sK0) | (~spl15_1 | ~spl15_16)),
  inference(resolution,[],[f188,f128])).
fof(f128,plain,(
  iof(sK0,c_sot) | ~spl15_1),
  inference(avatar_component_clause,[],[f127])).
fof(f226,plain,(
  spl15_24 | ~spl15_2 | ~spl15_11),
  inference(avatar_split_clause,[],[f222,f167,f131,f224])).
fof(f167,plain,(
  spl15_11 <=> ! [X1,X0] : (~specializes(X0,X1) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_11])])).
fof(f222,plain,(
  type_(sK0) | (~spl15_2 | ~spl15_11)),
  inference(resolution,[],[f168,f132])).
fof(f168,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) ) | ~spl15_11),
  inference(avatar_component_clause,[],[f167])).
fof(f205,plain,(
  spl15_20),
  inference(avatar_split_clause,[],[f82,f203])).
fof(f82,plain,(
  ( ! [X0] : (iof(sK8(X0),X0) | ~type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a3)).
fof(f201,plain,(
  spl15_19),
  inference(avatar_split_clause,[],[f79,f199])).
fof(f79,plain,(
  ( ! [X0] : (iof(sK7(X0),X0) | individual(X0)) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : (individual(X0) <=> ! [X1] : ~iof(X1,X0))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0] : (individual(X0) <=> ~? [X1] : iof(X1,X0))),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a2)).
fof(f189,plain,(
  spl15_16),
  inference(avatar_split_clause,[],[f115,f187])).
fof(f115,plain,(
  ( ! [X1] : (secondordertype(X1) | ~iof(X1,c_sot)) )),
  inference(equality_resolution,[],[f92])).
fof(f92,plain,(
  ( ! [X0,X1] : (secondordertype(X1) | ~iof(X1,X0) | c_sot != X0) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : (c_sot = X0 <=> ! [X1] : (iof(X1,X0) <=> secondordertype(X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10] : (c_sot = X10 <=> ! [X0] : (iof(X0,X10) <=> secondordertype(X0)))),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a15_sotConstantDef)).
fof(f169,plain,(
  spl15_11),
  inference(avatar_split_clause,[],[f110,f167])).
fof(f110,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) )),
  inference(cnf_transformation,[],[f58])).
fof(f153,plain,(
  spl15_7),
  inference(avatar_split_clause,[],[f80,f151])).
fof(f80,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | ~individual(X0)) )),
  inference(cnf_transformation,[],[f50])).
fof(f145,plain,(
  spl15_5),
  inference(avatar_split_clause,[],[f78,f143])).
fof(f78,plain,(
  ( ! [X0] : (type_(X0) | ~firstordertype(X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f133,plain,(
  spl15_2),
  inference(avatar_split_clause,[],[f60,f131])).
fof(f60,plain,(
  specializes(sK0,c_sot)),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ? [X0] : (specializes(X0,c_sot) & iof(X0,c_sot))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ~~? [X0] : (specializes(X0,c_sot) & iof(X0,c_sot))),
  inference(rectify,[],[f25])).
fof(f25,negated_conjecture,(
  ~~? [X10] : (specializes(X10,c_sot) & iof(X10,c_sot))),
  inference(negated_conjecture,[],[f24])).
fof(f24,conjecture,(
  ~? [X10] : (specializes(X10,c_sot) & iof(X10,c_sot))),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',t9_ap3)).
fof(f129,plain,(
  spl15_1),
  inference(avatar_split_clause,[],[f59,f127])).
fof(f59,plain,(
  iof(sK0,c_sot)),
  inference(cnf_transformation,[],[f27])).
% SZS output end Proof for SOT_hbsL
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 10106
% Time elapsed: 0.213 s
% ------------------------------
% ------------------------------
% Success in time 0.258 s
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_hbsL
% SZS output start Proof for SOT_hbsL
fof(f9378,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f129,f133,f145,f153,f169,f189,f201,f205,f226,f244,f336,f348,f396,f626,f728,f795,f802,f818,f917,f1239,f1249,f4014,f4093,f4537,f9353,f9360,f9368,f9377])).
fof(f9377,plain,(
  ~spl15_7 | ~spl15_20 | ~spl15_368 | ~spl15_369),
  inference(avatar_contradiction_clause,[],[f9376])).
fof(f9376,plain,(
  $false | (~spl15_7 | ~spl15_20 | ~spl15_368 | ~spl15_369)),
  inference(subsumption_resolution,[],[f9372,f9364])).
fof(f9364,plain,(
  ( ! [X0] : (~iof(X0,sK8(sK7(sK0)))) ) | (~spl15_7 | ~spl15_368)),
  inference(resolution,[],[f9359,f152])).
fof(f152,plain,(
  ( ! [X0,X1] : (~individual(X0) | ~iof(X1,X0)) ) | ~spl15_7),
  inference(avatar_component_clause,[],[f151])).
fof(f151,plain,(
  spl15_7 <=> ! [X1,X0] : (~iof(X1,X0) | ~individual(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_7])])).
fof(f9359,plain,(
  individual(sK8(sK7(sK0))) | ~spl15_368),
  inference(avatar_component_clause,[],[f9358])).
fof(f9358,plain,(
  spl15_368 <=> individual(sK8(sK7(sK0)))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_368])])).
fof(f9372,plain,(
  iof(sK8(sK8(sK7(sK0))),sK8(sK7(sK0))) | (~spl15_20 | ~spl15_369)),
  inference(resolution,[],[f9367,f204])).
fof(f204,plain,(
  ( ! [X0] : (~type_(X0) | iof(sK8(X0),X0)) ) | ~spl15_20),
  inference(avatar_component_clause,[],[f203])).
fof(f203,plain,(
  spl15_20 <=> ! [X0] : (iof(sK8(X0),X0) | ~type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_20])])).
fof(f9367,plain,(
  type_(sK8(sK7(sK0))) | ~spl15_369),
  inference(avatar_component_clause,[],[f9366])).
fof(f9366,plain,(
  spl15_369 <=> type_(sK8(sK7(sK0)))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_369])])).
fof(f9368,plain,(
  spl15_369 | ~spl15_5 | ~spl15_367),
  inference(avatar_split_clause,[],[f9356,f9351,f143,f9366])).
fof(f143,plain,(
  spl15_5 <=> ! [X0] : (type_(X0) | ~firstordertype(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_5])])).
fof(f9351,plain,(
  spl15_367 <=> firstordertype(sK8(sK7(sK0)))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_367])])).
fof(f9356,plain,(
  type_(sK8(sK7(sK0))) | (~spl15_5 | ~spl15_367)),
  inference(resolution,[],[f9352,f144])).
fof(f144,plain,(
  ( ! [X0] : (~firstordertype(X0) | type_(X0)) ) | ~spl15_5),
  inference(avatar_component_clause,[],[f143])).
fof(f9352,plain,(
  firstordertype(sK8(sK7(sK0))) | ~spl15_367),
  inference(avatar_component_clause,[],[f9351])).
fof(f9360,plain,(
  spl15_368 | ~spl15_298 | ~spl15_344),
  inference(avatar_split_clause,[],[f9345,f4535,f4012,f9358])).
fof(f4012,plain,(
  spl15_298 <=> ! [X0] : (individual(X0) | ~iof(X0,sK7(sK0)))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_298])])).
fof(f4535,plain,(
  spl15_344 <=> iof(sK8(sK7(sK0)),sK7(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_344])])).
fof(f9345,plain,(
  individual(sK8(sK7(sK0))) | (~spl15_298 | ~spl15_344)),
  inference(resolution,[],[f4536,f4013])).
fof(f4013,plain,(
  ( ! [X0] : (~iof(X0,sK7(sK0)) | individual(X0)) ) | ~spl15_298),
  inference(avatar_component_clause,[],[f4012])).
fof(f4536,plain,(
  iof(sK8(sK7(sK0)),sK7(sK0)) | ~spl15_344),
  inference(avatar_component_clause,[],[f4535])).
fof(f9353,plain,(
  spl15_367 | ~spl15_303 | ~spl15_344),
  inference(avatar_split_clause,[],[f9344,f4535,f4091,f9351])).
fof(f4091,plain,(
  spl15_303 <=> ! [X0] : (firstordertype(X0) | ~iof(X0,sK7(sK0)))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_303])])).
fof(f9344,plain,(
  firstordertype(sK8(sK7(sK0))) | (~spl15_303 | ~spl15_344)),
  inference(resolution,[],[f4536,f4092])).
fof(f4092,plain,(
  ( ! [X0] : (~iof(X0,sK7(sK0)) | firstordertype(X0)) ) | ~spl15_303),
  inference(avatar_component_clause,[],[f4091])).
fof(f4537,plain,(
  spl15_344 | ~spl15_20 | ~spl15_90),
  inference(avatar_split_clause,[],[f809,f800,f203,f4535])).
fof(f800,plain,(
  spl15_90 <=> type_(sK7(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_90])])).
fof(f809,plain,(
  iof(sK8(sK7(sK0)),sK7(sK0)) | (~spl15_20 | ~spl15_90)),
  inference(resolution,[],[f801,f204])).
fof(f801,plain,(
  type_(sK7(sK0)) | ~spl15_90),
  inference(avatar_component_clause,[],[f800])).
fof(f4093,plain,(
  spl15_303 | ~spl15_35 | ~spl15_149),
  inference(avatar_split_clause,[],[f1250,f1247,f334,f4091])).
fof(f334,plain,(
  spl15_35 <=> ! [X1,X0] : (~iof(X1,X0) | firstordertype(X1) | ~secondordertype(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_35])])).
fof(f1247,plain,(
  spl15_149 <=> secondordertype(sK7(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_149])])).
fof(f1250,plain,(
  ( ! [X0] : (firstordertype(X0) | ~iof(X0,sK7(sK0))) ) | (~spl15_35 | ~spl15_149)),
  inference(resolution,[],[f1248,f335])).
fof(f335,plain,(
  ( ! [X0,X1] : (~secondordertype(X0) | firstordertype(X1) | ~iof(X1,X0)) ) | ~spl15_35),
  inference(avatar_component_clause,[],[f334])).
fof(f1248,plain,(
  secondordertype(sK7(sK0)) | ~spl15_149),
  inference(avatar_component_clause,[],[f1247])).
fof(f4014,plain,(
  spl15_298 | ~spl15_38 | ~spl15_89),
  inference(avatar_split_clause,[],[f823,f793,f346,f4012])).
fof(f346,plain,(
  spl15_38 <=> ! [X1,X0] : (~iof(X1,X0) | individual(X1) | ~firstordertype(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_38])])).
fof(f793,plain,(
  spl15_89 <=> firstordertype(sK7(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_89])])).
fof(f823,plain,(
  ( ! [X0] : (individual(X0) | ~iof(X0,sK7(sK0))) ) | (~spl15_38 | ~spl15_89)),
  inference(resolution,[],[f794,f347])).
fof(f347,plain,(
  ( ! [X0,X1] : (~firstordertype(X0) | individual(X1) | ~iof(X1,X0)) ) | ~spl15_38),
  inference(avatar_component_clause,[],[f346])).
fof(f794,plain,(
  firstordertype(sK7(sK0)) | ~spl15_89),
  inference(avatar_component_clause,[],[f793])).
fof(f1249,plain,(
  spl15_149 | ~spl15_16 | ~spl15_147),
  inference(avatar_split_clause,[],[f1240,f1237,f187,f1247])).
fof(f187,plain,(
  spl15_16 <=> ! [X1] : (secondordertype(X1) | ~iof(X1,c_sot))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_16])])).
fof(f1237,plain,(
  spl15_147 <=> iof(sK7(sK0),c_sot)),
  introduced(avatar_definition,[new_symbols(naming,[spl15_147])])).
fof(f1240,plain,(
  secondordertype(sK7(sK0)) | (~spl15_16 | ~spl15_147)),
  inference(resolution,[],[f1238,f188])).
fof(f188,plain,(
  ( ! [X1] : (~iof(X1,c_sot) | secondordertype(X1)) ) | ~spl15_16),
  inference(avatar_component_clause,[],[f187])).
fof(f1238,plain,(
  iof(sK7(sK0),c_sot) | ~spl15_147),
  inference(avatar_component_clause,[],[f1237])).
fof(f1239,plain,(
  spl15_88 | spl15_147 | ~spl15_19 | ~spl15_107),
  inference(avatar_split_clause,[],[f918,f915,f199,f1237,f790])).
fof(f790,plain,(
  spl15_88 <=> individual(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl15_88])])).
fof(f199,plain,(
  spl15_19 <=> ! [X0] : (iof(sK7(X0),X0) | individual(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_19])])).
fof(f915,plain,(
  spl15_107 <=> ! [X0] : (~iof(X0,sK0) | iof(X0,c_sot))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_107])])).
fof(f918,plain,(
  iof(sK7(sK0),c_sot) | individual(sK0) | (~spl15_19 | ~spl15_107)),
  inference(resolution,[],[f916,f200])).
fof(f200,plain,(
  ( ! [X0] : (iof(sK7(X0),X0) | individual(X0)) ) | ~spl15_19),
  inference(avatar_component_clause,[],[f199])).
fof(f916,plain,(
  ( ! [X0] : (~iof(X0,sK0) | iof(X0,c_sot)) ) | ~spl15_107),
  inference(avatar_component_clause,[],[f915])).
fof(f917,plain,(
  spl15_107 | ~spl15_2 | ~spl15_68),
  inference(avatar_split_clause,[],[f861,f624,f131,f915])).
fof(f131,plain,(
  spl15_2 <=> specializes(sK0,c_sot)),
  introduced(avatar_definition,[new_symbols(naming,[spl15_2])])).
fof(f624,plain,(
  spl15_68 <=> ! [X1,X0,X2] : (~specializes(X0,X1) | ~iof(X2,X0) | iof(X2,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_68])])).
fof(f861,plain,(
  ( ! [X0] : (~iof(X0,sK0) | iof(X0,c_sot)) ) | (~spl15_2 | ~spl15_68)),
  inference(resolution,[],[f625,f132])).
fof(f132,plain,(
  specializes(sK0,c_sot) | ~spl15_2),
  inference(avatar_component_clause,[],[f131])).
fof(f625,plain,(
  ( ! [X2,X0,X1] : (~specializes(X0,X1) | ~iof(X2,X0) | iof(X2,X1)) ) | ~spl15_68),
  inference(avatar_component_clause,[],[f624])).
fof(f818,plain,(
  ~spl15_7 | ~spl15_45 | ~spl15_88),
  inference(avatar_contradiction_clause,[],[f817])).
fof(f817,plain,(
  $false | (~spl15_7 | ~spl15_45 | ~spl15_88)),
  inference(subsumption_resolution,[],[f395,f815])).
fof(f815,plain,(
  ( ! [X0] : (~iof(X0,sK0)) ) | (~spl15_7 | ~spl15_88)),
  inference(resolution,[],[f791,f152])).
fof(f791,plain,(
  individual(sK0) | ~spl15_88),
  inference(avatar_component_clause,[],[f790])).
fof(f395,plain,(
  iof(sK8(sK0),sK0) | ~spl15_45),
  inference(avatar_component_clause,[],[f394])).
fof(f394,plain,(
  spl15_45 <=> iof(sK8(sK0),sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl15_45])])).
fof(f802,plain,(
  spl15_90 | ~spl15_5 | ~spl15_89),
  inference(avatar_split_clause,[],[f798,f793,f143,f800])).
fof(f798,plain,(
  type_(sK7(sK0)) | (~spl15_5 | ~spl15_89)),
  inference(resolution,[],[f794,f144])).
fof(f795,plain,(
  spl15_88 | spl15_89 | ~spl15_19 | ~spl15_78),
  inference(avatar_split_clause,[],[f729,f726,f199,f793,f790])).
fof(f726,plain,(
  spl15_78 <=> ! [X0] : (firstordertype(X0) | ~iof(X0,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_78])])).
fof(f729,plain,(
  firstordertype(sK7(sK0)) | individual(sK0) | (~spl15_19 | ~spl15_78)),
  inference(resolution,[],[f727,f200])).
fof(f727,plain,(
  ( ! [X0] : (~iof(X0,sK0) | firstordertype(X0)) ) | ~spl15_78),
  inference(avatar_component_clause,[],[f726])).
fof(f728,plain,(
  spl15_78 | ~spl15_27 | ~spl15_35),
  inference(avatar_split_clause,[],[f417,f334,f242,f726])).
fof(f242,plain,(
  spl15_27 <=> secondordertype(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl15_27])])).
fof(f417,plain,(
  ( ! [X0] : (firstordertype(X0) | ~iof(X0,sK0)) ) | (~spl15_27 | ~spl15_35)),
  inference(resolution,[],[f335,f243])).
fof(f243,plain,(
  secondordertype(sK0) | ~spl15_27),
  inference(avatar_component_clause,[],[f242])).
fof(f626,plain,(
  spl15_68),
  inference(avatar_split_clause,[],[f109,f624])).
fof(f109,plain,(
  ( ! [X2,X0,X1] : (~specializes(X0,X1) | ~iof(X2,X0) | iof(X2,X1)) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : ((! [X2] : (iof(X2,X1) | ~iof(X2,X0)) & type_(X0)) | ~specializes(X0,X1))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : (specializes(X0,X1) => (! [X2] : (iof(X2,X0) => iof(X2,X1)) & type_(X0)))),
  inference(unused_predicate_definition_removal,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : (specializes(X0,X1) <=> (! [X2] : (iof(X2,X0) => iof(X2,X1)) & type_(X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7] : (specializes(X6,X7) <=> (! [X3] : (iof(X3,X6) => iof(X3,X7)) & type_(X6)))),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a7_specDef)).
fof(f396,plain,(
  spl15_45 | ~spl15_20 | ~spl15_24),
  inference(avatar_split_clause,[],[f359,f224,f203,f394])).
fof(f224,plain,(
  spl15_24 <=> type_(sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl15_24])])).
fof(f359,plain,(
  iof(sK8(sK0),sK0) | (~spl15_20 | ~spl15_24)),
  inference(resolution,[],[f204,f225])).
fof(f225,plain,(
  type_(sK0) | ~spl15_24),
  inference(avatar_component_clause,[],[f224])).
fof(f348,plain,(
  spl15_38),
  inference(avatar_split_clause,[],[f77,f346])).
fof(f77,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | individual(X1) | ~firstordertype(X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (individual(X1) | ~iof(X1,X0)) & type_(X0)))),
  inference(ennf_transformation,[],[f10])).
fof(f10,axiom,(
  ! [X0] : (firstordertype(X0) <=> (! [X1] : (iof(X1,X0) => individual(X1)) & type_(X0)))),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a5_firstOrderTypeDef)).
fof(f336,plain,(
  spl15_35),
  inference(avatar_split_clause,[],[f73,f334])).
fof(f73,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | firstordertype(X1) | ~secondordertype(X0)) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : (secondordertype(X0) <=> (! [X1] : (firstordertype(X1) | ~iof(X1,X0)) & type_(X0)))),
  inference(ennf_transformation,[],[f11])).
fof(f11,axiom,(
  ! [X0] : (secondordertype(X0) <=> (! [X1] : (iof(X1,X0) => firstordertype(X1)) & type_(X0)))),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a6_secondOrderTypeDef)).
fof(f244,plain,(
  spl15_27 | ~spl15_1 | ~spl15_16),
  inference(avatar_split_clause,[],[f240,f187,f127,f242])).
fof(f127,plain,(
  spl15_1 <=> iof(sK0,c_sot)),
  introduced(avatar_definition,[new_symbols(naming,[spl15_1])])).
fof(f240,plain,(
  secondordertype(sK0) | (~spl15_1 | ~spl15_16)),
  inference(resolution,[],[f188,f128])).
fof(f128,plain,(
  iof(sK0,c_sot) | ~spl15_1),
  inference(avatar_component_clause,[],[f127])).
fof(f226,plain,(
  spl15_24 | ~spl15_2 | ~spl15_11),
  inference(avatar_split_clause,[],[f222,f167,f131,f224])).
fof(f167,plain,(
  spl15_11 <=> ! [X1,X0] : (~specializes(X0,X1) | type_(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_11])])).
fof(f222,plain,(
  type_(sK0) | (~spl15_2 | ~spl15_11)),
  inference(resolution,[],[f168,f132])).
fof(f168,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) ) | ~spl15_11),
  inference(avatar_component_clause,[],[f167])).
fof(f205,plain,(
  spl15_20),
  inference(avatar_split_clause,[],[f82,f203])).
fof(f82,plain,(
  ( ! [X0] : (iof(sK8(X0),X0) | ~type_(X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (type_(X0) <=> ? [X1] : iof(X1,X0))),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a3)).
fof(f201,plain,(
  spl15_19),
  inference(avatar_split_clause,[],[f79,f199])).
fof(f79,plain,(
  ( ! [X0] : (iof(sK7(X0),X0) | individual(X0)) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : (individual(X0) <=> ! [X1] : ~iof(X1,X0))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0] : (individual(X0) <=> ~? [X1] : iof(X1,X0))),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a2)).
fof(f189,plain,(
  spl15_16),
  inference(avatar_split_clause,[],[f115,f187])).
fof(f115,plain,(
  ( ! [X1] : (secondordertype(X1) | ~iof(X1,c_sot)) )),
  inference(equality_resolution,[],[f92])).
fof(f92,plain,(
  ( ! [X0,X1] : (secondordertype(X1) | ~iof(X1,X0) | c_sot != X0) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : (c_sot = X0 <=> ! [X1] : (iof(X1,X0) <=> secondordertype(X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10] : (c_sot = X10 <=> ! [X0] : (iof(X0,X10) <=> secondordertype(X0)))),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',a15_sotConstantDef)).
fof(f169,plain,(
  spl15_11),
  inference(avatar_split_clause,[],[f110,f167])).
fof(f110,plain,(
  ( ! [X0,X1] : (~specializes(X0,X1) | type_(X0)) )),
  inference(cnf_transformation,[],[f58])).
fof(f153,plain,(
  spl15_7),
  inference(avatar_split_clause,[],[f80,f151])).
fof(f80,plain,(
  ( ! [X0,X1] : (~iof(X1,X0) | ~individual(X0)) )),
  inference(cnf_transformation,[],[f50])).
fof(f145,plain,(
  spl15_5),
  inference(avatar_split_clause,[],[f78,f143])).
fof(f78,plain,(
  ( ! [X0] : (type_(X0) | ~firstordertype(X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f133,plain,(
  spl15_2),
  inference(avatar_split_clause,[],[f60,f131])).
fof(f60,plain,(
  specializes(sK0,c_sot)),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ? [X0] : (specializes(X0,c_sot) & iof(X0,c_sot))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ~~? [X0] : (specializes(X0,c_sot) & iof(X0,c_sot))),
  inference(rectify,[],[f25])).
fof(f25,negated_conjecture,(
  ~~? [X10] : (specializes(X10,c_sot) & iof(X10,c_sot))),
  inference(negated_conjecture,[],[f24])).
fof(f24,conjecture,(
  ~? [X10] : (specializes(X10,c_sot) & iof(X10,c_sot))),
  file('/tmp/2fvw3M_yeU/SOT_hbsLD0',t9_ap3)).
fof(f129,plain,(
  spl15_1),
  inference(avatar_split_clause,[],[f59,f127])).
fof(f59,plain,(
  iof(sK0,c_sot)),
  inference(cnf_transformation,[],[f27])).
% SZS output end Proof for SOT_hbsL
% ------------------------------
% Version: Vampire 4.4.0 (commit 4fef31f on 2020-05-01 13:38:24 +0200)
% Termination reason: Refutation

% Memory used [KB]: 10106
% Time elapsed: 0.211 s
% ------------------------------
% ------------------------------
% Success in time 0.251 s
start to proof:SOT_hbsLD0

% END OF SYSTEM OUTPUT

% END OF SYSTEM OUTPUT

% END OF SYSTEM OUTPUT

% END OF SYSTEM OUTPUT

% END OF SYSTEM OUTPUT
% RESULT: SOT_hbsLD0 - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.25 
% OUTPUT: SOT_hbsLD0 - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.25 
% RESULT: SOT_hbsLD0 - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.26 
% OUTPUT: SOT_hbsLD0 - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.26 
% RESULT: SOT_hbsLD0 - E---FNT-2.5 says Theorem - CPU = 0.00 WC = 0.04 
% OUTPUT: SOT_hbsLD0 - E---FNT-2.5 says CNFRefutation - CPU = 0.00 WC = 0.04 
% RESULT: SOT_hbsLD0 - E---FNT-2.5 says Theorem - CPU = 0.00 WC = 0.04 
% OUTPUT: SOT_hbsLD0 - E---FNT-2.5 says CNFRefutation - CPU = 0.00 WC = 0.04 
% RESULT: SOT_hbsLD0 - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.25 
% OUTPUT: SOT_hbsLD0 - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.26 

% END OF SYSTEM OUTPUT
% RESULT: SOT_hbsLD0 - E---FNT-2.5 says Theorem - CPU = 0.00 WC = 0.08 
% OUTPUT: SOT_hbsLD0 - E---FNT-2.5 says CNFRefutation - CPU = 0.00 WC = 0.08 

% END OF SYSTEM OUTPUT
% RESULT: SOT_hbsLD0 - CSE_E---1.0 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_hbsLD0 - CSE_E---1.0 says None - CPU = 0 WC = 0 

% END OF SYSTEM OUTPUT
% RESULT: SOT_hbsLD0 - Vampire---SAT-4.5 says Theorem - CPU = 0.00 WC = 0.25 
% OUTPUT: SOT_hbsLD0 - Vampire---SAT-4.5 says Refutation - CPU = 0.00 WC = 0.25 
% 10 wall clock seconds gone, waiting for system output
% 20 wall clock seconds gone, waiting for system output
% 30 wall clock seconds gone, waiting for system output
% 40 wall clock seconds gone, waiting for system output

% END OF SYSTEM OUTPUT
% RESULT: SOT_hbsLD0 - Isabelle---2016 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_hbsLD0 - Isabelle---2016 says None - CPU = 0 WC = 0 

% END OF SYSTEM OUTPUT
% RESULT: SOT_hbsLD0 - Isabelle---2016 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_hbsLD0 - Isabelle---2016 says None - CPU = 0 WC = 0 
Vampire---SAT-4.5 has ended (PID 102627)
Vampire---SAT-4.5 has ended (PID 102628)
Vampire---SAT-4.5 has ended (PID 102629)
Vampire---SAT-4.5 has ended (PID 102630)
E---FNT-2.5 has ended (PID 102631)
E---FNT-2.5 has ended (PID 102632)
CSE_E---1.0 has ended (PID 102633)
E---FNT-2.5 has ended (PID 102634)
Isabelle---2016 has ended (PID 102635)
Isabelle---2016 has ended (PID 102636)
RESULT: SOT_hbsLD0 - SSCPA---0.0 says Theorem - CPU = 217.1 WC = 45.8  SolvingSystem = Vampire---SAT-4.5
</PRE>

</body>
</html>