<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>System Results</title>
<meta name="ROBOTS" content="INDEX,NOFOLLOW" />
<script type="text/javascript">//<![CDATA[

  var tstp_dump;
  function openSoTSTP (dump) {   
    var tstp_url = 'http://www.tptp.org/cgi-bin/SystemOnTSTP';
    var tstp_browser = window.open(tstp_url, '_blank');
    tstp_dump = dump;
  }
  function getTSTPDump () {
    return tstp_dump;
  }

//]]></script>
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>

% SZS start RequiredInformation<BR>
% Congratulations - you have become a registered power user of SystemOnTPTP, at IP address 191.6.38.118.<BR>
% Please consider donating to the TPTP project - see www.tptp.org for details.<BR>
% When you donate this message will disappear.<BR>
% If you do not donate a random delay might be added to your processing time.<BR>
% SZS end RequiredInformation<BR>
<PRE>
Getting system recommendations for generic problem /tmp/Rdg3WgVE6m/SOT_J_eUeW
This is a FOF_THM_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.95
Vampire---TF1-4.5    recommended  0.95
Vampire---4.5        recommended  0.95
Vampire---SAT-4.5    recommended  0.95
ET---2.0             recommended  0.80
E---2.5              recommended  0.78
E---FNT-2.5          recommended  0.78
CSE_E---1.0          recommended  0.76
E---FNT-2.5          recommended  0.74
Isabelle---2TH0      recommended  0.69
Isabelle---2TF0      recommended  0.69
Isabelle---2FOF      recommended  0.69
Isabelle---2016      recommended  0.69
Isabelle-HOT---2016  recommended  0.68
CVC4---1.7           recommended  0.68
CVC4---SAT-1.7       recommended  0.68
Leo-III---1.4        recommended  0.61
Vampire---THF-4.5    recommended  0.59
Vampire---TF1-4.5    recommended  0.59
Vampire---SAT-4.5    recommended  0.59
iProver---3.0        recommended  0.58
iProver---SAT-3.0    recommended  0.58
Fampire---1.3        recommended  0.57
SInE---0.4           recommended  0.53
iProver---SAT-3.0    recommended  0.52
ePrincess---1.0      recommended  0.50
Princess---170717    recommended  0.48
SPASS---3.9          recommended  0.48
Z3---4.4.1           recommended  0.46
Zipperpin---1.5      recommended  0.40
CVC4---SAT-1.7       recommended  0.40
Prover9---1109a      recommended  0.39
Equinox---6.0.1a     recommended  0.37
LEO-II---1.7.0       recommended  0.35
nanoCoP---1.1        recommended  0.34
Muscadet---4.5       recommended  0.25
Twee---2.2           recommended  0.23
Metis---2.4          recommended  0.23
Geo-III---2018C      recommended  0.18
iProverMo---2.5-0.1  recommended  0.14
Beagle---0.9.47      subsumed     0.41
leanCoP---2.2        subsumed     0.36
E-Darwin---1.5       subsumed     0.31
CSE---1.0            subsumed     0.28
Darwin---1.4.5       subsumed     0.28
Bliksem---1.12       subsumed     0.21
Zenon---0.7.1        subsumed     0.16
ZenonModulo---0.4.2  subsumed     0.15
SOS---2.0            subsumed     0.11
Otter---3.3          subsumed     0.04
This is a FOF_CSA_RFO_SEQ report.
Vampire---THF-4.5    recommended  0.98
Vampire---TF1-4.5    recommended  0.98
Vampire---SAT-4.5    recommended  0.98
Paradox---4.0        recommended  0.60
Nitpick---2016       subsumed     0.55
DarwinFM---1.4.5     subsumed     0.43
CVC4---SAT-1.7       subsumed     0.40
Vampire---THF-4.5    subsumed     0.34
Vampire---TF1-4.5    subsumed     0.34
Vampire---4.5        subsumed     0.34
Vampire---SAT-4.5    subsumed     0.34
E---2.5              subsumed     0.34
E---FNT-2.5          subsumed     0.34
E---FNT-2.5          subsumed     0.34
iProver---SAT-3.0    subsumed     0.32
Metis---2.4          subsumed     0.32
SPASS---3.9          subsumed     0.28
Beagle---0.9.47      subsumed     0.26
Geo-III---2018C      subsumed     0.25
Z3---4.4.1           subsumed     0.25
Princess---170717    subsumed     0.25
Equinox---6.0.1a     subsumed     0.19
iProver---3.0        subsumed     0.11
iProver---SAT-3.0    subsumed     0.11
Darwin---1.4.5       subsumed     0.06
E-Darwin---1.5       subsumed     0.04
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use ET,2.0
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use CSE_E,1.0
CSE_E---1.0 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use Isabelle,2TH0
Isabelle---2016 can be used
Check if we can use Isabelle,2TF0
Isabelle---2016 can be used
Check if we can use Isabelle,2FOF
Isabelle---2016 can be used
Check if we can use Isabelle,2016
Isabelle---2016 can be used
Check if we can use Isabelle-HOT,2016
Isabelle-HOT---2016 can be used
Check if we can use CVC4,1.7
CVC4---SAT-1.7 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Leo-III,1.4
Leo-III---1.4 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Fampire,1.3
Check if we can use SInE,0.4
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use ePrincess,1.0
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Zipperpin,1.5
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Prover9,1109a
Prover9---1109a can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use LEO-II,1.7.0
LEO-II---1.7.0 can be used
Check if we can use nanoCoP,1.1
nanoCoP---1.1 can be used
Check if we can use Muscadet,4.5
Muscadet---4.5 can be used
Check if we can use Twee,2.2
Twee---2.2 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use iProverMo,2.5-0.1
iProverMo---2.5-0.1 can be used
Check if we can use Beagle,0.9.47
Check if we can use leanCoP,2.2
leanCoP---2.2 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Check if we can use CSE,1.0
CSE---1.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use Bliksem,1.12
Bliksem---1.12 can be used
Check if we can use Zenon,0.7.1
Zenon---0.7.1 can be used
Check if we can use ZenonModulo,0.4.2
Check if we can use SOS,2.0
SOS---2.0 can be used
Check if we can use Otter,3.3
Otter---3.3 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use Paradox,4.0
Paradox---4.0 can be used
Check if we can use Nitpick,2016
Nitpick---2016 can be used
Check if we can use DarwinFM,1.4.5
DarwinFM---1.4.5 can be used
Check if we can use CVC4,SAT-1.7
CVC4---SAT-1.7 can be used
Check if we can use Vampire,THF-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,TF1-4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,4.5
Vampire---SAT-4.5 can be used
Check if we can use Vampire,SAT-4.5
Vampire---SAT-4.5 can be used
Check if we can use E,2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use E,FNT-2.5
E---FNT-2.5 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Metis,2.4
Metis---2.4 can be used
Check if we can use SPASS,3.9
SPASS---3.9 can be used
Check if we can use Beagle,0.9.47
Check if we can use Geo-III,2018C
Geo-III---2018C can be used
Check if we can use Z3,4.4.1
Z3---4.4.1 can be used
Check if we can use Princess,170717
Princess---170717 can be used
Check if we can use Equinox,6.0.1a
Equinox---6.0.1a can be used
Check if we can use iProver,3.0
iProver---SAT-3.0 can be used
Check if we can use iProver,SAT-3.0
iProver---SAT-3.0 can be used
Check if we can use Darwin,1.4.5
Darwin---1.4.5 can be used
Check if we can use E-Darwin,1.5
E-Darwin---1.5 can be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
Vampire---SAT-4.5 will be used
E---FNT-2.5 will be used
E---FNT-2.5 will be used
Vampire---SAT-4.5 is given 499 seconds in -cE mode
Vampire---SAT-4.5 is given 250 seconds in -cE mode
Vampire---SAT-4.5 is given 125 seconds in -cE mode
Vampire---SAT-4.5 is given 62 seconds in -cE mode
E---FNT-2.5 is given 31 seconds in -cE mode
E---FNT-2.5 is given 32 seconds in -cE mode
Vampire---SAT-4.5 has started (PID 97297)
Vampire---SAT-4.5 has started (PID 97298)
Vampire---SAT-4.5 has started (PID 97299)
Vampire---SAT-4.5 has started (PID 97300)
E---FNT-2.5 has started (PID 97301)
E---FNT-2.5 has started (PID 97303)
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  499  /tmp/Rdg3WgVE6m/SOT_J_eUeW
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  250  /tmp/Rdg3WgVE6m/SOT_J_eUeW
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  125  /tmp/Rdg3WgVE6m/SOT_J_eUeW
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 Vampire---SAT-4.5  62  /tmp/Rdg3WgVE6m/SOT_J_eUeW
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  31  /tmp/Rdg3WgVE6m/SOT_J_eUeW
Running /home/tptp/SystemExecution/SystemOnTPTP -q01 E---FNT-2.5  32  /tmp/Rdg3WgVE6m/SOT_J_eUeW
% START OF SYSTEM OUTPUT
# Version: 2.4dev006
# Trying AutoSched0 for 15 seconds
% START OF SYSTEM OUTPUT
% START OF SYSTEM OUTPUT
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_J_eU
# Version: 2.4dev006
# Trying AutoSched0 for 15 seconds
% START OF SYSTEM OUTPUT
% START OF SYSTEM OUTPUT
% START OF SYSTEM OUTPUT
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_J_eU
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_J_eU
# AutoSched0-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S5PRR_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.018 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3, conjecture, ![X7, X8, X9, X10]:(((specializes(X7,X8)&isPowertypeOf(X9,X7))&isPowertypeOf(X10,X8))=>specializes(X9,X10)), file('/tmp/Rdg3WgVE6m/SOT_J_eUeW', t3)).
fof(a10_isPowertypeOfDef, axiom, ![X7, X8]:(isPowertypeOf(X7,X8)<=>(type_(X7)&![X9]:(iof(X9,X7)<=>specializes(X9,X8)))), file('/tmp/Rdg3WgVE6m/SOT_J_eUeW', a10_isPowertypeOfDef)).
fof(a7_specDef, axiom, ![X7, X8]:(specializes(X7,X8)<=>(type_(X7)&![X4]:(iof(X4,X7)=>iof(X4,X8)))), file('/tmp/Rdg3WgVE6m/SOT_J_eUeW', a7_specDef)).
fof(c_0_3, negated_conjecture, ~(![X7, X8, X9, X10]:(((specializes(X7,X8)&isPowertypeOf(X9,X7))&isPowertypeOf(X10,X8))=>specializes(X9,X10))), inference(assume_negation,[status(cth)],[t3])).
fof(c_0_4, plain, ![X62, X63, X64, X65, X66, X67]:(((type_(X62)|~isPowertypeOf(X62,X63))&((~iof(X64,X62)|specializes(X64,X63)|~isPowertypeOf(X62,X63))&(~specializes(X65,X63)|iof(X65,X62)|~isPowertypeOf(X62,X63))))&((~iof(esk13_2(X66,X67),X66)|~specializes(esk13_2(X66,X67),X67)|~type_(X66)|isPowertypeOf(X66,X67))&(iof(esk13_2(X66,X67),X66)|specializes(esk13_2(X66,X67),X67)|~type_(X66)|isPowertypeOf(X66,X67)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a10_isPowertypeOfDef])])])])])])).
fof(c_0_5, negated_conjecture, (((specializes(esk22_0,esk23_0)&isPowertypeOf(esk24_0,esk22_0))&isPowertypeOf(esk25_0,esk23_0))&~specializes(esk24_0,esk25_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
cnf(c_0_6, plain, (specializes(X1,X3)|~iof(X1,X2)|~isPowertypeOf(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, negated_conjecture, (isPowertypeOf(esk24_0,esk22_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_8, plain, ![X49, X50, X51, X52, X53]:(((type_(X49)|~specializes(X49,X50))&(~iof(X51,X49)|iof(X51,X50)|~specializes(X49,X50)))&((iof(esk11_2(X52,X53),X52)|~type_(X52)|specializes(X52,X53))&(~iof(esk11_2(X52,X53),X53)|~type_(X52)|specializes(X52,X53)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a7_specDef])])])])])])).
cnf(c_0_9, plain, (type_(X1)|~isPowertypeOf(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10, negated_conjecture, (specializes(X1,esk22_0)|~iof(X1,esk24_0)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_11, plain, (iof(esk11_2(X1,X2),X1)|specializes(X1,X2)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, negated_conjecture, (type_(esk24_0)), inference(spm,[status(thm)],[c_0_9, c_0_7])).
cnf(c_0_13, plain, (iof(X1,X3)|~iof(X1,X2)|~specializes(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, negated_conjecture, (specializes(esk22_0,esk23_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_15, negated_conjecture, (specializes(esk11_2(esk24_0,X1),esk22_0)|specializes(esk24_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])])).
cnf(c_0_16, plain, (type_(X1)|~specializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_17, plain, (iof(X1,X3)|~specializes(X1,X2)|~isPowertypeOf(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_18, negated_conjecture, (isPowertypeOf(esk25_0,esk23_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_19, plain, (specializes(X1,X2)|~iof(esk11_2(X1,X2),X2)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_20, negated_conjecture, (iof(X1,esk23_0)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_21, negated_conjecture, (specializes(esk24_0,X1)|iof(X2,esk22_0)|~iof(X2,esk11_2(esk24_0,X1))), inference(spm,[status(thm)],[c_0_13, c_0_15])).
cnf(c_0_22, negated_conjecture, (specializes(esk24_0,X1)|type_(esk11_2(esk24_0,X1))), inference(spm,[status(thm)],[c_0_16, c_0_15])).
cnf(c_0_23, negated_conjecture, (iof(X1,esk25_0)|~specializes(X1,esk23_0)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24, negated_conjecture, (specializes(X1,esk23_0)|~type_(X1)|~iof(esk11_2(X1,esk23_0),esk22_0)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_25, negated_conjecture, (specializes(esk11_2(esk24_0,X1),X2)|specializes(esk24_0,X1)|iof(esk11_2(esk11_2(esk24_0,X1),X2),esk22_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_11]), c_0_22])).
cnf(c_0_26, negated_conjecture, (specializes(X1,esk25_0)|~specializes(esk11_2(X1,esk25_0),esk23_0)|~type_(X1)), inference(spm,[status(thm)],[c_0_19, c_0_23])).
cnf(c_0_27, negated_conjecture, (specializes(esk11_2(esk24_0,X1),esk23_0)|specializes(esk24_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_22])).
cnf(c_0_28, negated_conjecture, (~specializes(esk24_0,esk25_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_29, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_12])]), c_0_28]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 30
# Proof object clause steps            : 23
# Proof object formula steps           : 7
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 78
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 73
# Processed clauses                    : 4007
# ...of these trivial                  : 0
# ...subsumed                          : 1600
# ...remaining for further processing  : 2407
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 66
# Backward-rewritten                   : 0
# Generated clauses                    : 9730
# ...of the previous two non-trivial   : 8967
# Contextual simplify-reflections      : 233
# Paramodulations                      : 9724
# Factorizations                       : 4
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2339
#    Positive orientable unit clauses  : 42
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 14
#    Non-unit-clauses                  : 2283
# Current number of unprocessed clauses: 4986
# ...number of literals in the above   : 26831
# Current number of archived formulas  : 0
# Current number of archived clauses   : 68
# Clause-clause subsumption calls (NU) : 1220084
# Rec. Clause-clause subsumption calls : 267488
# Non-unit clause-clause subsumptions  : 1387
# Unit Clause-clause subsumption calls : 579
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 202572

# -------------------------------------------------
# User time                : 0.568 s
# System time              : 0.029 s
# Total time               : 0.597 s
# Maximum resident set size: 1608 pages
# AutoSched0-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S5PRR_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.017 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3, conjecture, ![X7, X8, X9, X10]:(((specializes(X7,X8)&isPowertypeOf(X9,X7))&isPowertypeOf(X10,X8))=>specializes(X9,X10)), file('/tmp/Rdg3WgVE6m/SOT_J_eUeW', t3)).
fof(a10_isPowertypeOfDef, axiom, ![X7, X8]:(isPowertypeOf(X7,X8)<=>(type_(X7)&![X9]:(iof(X9,X7)<=>specializes(X9,X8)))), file('/tmp/Rdg3WgVE6m/SOT_J_eUeW', a10_isPowertypeOfDef)).
fof(a7_specDef, axiom, ![X7, X8]:(specializes(X7,X8)<=>(type_(X7)&![X4]:(iof(X4,X7)=>iof(X4,X8)))), file('/tmp/Rdg3WgVE6m/SOT_J_eUeW', a7_specDef)).
fof(c_0_3, negated_conjecture, ~(![X7, X8, X9, X10]:(((specializes(X7,X8)&isPowertypeOf(X9,X7))&isPowertypeOf(X10,X8))=>specializes(X9,X10))), inference(assume_negation,[status(cth)],[t3])).
fof(c_0_4, plain, ![X62, X63, X64, X65, X66, X67]:(((type_(X62)|~isPowertypeOf(X62,X63))&((~iof(X64,X62)|specializes(X64,X63)|~isPowertypeOf(X62,X63))&(~specializes(X65,X63)|iof(X65,X62)|~isPowertypeOf(X62,X63))))&((~iof(esk13_2(X66,X67),X66)|~specializes(esk13_2(X66,X67),X67)|~type_(X66)|isPowertypeOf(X66,X67))&(iof(esk13_2(X66,X67),X66)|specializes(esk13_2(X66,X67),X67)|~type_(X66)|isPowertypeOf(X66,X67)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a10_isPowertypeOfDef])])])])])])).
fof(c_0_5, negated_conjecture, (((specializes(esk22_0,esk23_0)&isPowertypeOf(esk24_0,esk22_0))&isPowertypeOf(esk25_0,esk23_0))&~specializes(esk24_0,esk25_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
cnf(c_0_6, plain, (specializes(X1,X3)|~iof(X1,X2)|~isPowertypeOf(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, negated_conjecture, (isPowertypeOf(esk24_0,esk22_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_8, plain, ![X49, X50, X51, X52, X53]:(((type_(X49)|~specializes(X49,X50))&(~iof(X51,X49)|iof(X51,X50)|~specializes(X49,X50)))&((iof(esk11_2(X52,X53),X52)|~type_(X52)|specializes(X52,X53))&(~iof(esk11_2(X52,X53),X53)|~type_(X52)|specializes(X52,X53)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[a7_specDef])])])])])])).
cnf(c_0_9, plain, (type_(X1)|~isPowertypeOf(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10, negated_conjecture, (specializes(X1,esk22_0)|~iof(X1,esk24_0)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_11, plain, (iof(esk11_2(X1,X2),X1)|specializes(X1,X2)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, negated_conjecture, (type_(esk24_0)), inference(spm,[status(thm)],[c_0_9, c_0_7])).
cnf(c_0_13, plain, (iof(X1,X3)|~iof(X1,X2)|~specializes(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, negated_conjecture, (specializes(esk22_0,esk23_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_15, negated_conjecture, (specializes(esk11_2(esk24_0,X1),esk22_0)|specializes(esk24_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])])).
cnf(c_0_16, plain, (type_(X1)|~specializes(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_17, plain, (iof(X1,X3)|~specializes(X1,X2)|~isPowertypeOf(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_18, negated_conjecture, (isPowertypeOf(esk25_0,esk23_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_19, plain, (specializes(X1,X2)|~iof(esk11_2(X1,X2),X2)|~type_(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_20, negated_conjecture, (iof(X1,esk23_0)|~iof(X1,esk22_0)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_21, negated_conjecture, (specializes(esk24_0,X1)|iof(X2,esk22_0)|~iof(X2,esk11_2(esk24_0,X1))), inference(spm,[status(thm)],[c_0_13, c_0_15])).
cnf(c_0_22, negated_conjecture, (specializes(esk24_0,X1)|type_(esk11_2(esk24_0,X1))), inference(spm,[status(thm)],[c_0_16, c_0_15])).
cnf(c_0_23, negated_conjecture, (iof(X1,esk25_0)|~specializes(X1,esk23_0)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24, negated_conjecture, (specializes(X1,esk23_0)|~type_(X1)|~iof(esk11_2(X1,esk23_0),esk22_0)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_25, negated_conjecture, (specializes(esk11_2(esk24_0,X1),X2)|specializes(esk24_0,X1)|iof(esk11_2(esk11_2(esk24_0,X1),X2),esk22_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_11]), c_0_22])).
cnf(c_0_26, negated_conjecture, (specializes(X1,esk25_0)|~specializes(esk11_2(X1,esk25_0),esk23_0)|~type_(X1)), inference(spm,[status(thm)],[c_0_19, c_0_23])).
cnf(c_0_27, negated_conjecture, (specializes(esk11_2(esk24_0,X1),esk23_0)|specializes(esk24_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_22])).
cnf(c_0_28, negated_conjecture, (~specializes(esk24_0,esk25_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_29, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_12])]), c_0_28]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 30
# Proof object clause steps            : 23
# Proof object formula steps           : 7
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 78
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 73
# Processed clauses                    : 4007
# ...of these trivial                  : 0
# ...subsumed                          : 1600
# ...remaining for further processing  : 2407
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 66
# Backward-rewritten                   : 0
# Generated clauses                    : 9730
# ...of the previous two non-trivial   : 8967
# Contextual simplify-reflections      : 233
# Paramodulations                      : 9724
# Factorizations                       : 4
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
% ott+11_3_aac=none:afr=on:afp=4000:afq=1.4:amm=off:anc=all:bs=unit_only:bsr=on:bce=on:fde=unused:irw=on:nm=64:newcnf=on:nwc=1:nicw=on:sac=on:sp=reverse_arity:uhcvi=on_31 on SOT_J_eU
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2339
#    Positive orientable unit clauses  : 42
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 14
#    Non-unit-clauses                  : 2283
# Current number of unprocessed clauses: 4986
# ...number of literals in the above   : 26831
# Current number of archived formulas  : 0
# Current number of archived clauses   : 68
# Clause-clause subsumption calls (NU) : 1220084
# Rec. Clause-clause subsumption calls : 267488
# Non-unit clause-clause subsumptions  : 1387
# Unit Clause-clause subsumption calls : 579
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 202572

# -------------------------------------------------
# User time                : 0.610 s
# System time              : 0.014 s
# Total time               : 0.624 s
# Maximum resident set size: 1608 pages

% END OF SYSTEM OUTPUT
% RESULT: SOT_J_eUeW - E---FNT-2.5 says Theorem - CPU = 0.00 WC = 0.63 
% OUTPUT: SOT_J_eUeW - E---FNT-2.5 says CNFRefutation - CPU = 0.00 WC = 0.63 

% END OF SYSTEM OUTPUT
% RESULT: SOT_J_eUeW - E---FNT-2.5 says Theorem - CPU = 0.00 WC = 0.61 
% OUTPUT: SOT_J_eUeW - E---FNT-2.5 says CNFRefutation - CPU = 0.00 WC = 0.61 

% END OF SYSTEM OUTPUT
% RESULT: SOT_J_eUeW - Vampire---SAT-4.5 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_J_eUeW - Vampire---SAT-4.5 says None - CPU = 0 WC = 0 

% END OF SYSTEM OUTPUT
% RESULT: SOT_J_eUeW - Vampire---SAT-4.5 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_J_eUeW - Vampire---SAT-4.5 says None - CPU = 0 WC = 0 

% END OF SYSTEM OUTPUT
% RESULT: SOT_J_eUeW - Vampire---SAT-4.5 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_J_eUeW - Vampire---SAT-4.5 says None - CPU = 0 WC = 0 

% END OF SYSTEM OUTPUT
% RESULT: SOT_J_eUeW - Vampire---SAT-4.5 says Unknown - CPU = 0 WC = 0 
% OUTPUT: SOT_J_eUeW - Vampire---SAT-4.5 says None - CPU = 0 WC = 0 
Vampire---SAT-4.5 has ended (PID 97297)
Vampire---SAT-4.5 has ended (PID 97298)
Vampire---SAT-4.5 has ended (PID 97299)
Vampire---SAT-4.5 has ended (PID 97300)
E---FNT-2.5 has ended (PID 97301)
E---FNT-2.5 has ended (PID 97303)
RESULT: SOT_J_eUeW - SSCPA---0.0 says Theorem - CPU = 4.9 WC = 6.8  SolvingSystem = E---FNT-2.5
</PRE>

</body>
</html>